[1493785801] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1493787415] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;1;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1493787535] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;2;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1493787655] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;HARD;3;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1493788559] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1493788618] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1493789218] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1493789278] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1493789341] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1493789939] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1493791138] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1493792059] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1493792119] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1493792179] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 89% (15 min average)
[1493792313] SERVICE ALERT: FITNESSGRAM;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1493792338] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1493792338] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1493792373] SERVICE ALERT: FITNESSGRAM;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1493792433] SERVICE ALERT: FITNESSGRAM;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1493792779] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1493792839] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1493792899] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1493792938] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1493792998] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1493793033] SERVICE ALERT: FITNESSGRAM;CPU Load;OK;HARD;3;CPU Load 9% (15 min average)
[1493793058] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1493794258] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1493794699] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 72% (15 min average)
[1493794855] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;OK;HARD;3;OK: No entries found
[1493794859] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1493795899] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1493795959] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 82% (15 min average)
[1493797159] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1493797219] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1493797279] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1493798458] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1493798479] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1493798479] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1493799058] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1493799118] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1493799178] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1493799778] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1493800200] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service has exited from a period of scheduled downtime
[1493800203] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1493800203] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (50.7% change >= 20.0% threshold)
[1493800203] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1493800978] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1493801479] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1493801539] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 37% (15 min average)
[1493802003] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1493802003] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (49.2% change >= 20.0% threshold)
[1493802003] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1493806339] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1493806399] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1493806459] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1493807059] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 62% (15 min average)
[1493807659] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1493807719] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1493807779] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1493808379] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1493808979] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1493809039] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1493809099] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1493810299] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 78% (15 min average)
[1493810899] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1493810959] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 90% (15 min average)
[1493811778] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1493813358] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1493813417] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1493813475] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1493814075] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 41% (15 min average)
[1493815693] SERVICE ALERT: ADDC5;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1493815733] SERVICE ALERT: ADDC5;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1493816120] SERVICE ALERT: ADDC5;Service: DNS Service;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1493816148] SERVICE ALERT: ADDC5;Service: DFS Namespace;CRITICAL;SOFT;1;Dfs: Stopped
[1493816179] SERVICE ALERT: ADDC5;Service: DNS Service;OK;SOFT;2;DNS: Started
[1493816208] SERVICE ALERT: ADDC5;Service: DFS Namespace;OK;SOFT;2;Dfs: Started
[1493818203] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (36.2% change >= 20.0% threshold)
[1493818203] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1493820566] SERVICE ALERT: ADDC5;GPO Replication;WARNING;SOFT;1;Warning! 2 policies are not replicating properly!
[1493820867] SERVICE ALERT: ADDC5;GPO Replication;WARNING;SOFT;2;Warning! 2 policies are not replicating properly!
[1493821165] SERVICE ALERT: ADDC5;GPO Replication;WARNING;HARD;3;Warning! 2 policies are not replicating properly!
[1493821863] SERVICE ALERT: ADWU1;Service: User Profile;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1493821923] SERVICE ALERT: ADWU1;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1493822453] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;1;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1493822573] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;2;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1493822693] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;HARD;3;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1493823108] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1493823168] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1493823228] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1493824595] SERVICE ALERT: hallesecondsouth;Device Status;OK;HARD;3;Lower green handle.
[1493825333] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1493825394] SERVICE ALERT: marshallpk;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! 1 of 2 Trays Still Full
[1493825475] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1493827333] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 99.90 Gb - used: 90.00 Gb (90%) - free 9.90 Gb (10%)
[1493827393] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 99.90 Gb - used: 90.00 Gb (90%) - free 9.90 Gb (10%)
[1493827453] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 99.90 Gb - used: 90.00 Gb (90%) - free 9.90 Gb (10%)
[1493827461] SERVICE ALERT: CAVA2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1493827521] SERVICE ALERT: CAVA2;CPU Load;OK;SOFT;2;CPU Load 63% (15 min average)
[1493827875] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1493827935] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1493827993] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1493828593] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 73% (15 min average)
[1493829003] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1493829293] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;OK;HARD;3;OK: No entries found
[1493829396] SERVICE ALERT: owen101labcolor;Toner Collection Unit HP CE265A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1493829426] HOST ALERT: owen101labcolor;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1493829431] SERVICE ALERT: owen101labcolor;Black Cartridge HP CE260X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1493829456] SERVICE ALERT: owen101labcolor;Toner Collection Unit HP CE265A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1493829461] HOST ALERT: owen101labcolor;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1493829491] SERVICE ALERT: owen101labcolor;Black Cartridge HP CE260X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1493829516] SERVICE ALERT: owen101labcolor;Toner Collection Unit HP CE265A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1493829521] HOST ALERT: owen101labcolor;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1493829551] SERVICE ALERT: owen101labcolor;Black Cartridge HP CE260X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1493829581] HOST ALERT: owen101labcolor;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1493829612] SERVICE ALERT: owen101labcolor;Yellow Cartridge HP CE262A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1493829620] SERVICE ALERT: owen101labcolor;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1493829642] HOST ALERT: owen101labcolor;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1493829671] SERVICE ALERT: owen101labcolor;Yellow Cartridge HP CE262A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1493829680] SERVICE ALERT: owen101labcolor;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1493829701] HOST ALERT: owen101labcolor;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1493829730] SERVICE ALERT: owen101labcolor;Yellow Cartridge HP CE262A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1493829740] SERVICE ALERT: owen101labcolor;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1493829760] HOST ALERT: owen101labcolor;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1493829782] SERVICE ALERT: owen101labcolor;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1493829812] HOST ALERT: owen101labcolor;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1493829815] SERVICE ALERT: owen101labcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1493829824] SERVICE ALERT: owen101labcolor;Cyan Cartridge HP CE261A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1493829842] SERVICE ALERT: owen101labcolor;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1493829845] HOST ALERT: owen101labcolor;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1493829875] SERVICE ALERT: owen101labcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1493829879] SERVICE ALERT: owen101labcolor;Magenta Cartridge HP CE263A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1493829884] SERVICE ALERT: owen101labcolor;Cyan Cartridge HP CE261A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1493829885] SERVICE ALERT: owen101labcolor;Image Transfer Kit HP CE249A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1493829902] SERVICE ALERT: owen101labcolor;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1493829905] HOST ALERT: owen101labcolor;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1493829935] SERVICE ALERT: owen101labcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1493829939] SERVICE ALERT: owen101labcolor;Magenta Cartridge HP CE263A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1493829944] SERVICE ALERT: owen101labcolor;Cyan Cartridge HP CE261A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1493829945] SERVICE ALERT: owen101labcolor;Image Transfer Kit HP CE249A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1493829999] SERVICE ALERT: owen101labcolor;Magenta Cartridge HP CE263A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1493830005] SERVICE ALERT: owen101labcolor;Image Transfer Kit HP CE249A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1493833833] SERVICE ALERT: FITNESSGRAM;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1493833893] SERVICE ALERT: FITNESSGRAM;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1493833953] SERVICE ALERT: FITNESSGRAM;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1493835153] SERVICE ALERT: FITNESSGRAM;CPU Load;OK;HARD;3;CPU Load 3% (15 min average)