[1454608694] SERVICE ALERT: SANDBOX-DB;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 39.66 Gb - used: 32.52 Gb (82%) - free 7.14 Gb (18%)
[1454608720] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1454608765] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1454608814] SERVICE ALERT: SANDBOX-DB;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 39.66 Gb - used: 32.52 Gb (82%) - free 7.14 Gb (18%)
[1454608825] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1454608835] SERVICE ALERT: owen101lab3;Device Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1454608842] HOST ALERT: hallethirdsouth;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1454608865] HOST ALERT: owen101lab3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1454608881] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1454608885] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1454608889] SERVICE ALERT: hallethirdsouth;Device Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.107.213! Make sure host is up and SNMP is configured properly.
[1454608895] SERVICE ALERT: owen101lab3;Device Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1454608919] HOST ALERT: hallethirdsouth;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1454608925] HOST ALERT: owen101lab3;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1454608934] SERVICE ALERT: SANDBOX-DB;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 39.66 Gb - used: 32.52 Gb (82%) - free 7.14 Gb (18%)
[1454608949] SERVICE ALERT: hallethirdsouth;Paper Tray Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.107.213! Make sure host is up and SNMP is configured properly.
[1454608949] SERVICE ALERT: hallethirdsouth;Device Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.107.213! Make sure host is up and SNMP is configured properly.
[1454608955] SERVICE ALERT: owen101lab3;Device Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1454608979] HOST ALERT: hallethirdsouth;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1454608985] HOST ALERT: owen101lab3;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1454609009] SERVICE ALERT: hallethirdsouth;Paper Tray Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.107.213! Make sure host is up and SNMP is configured properly.
[1454609009] SERVICE ALERT: hallethirdsouth;Device Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.107.213! Make sure host is up and SNMP is configured properly.
[1454609039] HOST ALERT: hallethirdsouth;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1454609051] SERVICE ALERT: owen101lab3;Paper Tray Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1454609069] SERVICE ALERT: hallethirdsouth;Paper Tray Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.107.213! Make sure host is up and SNMP is configured properly.
[1454609075] HOST ALERT: owen101lab3;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1454609099] HOST ALERT: hallethirdsouth;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1454609111] SERVICE ALERT: owen101lab3;Paper Tray Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1454609141] HOST ALERT: owen101lab3;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1454609171] SERVICE ALERT: owen101lab3;Paper Tray Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1454609189] HOST ALERT: hallethirdsouth;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1454609192] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1454609201] HOST ALERT: owen101lab3;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1454609213] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;CRITICAL;HARD;3;Black Print Cartridge HP C8543X is at 1% - CRITICAL!
[1454609252] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1454609279] HOST ALERT: hallethirdsouth;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1454609282] HOST ALERT: owen101lab3;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1454609312] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;WARNING;HARD;3;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1454609342] HOST ALERT: owen101lab3;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1454609369] HOST ALERT: hallethirdsouth;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1454609432] HOST ALERT: owen101lab3;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1454609459] HOST ALERT: hallethirdsouth;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1454609485] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;CRITICAL;HARD;3;Black Print Cartridge HP C8543X is at 1% - CRITICAL!
[1454609522] HOST ALERT: owen101lab3;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1454609549] HOST ALERT: hallethirdsouth;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1454609621] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 36% - OK!
[1454610681] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1454610741] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 100% - OK!
[1454610779] SERVICE ALERT: DB01-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1454610899] SERVICE ALERT: DB01-2012;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1454611019] SERVICE ALERT: DB01-2012;CPU Load;OK;SOFT;3;CPU Load 0% (5 min average)
[1454611933] SERVICE ALERT: SANDBOX-DB;C:\ Drive Space;OK;HARD;3;c:\ - total: 39.66 Gb - used: 29.74 Gb (75%) - free 9.92 Gb (25%)
[1454612036] SERVICE ALERT: SANDBOX-DB;C:\ Drive Space;OK;HARD;3;c:\ - total: 39.66 Gb - used: 29.74 Gb (75%) - free 9.92 Gb (25%)
[1454612660] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1454612719] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1454612780] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1454612920] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1454612981] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1454613040] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1454613641] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1454613980] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1454614467] SERVICE ALERT: prayharroldpk;Device Status;WARNING;SOFT;1;Possible paper jam detected!
[1454614527] SERVICE ALERT: prayharroldpk;Device Status;OK;SOFT;2;Clearing paper path
[1454621578] SERVICE ALERT: marshallpk;Black Cartridge HP CC364X;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.137.101! Make sure host is up and SNMP is configured properly.
[1454621608] HOST ALERT: marshallpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1454621638] SERVICE ALERT: marshallpk;Black Cartridge HP CC364X;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.137.101! Make sure host is up and SNMP is configured properly.
[1454621642] HOST ALERT: marshallpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1454621671] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.137.101! Make sure host is up and SNMP is configured properly.
[1454621698] SERVICE ALERT: marshallpk;Black Cartridge HP CC364X;WARNING;HARD;3;WARNING: No SNMP response from 164.76.137.101! Make sure host is up and SNMP is configured properly.
[1454621729] SERVICE ALERT: marshallpk;Black Cartridge HP CC364X;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.137.101! Make sure host is up and SNMP is configured properly.
[1454621731] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.137.101! Make sure host is up and SNMP is configured properly.
[1454621748] SERVICE ALERT: ADCA2;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1454621790] SERVICE ALERT: marshallpk;Black Cartridge HP CC364X;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.137.101! Make sure host is up and SNMP is configured properly.
[1454621791] SERVICE ALERT: marshallpk;Device Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.137.101! Make sure host is up and SNMP is configured properly.
[1454621792] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.137.101! Make sure host is up and SNMP is configured properly.
[1454621849] SERVICE ALERT: marshallpk;Black Cartridge HP CC364X;WARNING;HARD;3;WARNING: No SNMP response from 164.76.137.101! Make sure host is up and SNMP is configured properly.
[1454621851] SERVICE ALERT: marshallpk;Device Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.137.101! Make sure host is up and SNMP is configured properly.
[1454621868] SERVICE ALERT: ADCA2;CPU Load;OK;SOFT;2;CPU Load 2% (5 min average)
[1454621911] SERVICE ALERT: marshallpk;Device Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.137.101! Make sure host is up and SNMP is configured properly.
[1454622020] SERVICE ALERT: marshallpk;Device Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.137.101! Make sure host is up and SNMP is configured properly.
[1454622024] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.137.101! Make sure host is up and SNMP is configured properly.
[1454622080] SERVICE ALERT: marshallpk;Device Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.137.101! Make sure host is up and SNMP is configured properly.
[1454622084] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.137.101! Make sure host is up and SNMP is configured properly.
[1454622140] SERVICE ALERT: marshallpk;Device Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.137.101! Make sure host is up and SNMP is configured properly.
[1454622144] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.137.101! Make sure host is up and SNMP is configured properly.
[1454622845] HOST ALERT: marshallpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1454622892] SERVICE ALERT: marshallpk;Black Cartridge HP CC364X;OK;HARD;3;Black Cartridge HP CC364X is at 65% - OK!
[1454622909] HOST ALERT: marshallpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.47 ms
[1454622984] SERVICE ALERT: marshallpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 3 of 4 Trays Still Full
[1454623042] SERVICE ALERT: marshallpk;Black Cartridge HP CC364X;OK;HARD;3;Black Cartridge HP CC364X is at 65% - OK!
[1454623105] SERVICE ALERT: marshallpk;Device Status;OK;HARD;3;TRAY 5 EMPTY PLAIN LETTER   Ready
[1454623223] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1454623284] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 37% - OK!
[1454623335] SERVICE ALERT: marshallpk;Device Status;OK;HARD;3;TRAY 5 EMPTY PLAIN LETTER   Ready
[1454623339] SERVICE ALERT: marshallpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 3 of 4 Trays Still Full
[1454623884] SERVICE ALERT: owen101lab1;Device Status;WARNING;SOFT;1;Possible paper jam detected!
[1454623929] SERVICE ALERT: owen101lab1;Device Status;WARNING;SOFT;1;Possible paper jam detected!
[1454623944] SERVICE ALERT: owen101lab1;Device Status;WARNING;SOFT;2;Possible paper jam detected!
[1454623989] SERVICE ALERT: owen101lab1;Device Status;OK;SOFT;2;Remove any  jammed  paper.
[1454624004] SERVICE ALERT: owen101lab1;Device Status;OK;SOFT;3;Remove any  jammed  paper.
[1454624589] SERVICE ALERT: owen101lab1;Device Status;WARNING;SOFT;1;Possible paper jam detected!
[1454624605] SERVICE ALERT: owen101lab1;Device Status;WARNING;SOFT;1;Possible paper jam detected!
[1454624649] SERVICE ALERT: owen101lab1;Device Status;WARNING;SOFT;2;Possible paper jam detected!
[1454624664] SERVICE ALERT: owen101lab1;Device Status;WARNING;SOFT;2;Possible paper jam detected!
[1454624710] SERVICE ALERT: owen101lab1;Device Status;WARNING;HARD;3;Possible paper jam detected!
[1454624724] SERVICE ALERT: owen101lab1;Device Status;WARNING;HARD;3;Possible paper jam detected!
[1454625264] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1454625324] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! 2 of 3 Trays Still Full
[1454625909] SERVICE ALERT: owen101lab1;Device Status;OK;HARD;3;TRAY 4 EMPTY PLAIN LETTER   Processing job from tray 3
[1454625924] SERVICE ALERT: owen101lab1;Device Status;OK;HARD;3;TRAY 4 EMPTY PLAIN LETTER   Processing job from tray 3
[1454628228] SERVICE ALERT: hallegroundcolor;Magenta Cartridge HP CE263A;WARNING;SOFT;1;Magenta Cartridge HP CE263A is at 5% - WARNING!
[1454628289] SERVICE ALERT: hallegroundcolor;Magenta Cartridge HP CE263A;WARNING;SOFT;2;Magenta Cartridge HP CE263A is at 5% - WARNING!
[1454628348] SERVICE ALERT: hallegroundcolor;Magenta Cartridge HP CE263A;WARNING;HARD;3;Magenta Cartridge HP CE263A is at 5% - WARNING!
[1454628561] SERVICE ALERT: hallegroundcolor;Magenta Cartridge HP CE263A;WARNING;SOFT;1;Magenta Cartridge HP CE263A is at 4% - WARNING!
[1454628622] SERVICE ALERT: hallegroundcolor;Magenta Cartridge HP CE263A;WARNING;SOFT;2;Magenta Cartridge HP CE263A is at 4% - WARNING!
[1454628681] SERVICE ALERT: hallegroundcolor;Magenta Cartridge HP CE263A;WARNING;HARD;3;Magenta Cartridge HP CE263A is at 4% - WARNING!
[1454638618] SERVICE ALERT: DB01-2012;CPU Load;WARNING;SOFT;1;CPU Load 86% (5 min average)
[1454638738] SERVICE ALERT: DB01-2012;CPU Load;CRITICAL;SOFT;2;CPU Load 90% (5 min average)
[1454638858] SERVICE ALERT: DB01-2012;CPU Load;OK;SOFT;3;CPU Load 72% (5 min average)