[1491796941] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1491797000] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1491797061] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1491797981] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1491798041] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1491798101] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1491798132] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1491798192] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1491798252] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1491798260] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 51% (15 min average)
[1491798449] SERVICE ALERT: halleground2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1491798509] SERVICE ALERT: halleground2;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1491798569] SERVICE ALERT: halleground2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1491798601] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1491800658] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1491800718] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 71% (15 min average)
[1491801184] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1491801243] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1491801318] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1491801378] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 83% (15 min average)
[1491801845] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1491801903] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1491801964] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1491801978] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1491802038] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1491802098] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 92% (15 min average)
[1491803898] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1491803958] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 92% (15 min average)
[1491804363] SERVICE ALERT: FITNESSGRAM;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1491804423] SERVICE ALERT: FITNESSGRAM;CPU Load;OK;SOFT;2;CPU Load 3% (15 min average)
[1491804963] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1491805074] SERVICE ALERT: ADCA2;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1491805134] SERVICE ALERT: ADCA2;CPU Load;OK;SOFT;2;CPU Load 3% (15 min average)
[1491805158] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1491805218] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1491805278] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1491806163] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1491806478] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 80% (15 min average)
[1491807078] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1491807138] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1491807198] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 94% (15 min average)
[1491807363] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1491807363] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1491807963] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1491808398] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1491808458] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1491808518] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1491810318] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1491810918] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1491810964] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1491811563] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1491812716] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 68% (15 min average)
[1491813000] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service has exited from a period of scheduled downtime
[1491813003] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (32.8% change >= 20.0% threshold)
[1491813003] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (35.3% change >= 20.0% threshold)
[1491813003] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1491813003] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1491813003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1491813003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1491813003] HOST DOWNTIME ALERT: INTLAPP;STARTED; Host has entered a period of scheduled downtime
[1491813003] HOST DOWNTIME ALERT: INTLDB;STARTED; Host has entered a period of scheduled downtime
[1491813003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1491813003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1491813003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1491813003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1491813003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1491813003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1491813003] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1491813003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1491813003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1491813003] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1491813003] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1491814802] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (30.4% change >= 20.0% threshold)
[1491814802] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (29.1% change >= 20.0% threshold)
[1491814803] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1491814803] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1491814803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1491814803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1491814803] HOST DOWNTIME ALERT: INTLAPP;STARTED; Host has entered a period of scheduled downtime
[1491814803] HOST DOWNTIME ALERT: INTLDB;STARTED; Host has entered a period of scheduled downtime
[1491814803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1491814803] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1491814803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1491814803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1491814803] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1491814803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1491814803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1491814803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1491814803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1491814803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1491814803] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1491815163] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1491815223] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1491815823] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1491815883] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1491816483] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1491816543] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1491819315] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1491819375] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1491819975] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1491820035] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1491820095] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1491821295] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1491822129] HOST ALERT: halleground2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1491822219] HOST ALERT: halleground2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1491822263] HOST ALERT: halleground2;UP;SOFT;3;PING OK - Packet loss = 66%, RTA = 0.50 ms
[1491822516] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Warming up
[1491822520] HOST ALERT: hallefirst2;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.56 ms
[1491822521] HOST ALERT: hallefirst3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1491822585] HOST ALERT: hallefirst3;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1491822617] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 8% - OK!
[1491822697] SERVICE ALERT: hallefirst1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1491822701] HOST ALERT: hallefirst1;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.49 ms
[1491822797] HOST ALERT: halleground1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1491822817] HOST ALERT: hallefirst3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1491822843] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 3% - OK!
[1491822861] HOST ALERT: halleground1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1491822887] SERVICE ALERT: halleground1;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1491822907] HOST ALERT: hallefirst3;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1491822913] HOST ALERT: hallefirstcolor;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1491822917] HOST ALERT: halleground1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1491822947] SERVICE ALERT: halleground1;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1491822952] HOST ALERT: hallefirstcolor;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1491822971] HOST ALERT: hallefirst3;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.46 ms
[1491822971] HOST FLAPPING ALERT: hallefirst3;STARTED; Host appears to have started flapping (23.3% change > 20.0% threshold)
[1491822977] HOST ALERT: halleground1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1491822993] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1491822998] HOST ALERT: hallefirstcolor;UP;SOFT;3;PING OK - Packet loss = 28%, RTA = 0.40 ms
[1491823007] SERVICE ALERT: halleground1;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1491823037] HOST ALERT: halleground1;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1491823048] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;ORDER CARTRIDGE LESS THAN 700 PAGES 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1491823056] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;OK;HARD;3;Magenta Cartridge HP CE263A is at 50% - OK!
[1491823088] SERVICE ALERT: hallefirst3;Device Status;OK;HARD;3;Ready
[1491823089] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;OK;HARD;3;Yellow Cartridge HP CE262A is at 31% - OK!
[1491823108] SERVICE ALERT: halleground1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1491823127] HOST ALERT: halleground1;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1491823144] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1491823164] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1491823168] SERVICE ALERT: halleground1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1491823181] SERVICE ALERT: hallefirstcolor;Device Status;OK;HARD;3;Ready
[1491823187] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 9% - OK!
[1491823194] HOST ALERT: halleground1;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1491823224] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1491823228] SERVICE ALERT: halleground1;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1491823254] HOST ALERT: halleground1;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1491823282] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;OK;HARD;3;Black Cartridge HP CE260X is at 31% - OK!
[1491823284] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1491823314] HOST ALERT: halleground1;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1491823330] SERVICE ALERT: hallefirstcolor;Toner Collection Unit HP CE265A;OK;HARD;3;
[1491823381] SERVICE ALERT: hallefirst3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1491823404] HOST ALERT: halleground1;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1491823477] SERVICE ALERT: hallefirstcolor;Image Fuser Kit HP 110V-CE246A;OK;HARD;3;Image Fuser Kit HP 110V-CE246A, 220V-CE247A is at 100% - OK!
[1491823494] HOST ALERT: halleground1;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1491823516] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;OK;HARD;3;Cyan Cartridge HP CE261A is at 79% - OK!
[1491823538] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1491823558] HOST ALERT: halleground1;UP;SOFT;10;PING OK - Packet loss = 0%, RTA = 0.58 ms
[1491823586] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;OK;HARD;3;Image Transfer Kit HP CE249A is at 12% - OK!
[1491823601] SERVICE ALERT: halleground1;Device Status;OK;HARD;3;Processing job from tray 4
[1491823695] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 90% (15 min average)
[1491823760] SERVICE ALERT: halleground2;Device Status;OK;HARD;3;Ready 164.76.106.52
[1491823823] SERVICE ALERT: halleground1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1491823878] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 73% - OK!
[1491823893] SERVICE ALERT: halleground2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1491824043] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 74% - OK!
[1491824248] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1491824308] SERVICE ALERT: hallefirst1;Device Status;OK;SOFT;2;Clearing paper path
[1491824649] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1491824679] HOST ALERT: hallefirst1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1491824709] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1491824739] HOST ALERT: hallefirst1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1491824769] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1491824799] HOST ALERT: hallefirst1;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1491824889] HOST ALERT: hallefirst1;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1491824895] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1491824914] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1491824944] HOST ALERT: hallefirst1;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1491824955] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1491824974] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1491825004] HOST ALERT: hallefirst1;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1491825034] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1491825064] HOST ALERT: hallefirst1;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1491825099] HOST FLAPPING ALERT: hallefirst3;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1491825103] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1491825133] HOST ALERT: hallefirst1;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1491825163] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1491825193] HOST ALERT: hallefirst1;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1491825223] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1491825253] HOST ALERT: hallefirst1;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1491826187] SERVICE ALERT: hallefirstcolor;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1491826229] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;Clearing paper path
[1491826233] HOST ALERT: hallefirst1;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.52 ms
[1491826242] SERVICE ALERT: hallefirstcolor;Device Status;OK;SOFT;2;Ready
[1491826418] SERVICE ALERT: hallefirst1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1491826563] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1491826834] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1491826864] HOST ALERT: hallefirst1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1491826894] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1491826924] HOST ALERT: hallefirst1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1491826954] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1491826954] SERVICE FLAPPING ALERT: hallefirst1;Device Status;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1491826984] HOST ALERT: hallefirst1;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1491827023] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1491827053] HOST ALERT: hallefirst1;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1491827083] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1491827113] HOST ALERT: hallefirst1;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1491827143] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1491827143] SERVICE FLAPPING ALERT: hallefirst1;Paper Tray Status;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1491827169] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1491827173] HOST ALERT: hallefirst1;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1491827229] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1491827259] HOST ALERT: hallefirst1;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1491827289] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1491827289] SERVICE FLAPPING ALERT: hallefirst1;Black Print Cartridge HP C8543X;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1491827319] HOST ALERT: hallefirst1;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1491827409] HOST ALERT: hallefirst1;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1491827499] HOST ALERT: hallefirst1;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1491828678] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5456 found since last run. oneLevelSearchOps OK. searchOps is too high! 5524 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491828738] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491829040] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5286 found since last run. oneLevelSearchOps OK. searchOps is too high! 5500 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491829100] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491829700] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6541 found since last run. oneLevelSearchOps OK. searchOps is too high! 6823 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491829760] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491830138] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 1 Trays Empty
[1491830198] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 1 Trays Empty
[1491830259] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 1 Trays Empty
[1491830360] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9648 found since last run. oneLevelSearchOps OK. searchOps is too high! 9936 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491830420] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491831003] SERVICE FLAPPING ALERT: hallefirst1;Black Print Cartridge HP C8543X;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1491831003] SERVICE FLAPPING ALERT: hallefirst1;Device Status;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1491831003] SERVICE FLAPPING ALERT: hallefirst1;Paper Tray Status;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1491831003] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1491831003] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1491831003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1491831003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1491831003] HOST DOWNTIME ALERT: INTLAPP;STARTED; Host has entered a period of scheduled downtime
[1491831003] HOST DOWNTIME ALERT: INTLDB;STARTED; Host has entered a period of scheduled downtime
[1491831003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1491831003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1491831003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1491831003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1491831003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1491831003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1491831003] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1491831003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1491831003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1491831003] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1491831003] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1491831020] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9496 found since last run. oneLevelSearchOps OK. searchOps is too high! 9780 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491831080] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491831138] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5727 found since last run. oneLevelSearchOps OK. searchOps is too high! 5949 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491831198] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491831519] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1491831549] HOST ALERT: prayharroldpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1491831572] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1491831579] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1491831602] HOST ALERT: prayharroldpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1491831632] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1491831639] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1491831662] HOST ALERT: prayharroldpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1491831680] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7094 found since last run. oneLevelSearchOps OK. searchOps is too high! 7360 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491831692] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1491831722] HOST ALERT: prayharroldpk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1491831740] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491831798] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5050 found since last run. oneLevelSearchOps OK. searchOps is too high! 5348 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491831812] HOST ALERT: prayharroldpk;DOWN;SOFT;5;(Host check timed out after 30.00 seconds)
[1491831835] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1491831858] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491831865] HOST ALERT: prayharroldpk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1491831905] SERVICE ALERT: CAVA2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1491831955] HOST ALERT: prayharroldpk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1491831965] SERVICE ALERT: CAVA2;CPU Load;OK;SOFT;2;CPU Load 0% (15 min average)
[1491832045] HOST ALERT: prayharroldpk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1491832135] HOST ALERT: prayharroldpk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1491832225] HOST ALERT: prayharroldpk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1491832340] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8205 found since last run. oneLevelSearchOps OK. searchOps is too high! 8470 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491832366] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1491832400] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491833000] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5207 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491833058] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5060 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491833060] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491833118] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491833361] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1491833420] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1491833480] SERVICE ALERT: halleground2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1491833660] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5081 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491833720] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491833858] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1491834230] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1491834234] HOST ALERT: prayharroldpk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1491834318] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5274 found since last run. oneLevelSearchOps OK. searchOps is too high! 5547 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491834320] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6488 found since last run. oneLevelSearchOps OK. searchOps is too high! 6789 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491834378] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491834380] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491834471] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1491834532] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1491834592] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1491834633] SERVICE ALERT: prayharroldpk;Device Status;OK;HARD;3;Processing job
[1491834687] SERVICE ALERT: prayharroldpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1491834978] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5324 found since last run. oneLevelSearchOps OK. searchOps is too high! 5646 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491834980] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9340 found since last run. oneLevelSearchOps OK. searchOps is too high! 9620 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491835038] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491835040] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491835280] SERVICE ALERT: halleground2;Device Status;OK;HARD;3;Ready 164.76.106.52
[1491835638] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 12838 found since last run. oneLevelSearchOps OK. searchOps is too high! 13149 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491835640] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8534 found since last run. oneLevelSearchOps OK. searchOps is too high! 8778 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491835698] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491835700] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491835880] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1491835941] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1491836000] SERVICE ALERT: halleground2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1491836298] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 16674 found since last run. oneLevelSearchOps OK. searchOps is too high! 16943 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491836300] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9864 found since last run. oneLevelSearchOps OK. searchOps is too high! 10150 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491836358] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491836360] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491836391] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1491836955] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1491836958] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8170 found since last run. oneLevelSearchOps OK. searchOps is too high! 8340 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491836960] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9773 found since last run. oneLevelSearchOps OK. searchOps is too high! 10068 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491837015] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1491837018] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491837020] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491837075] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1491837136] SERVICE ALERT: hallefirstcolor;Toner Collection Unit HP CE265A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1491837166] HOST ALERT: hallefirstcolor;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1491837196] SERVICE ALERT: hallefirstcolor;Toner Collection Unit HP CE265A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1491837226] HOST ALERT: hallefirstcolor;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1491837256] SERVICE ALERT: hallefirstcolor;Toner Collection Unit HP CE265A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1491837283] SERVICE ALERT: hallefirstcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1491837286] HOST ALERT: hallefirstcolor;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1491837321] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1491837343] SERVICE ALERT: hallefirstcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1491837351] HOST ALERT: hallefirstcolor;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1491837381] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1491837392] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1491837403] SERVICE ALERT: hallefirstcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1491837411] HOST ALERT: hallefirstcolor;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1491837441] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1491837452] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1491837461] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1491837464] SERVICE ALERT: hallefirstcolor;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1491837471] HOST ALERT: hallefirstcolor;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1491837492] HOST ALERT: hallefirstcolor;UP;SOFT;7;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1491837506] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;OK;SOFT;3;Image Transfer Kit HP CE249A is at 11% - OK!
[1491837515] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;OK;SOFT;2;Magenta Cartridge HP CE263A is at 49% - OK!
[1491837519] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! All Trays Still Full
[1491837618] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7304 found since last run. oneLevelSearchOps OK. searchOps is too high! 7395 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491837620] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8401 found since last run. oneLevelSearchOps OK. searchOps is too high! 8668 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491837640] SERVICE ALERT: hallefirstcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1491837675] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1491837678] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491837680] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491837700] SERVICE ALERT: hallefirstcolor;Device Status;OK;SOFT;2;Processing copy 1 of 4
[1491837850] SERVICE ALERT: hallefirstcolor;Toner Collection Unit HP CE265A;OK;HARD;3;
[1491837997] SERVICE ALERT: hallefirstcolor;Image Fuser Kit HP 110V-CE246A;OK;HARD;3;Image Fuser Kit HP 110V-CE246A, 220V-CE247A is at 100% - OK!
[1491838035] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;OK;HARD;3;Cyan Cartridge HP CE261A is at 78% - OK!
[1491838275] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1491838278] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7654 found since last run. oneLevelSearchOps OK. searchOps is too high! 7830 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491838280] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5107 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491838338] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491838340] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491838354] SERVICE FLAPPING ALERT: hallefirst1;Device Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1491838435] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1491838495] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1491838543] SERVICE FLAPPING ALERT: hallefirst1;Paper Tray Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1491838555] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1491838689] SERVICE FLAPPING ALERT: hallefirst1;Black Print Cartridge HP C8543X;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1491838875] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 90% (15 min average)
[1491838875] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1491838938] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5692 found since last run. oneLevelSearchOps OK. searchOps is too high! 5894 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491838940] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8234 found since last run. oneLevelSearchOps OK. searchOps is too high! 8450 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491838998] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491838999] SERVICE ALERT: halleground2;Device Status;OK;HARD;3;Processing job from tray 3
[1491838999] SERVICE FLAPPING ALERT: halleground2;Device Status;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1491839000] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491839185] HOST ALERT: prayharroldpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1491839249] HOST ALERT: prayharroldpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1491839475] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1491839535] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 61% (15 min average)
[1491839598] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5712 found since last run. oneLevelSearchOps OK. searchOps is too high! 5956 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491839600] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10830 found since last run. oneLevelSearchOps OK. searchOps is too high! 11096 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491839658] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491839660] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491839749] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 84% - OK!
[1491839749] SERVICE FLAPPING ALERT: prayharroldpk;Black Print Cartridge HP C8543X;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1491839914] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1491839974] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1491840035] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1491840084] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1491840144] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1491840204] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1491840258] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9734 found since last run. oneLevelSearchOps OK. searchOps is too high! 10032 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491840260] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9605 found since last run. oneLevelSearchOps OK. searchOps is too high! 9888 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491840319] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491840320] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491840705] SERVICE ALERT: eateriespk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1491840766] SERVICE ALERT: eateriespk;Device Status;OK;SOFT;2;50 72 65 73 73 20 17 20 74 6F 20 73 65 65 20 73  74 65 70  by step information
[1491840918] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 12129 found since last run. oneLevelSearchOps OK. searchOps is too high! 12224 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491840920] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9614 found since last run. oneLevelSearchOps OK. searchOps is too high! 9882 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491840978] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491840980] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491841577] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 15380 found since last run. oneLevelSearchOps OK. searchOps is too high! 15583 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491841580] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 12837 found since last run. oneLevelSearchOps OK. searchOps is too high! 13135 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491841637] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491841640] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491841803] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1491841803] SERVICE FLAPPING ALERT: halleground2;Device Status;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1491841803] SERVICE FLAPPING ALERT: prayharroldpk;Black Print Cartridge HP C8543X;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1491841803] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1491841803] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1491841803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1491841803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1491841803] HOST DOWNTIME ALERT: INTLAPP;STARTED; Host has entered a period of scheduled downtime
[1491841803] HOST DOWNTIME ALERT: INTLDB;STARTED; Host has entered a period of scheduled downtime
[1491841803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1491841803] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1491841803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1491841803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1491841803] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1491841803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1491841803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1491841803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1491841803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1491841803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1491841803] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1491841808] SERVICE ALERT: cava2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1491841868] SERVICE ALERT: cava2;CPU Load;OK;SOFT;2;CPU Load 1% (15 min average)
[1491842237] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10316 found since last run. oneLevelSearchOps OK. searchOps is too high! 10569 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491842240] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5368 found since last run. oneLevelSearchOps OK. searchOps is too high! 5629 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491842297] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491842300] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491842897] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6500 found since last run. oneLevelSearchOps OK. searchOps is too high! 6648 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491842900] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6583 found since last run. oneLevelSearchOps OK. searchOps is too high! 6832 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491842957] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491842960] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491843034] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;CRITICAL;HARD;3;Black Print Cartridge HP C8543X is at 1% - CRITICAL!
[1491843557] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5619 found since last run. oneLevelSearchOps OK. searchOps is too high! 5738 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491843560] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10106 found since last run. oneLevelSearchOps OK. searchOps is too high! 10395 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491843617] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491843620] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491843735] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1491843795] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 71% (15 min average)
[1491844217] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9592 found since last run. oneLevelSearchOps OK. searchOps is too high! 9703 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491844277] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491844404] SERVICE ALERT: studentcenterpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1491844820] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5021 found since last run. oneLevelSearchOps OK. searchOps is too high! 5271 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491844877] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 12311 found since last run. oneLevelSearchOps OK. searchOps is too high! 12460 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491844880] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491844937] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491844995] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1491845055] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1491845115] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1491845468] SERVICE ALERT: cava2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1491845480] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6700 found since last run. oneLevelSearchOps OK. searchOps is too high! 6976 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491845528] SERVICE ALERT: cava2;CPU Load;OK;SOFT;2;CPU Load 0% (15 min average)
[1491845537] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6243 found since last run. oneLevelSearchOps OK. searchOps is too high! 6415 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491845540] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491845597] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491846140] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10476 found since last run. oneLevelSearchOps OK. searchOps is too high! 10747 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491846197] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9575 found since last run. oneLevelSearchOps OK. searchOps is too high! 9729 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491846200] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491846257] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491846315] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1491846571] SERVICE ALERT: ADDC5;GPO Replication;WARNING;SOFT;1;Warning! 1 policies are not replicating properly!
[1491846800] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 11930 found since last run. oneLevelSearchOps OK. searchOps is too high! 12126 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491846857] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 14498 found since last run. oneLevelSearchOps OK. searchOps is too high! 14855 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491846859] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491846869] SERVICE ALERT: ADDC5;GPO Replication;WARNING;SOFT;2;Warning! 1 policies are not replicating properly!
[1491846915] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1491846917] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491847170] SERVICE ALERT: ADDC5;GPO Replication;WARNING;HARD;3;Warning! 1 policies are not replicating properly!
[1491847397] SERVICE ALERT: ADARC;Service: Group Policy;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1491847457] SERVICE ALERT: ADARC;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1491847459] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10571 found since last run. oneLevelSearchOps OK. searchOps is too high! 10707 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491847515] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 70% (15 min average)
[1491847517] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 11998 found since last run. oneLevelSearchOps OK. searchOps is too high! 12451 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491847519] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491847577] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491847834] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 76% - OK!
[1491847998] SERVICE FLAPPING ALERT: halleground2;Device Status;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1491848119] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8425 found since last run. oneLevelSearchOps OK. searchOps is too high! 8435 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491848177] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7075 found since last run. oneLevelSearchOps OK. searchOps is too high! 7358 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491848179] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491848237] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491848499] SERVICE ALERT: hallefirstcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1491848559] SERVICE ALERT: hallefirstcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1491848619] SERVICE ALERT: hallefirstcolor;Device Status;OK;SOFT;3;Open tray indicated.
[1491848715] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1491848775] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1491848779] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7869 found since last run. oneLevelSearchOps OK. searchOps is too high! 7883 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491848794] SERVICE ALERT: hallefirst2;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1491848835] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1491848837] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7268 found since last run. oneLevelSearchOps OK. searchOps is too high! 7456 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491848839] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491848854] SERVICE ALERT: hallefirst2;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1491848897] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491848914] SERVICE ALERT: hallefirst2;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1491849219] SERVICE ALERT: hallefirstcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1491849279] SERVICE ALERT: hallefirstcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1491849339] SERVICE ALERT: hallefirstcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1491849439] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8249 found since last run. oneLevelSearchOps OK. searchOps is too high! 8255 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491849497] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7682 found since last run. oneLevelSearchOps OK. searchOps is too high! 7953 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491849499] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491849557] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491849940] SERVICE ALERT: hallefirstcolor;Device Status;OK;HARD;3;Open tray indicated.
[1491850157] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9922 found since last run. oneLevelSearchOps OK. searchOps is too high! 10094 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491850217] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491850549] SERVICE FLAPPING ALERT: prayharroldpk;Black Print Cartridge HP C8543X;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1491850699] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5756 found since last run. oneLevelSearchOps OK. searchOps is too high! 5764 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491850759] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491850817] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9305 found since last run. oneLevelSearchOps OK. searchOps is too high! 9442 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491850877] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491851235] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 23% (15 min average)
[1491851314] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1491851359] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7706 found since last run. oneLevelSearchOps OK. searchOps is too high! 7713 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491851419] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491851477] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9968 found since last run. oneLevelSearchOps OK. searchOps is too high! 10193 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491851537] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491852019] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8720 found since last run. oneLevelSearchOps OK. searchOps is too high! 8729 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491852079] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491852137] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 11907 found since last run. oneLevelSearchOps OK. searchOps is too high! 12169 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491852197] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491852679] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10193 found since last run. oneLevelSearchOps OK. searchOps is too high! 10206 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491852739] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491852759] SERVICE ALERT: owen101lab1;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1491852762] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1491852771] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1491852797] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9755 found since last run. oneLevelSearchOps OK. searchOps is too high! 9954 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491852799] SERVICE ALERT: owen101lab1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1491852819] SERVICE ALERT: owen101lab1;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1491852822] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1491852830] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1491852857] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491852859] SERVICE ALERT: owen101lab1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1491852879] SERVICE ALERT: owen101lab1;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1491852882] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1491852889] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1491852895] SERVICE ALERT: owen101lab2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1491852913] SERVICE ALERT: owen101lab3;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1491852919] SERVICE ALERT: owen101lab1;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1491852939] SERVICE ALERT: hallefirstcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1491852955] SERVICE ALERT: owen101lab2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1491852973] SERVICE ALERT: owen101lab3;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1491852999] SERVICE ALERT: hallefirstcolor;Device Status;OK;SOFT;2;Then reinsert or close tray.
[1491853015] SERVICE ALERT: owen101lab2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1491853029] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1491853033] SERVICE ALERT: owen101lab3;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1491853035] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1491853089] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1491853095] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1491853149] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1491853155] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1491853339] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8460 found since last run. oneLevelSearchOps OK. searchOps is too high! 8468 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491853399] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491853457] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6113 found since last run. oneLevelSearchOps OK. searchOps is too high! 6338 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491853473] SERVICE ALERT: owen101lab1;Device Status;OK;HARD;3;Ready
[1491853477] SERVICE ALERT: owen101lab3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1491853483] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 82% - OK!
[1491853514] SERVICE ALERT: owen101lab1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1491853517] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491853610] SERVICE ALERT: owen101lab2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1491853627] SERVICE ALERT: owen101lab3;Device Status;OK;HARD;3;Ready
[1491853743] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 33% - OK!
[1491853999] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8456 found since last run. oneLevelSearchOps OK. searchOps is too high! 8466 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491854059] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491854355] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 73% (15 min average)
[1491854372] SERVICE ALERT: ADDC5;GPO Replication;OK;HARD;3;OK! All GPOs are properly replicated.
[1491854769] SERVICE ALERT: markjeffersonpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1491854799] HOST ALERT: markjeffersonpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1491854828] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! 2 of 3 Trays Still Full
[1491854832] HOST ALERT: markjeffersonpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.60 ms
[1491855355] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1491855385] HOST ALERT: prayharroldpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1491855400] SERVICE ALERT: hallefirstcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1491855409] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 39% - OK!
[1491855413] HOST ALERT: prayharroldpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1491855459] SERVICE ALERT: hallefirstcolor;Device Status;OK;SOFT;2;Checking printer
[1491855639] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1491855689] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1491855699] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1491855749] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1491855759] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1491855781] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.150.252. Is the printer off?  Is the network down?
[1491855798] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1491855809] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1491855836] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 47% - OK!
[1491855858] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1491855918] SERVICE ALERT: halleground2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1491856015] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1491856075] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1491856105] HOST ALERT: prayharroldpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1491856135] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1491856140] HOST ALERT: prayharroldpk;UP;SOFT;2;PING OK - Packet loss = 16%, RTA = 0.64 ms
[1491856140] HOST FLAPPING ALERT: prayharroldpk;STARTED; Host appears to have started flapping (20.3% change > 20.0% threshold)
[1491856155] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1491856215] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1491856275] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 60% (15 min average)
[1491856353] SERVICE ALERT: prayharroldpk;Device Status;OK;HARD;3;TRAY 3 EMPTY PLAIN LETTER   Ready
[1491856404] SERVICE ALERT: prayharroldpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1491856459] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6838 found since last run. oneLevelSearchOps OK. searchOps is too high! 6845 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491856517] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5393 found since last run. oneLevelSearchOps OK. searchOps is too high! 5749 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491856519] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491856577] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491856730] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 39% - OK!
[1491857119] SERVICE ALERT: halleground2;Device Status;OK;HARD;3;50 72 65 73 73 20 17 20 74 6F 20 73 65 65 20 73  74 65 70  by step information
[1491857119] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8654 found since last run. oneLevelSearchOps OK. searchOps is too high! 8660 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491857177] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8618 found since last run. oneLevelSearchOps OK. searchOps is too high! 8819 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491857179] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491857237] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491857718] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1491857777] SERVICE ALERT: halleground2;Device Status;OK;SOFT;2;Clearing paper path
[1491857779] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8419 found since last run. oneLevelSearchOps OK. searchOps is too high! 8427 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491857826] SERVICE ALERT: owen101lab3;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1491857837] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9567 found since last run. oneLevelSearchOps OK. searchOps is too high! 9985 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491857839] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491857887] SERVICE ALERT: owen101lab3;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1491857897] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491857946] SERVICE ALERT: owen101lab3;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1491858052] HOST ALERT: halleground2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1491858075] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1491858098] HOST ALERT: halleground2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.56 ms
[1491858135] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1491858195] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 79% (15 min average)
[1491858268] HOST FLAPPING ALERT: prayharroldpk;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1491858439] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7164 found since last run. oneLevelSearchOps OK. searchOps is too high! 7172 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491858497] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7268 found since last run. oneLevelSearchOps OK. searchOps is too high! 7506 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491858499] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491858557] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491859099] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5681 found since last run. oneLevelSearchOps OK. searchOps is too high! 5687 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491859157] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8577 found since last run. oneLevelSearchOps OK. searchOps is too high! 8844 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491859159] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491859217] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491859220] SERVICE ALERT: halleground1;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1491859280] SERVICE ALERT: halleground1;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! 1 of 3 Trays Still Full
[1491859395] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1491859455] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1491859515] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1491859624] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1491859684] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1491859744] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1491859759] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6838 found since last run. oneLevelSearchOps OK. searchOps is too high! 6846 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491859791] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1491859817] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9062 found since last run. oneLevelSearchOps OK. searchOps is too high! 9406 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491859819] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491859851] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1491859877] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491859880] SERVICE ALERT: halleground1;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1491859910] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1491859940] SERVICE ALERT: halleground1;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1491859999] SERVICE ALERT: halleground1;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1491860115] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 84% (15 min average)
[1491860478] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5134 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491860537] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491860604] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1491860664] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1491860715] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1491860724] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1491860775] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1491860777] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1491860835] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1491860838] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1491860897] SERVICE ALERT: halleground2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1491861019] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5305 found since last run. oneLevelSearchOps OK. searchOps is too high! 5309 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491861079] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491861546] SERVICE ALERT: owen101lab3;Device Status;OK;HARD;3;Clearing paper path
[1491861546] SERVICE FLAPPING ALERT: owen101lab3;Device Status;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1491861727] SERVICE ALERT: cava2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1491861736] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5382 found since last run. oneLevelSearchOps OK. searchOps is too high! 5516 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491861787] SERVICE ALERT: cava2;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1491861796] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491861847] SERVICE ALERT: cava2;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1491862035] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1491862279] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7193 found since last run. oneLevelSearchOps OK. searchOps is too high! 7203 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491862339] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491862447] SERVICE ALERT: cava2;CPU Load;OK;HARD;3;CPU Load 1% (15 min average)
[1491862635] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1491862695] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1491862996] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6159 found since last run. oneLevelSearchOps OK. searchOps is too high! 6301 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491863056] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491863295] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1491863355] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1491863415] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1491863599] SERVICE ALERT: halleground1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1491863656] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7274 found since last run. oneLevelSearchOps OK. searchOps is too high! 7575 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491863716] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491864139] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6304 found since last run. oneLevelSearchOps OK. searchOps is too high! 6313 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491864199] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491864799] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8041 found since last run. oneLevelSearchOps OK. searchOps is too high! 8051 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491864859] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491865144] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1491865398] SERVICE ALERT: halleground1;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1491865458] SERVICE ALERT: halleground1;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1491865518] SERVICE ALERT: halleground1;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1491865744] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1491865804] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1491865864] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1491866153] SERVICE ALERT: hallegroundcolor;Image Transfer Kit HP CE249A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1491866207] SERVICE ALERT: hallegroundcolor;Image Transfer Kit HP CE249A;OK;SOFT;2;Image Transfer Kit HP CE249A is at 45% - OK!
[1491866415] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1491866643] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1491866703] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 73% - OK!
[1491866716] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7456 found since last run. oneLevelSearchOps OK. searchOps is too high! 7739 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491866776] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491867615] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1491867675] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1491867735] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1491867976] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5044 found since last run. oneLevelSearchOps OK. searchOps is too high! 5213 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491868036] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491868636] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5054 found since last run. oneLevelSearchOps OK. searchOps is too high! 5219 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491868696] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491869059] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6117 found since last run. oneLevelSearchOps OK. searchOps is too high! 6126 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491869119] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491869296] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5336 found since last run. oneLevelSearchOps OK. searchOps is too high! 5523 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491869356] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491869719] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7522 found since last run. oneLevelSearchOps OK. searchOps is too high! 7546 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491869779] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491869945] SERVICE FLAPPING ALERT: owen101lab3;Device Status;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1491870135] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 87% (15 min average)
[1491870318] SERVICE ALERT: halleground1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1491870318] SERVICE FLAPPING ALERT: halleground1;Paper Tray Status;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1491870379] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9625 found since last run. oneLevelSearchOps OK. searchOps is too high! 9638 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491870439] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491870496] SERVICE ALERT: halleground2;Device Status;OK;HARD;3;Ready 164.76.106.52
[1491871039] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5800 found since last run. oneLevelSearchOps OK. searchOps is too high! 5808 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491871099] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491871335] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1491871395] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1491871455] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1491871524] SERVICE ALERT: studentcenterpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1491871699] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7804 found since last run. oneLevelSearchOps OK. searchOps is too high! 7820 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491871759] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491872154] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1491872213] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1491872274] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1491872359] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5443 found since last run. oneLevelSearchOps OK. searchOps is too high! 5449 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491872419] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491873255] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 82% (15 min average)
[1491873855] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1491873915] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1491873975] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1491874143] SERVICE ALERT: TARGETX-2012;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1491874200] HOST DOWNTIME ALERT: IMAGES;STARTED; Host has entered a period of scheduled downtime
[1491874203] SERVICE ALERT: TARGETX-2012;CPU Load;OK;SOFT;2;CPU Load 88% (15 min average)
[1491874575] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 65% (15 min average)
[1491875175] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1491875199] SERVICE ALERT: IMAGES;MDT Monitor;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1491875205] HOST ALERT: IMAGES;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1491875235] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1491875295] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1491875295] HOST ALERT: IMAGES;DOWN;HARD;2;PING CRITICAL - Packet loss = 100%
[1491875305] SERVICE ALERT: IMAGES;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1491875313] HOST ALERT: IMAGES;UP;HARD;2;PING OK - Packet loss = 44%, RTA = 0.22 ms
[1491875345] SERVICE ALERT: IMAGES;CPU Load;OK;SOFT;2;CPU Load 0% (15 min average)
[1491875356] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6486 found since last run. oneLevelSearchOps OK. searchOps is too high! 6688 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491875416] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491875779] SERVICE ALERT: IMAGES;MDT Monitor;UNKNOWN;HARD;3;Failed to open service MDT_Monitor: 1060: The specified service does not exist as an installed service.
[1491875858] SERVICE ALERT: hallefirstcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1491875895] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 84% (15 min average)
[1491875918] SERVICE ALERT: hallefirstcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1491875978] SERVICE ALERT: hallefirstcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1491876579] SERVICE ALERT: hallefirstcolor;Device Status;OK;HARD;3;Cleaning...
[1491877216] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5233 found since last run. oneLevelSearchOps OK. searchOps is too high! 5539 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491877276] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491877518] SERVICE FLAPPING ALERT: halleground1;Paper Tray Status;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1491877674] SERVICE ALERT: eateriespk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1491878417] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5904 found since last run. oneLevelSearchOps OK. searchOps is too high! 5917 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491878477] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491878895] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1491878955] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 86% (15 min average)
[1491878978] SERVICE ALERT: hallefirstcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1491879038] SERVICE ALERT: hallefirstcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1491879098] SERVICE ALERT: hallefirstcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1491879555] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1491879615] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1491879675] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1491879698] SERVICE ALERT: hallefirstcolor;Device Status;OK;HARD;3;Open right door.
[1491879698] SERVICE FLAPPING ALERT: hallefirstcolor;Device Status;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1491879971] HOST ALERT: hallefirst1;UP;HARD;10;PING WARNING - Packet loss = 93%, RTA = 0.53 ms
[1491880083] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1491880348] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;41.3 UNEXPECTED SIZE IN TRAY 2 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   LOAD TRAY 2 PLAIN LETTER To use another 74 72 61 79 20 70 72 65 73 73 20 1E
[1491880536] SERVICE ALERT: hallefirst1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1491880716] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 1 Trays Empty
[1491880776] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 1 Trays Empty
[1491880837] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;SOFT;3;Printer Paper Is OK! All Trays Still Full
[1491881400] HOST DOWNTIME ALERT: IMAGES;STOPPED; Host has exited from a period of scheduled downtime
[1491882098] SERVICE ALERT: hallefirstcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1491882158] SERVICE ALERT: hallefirstcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1491882218] SERVICE ALERT: hallefirstcolor;Device Status;OK;SOFT;3;BLACK CARTRIDGE LOW 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Processing...
[1491882749] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1491882808] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1491882867] SERVICE ALERT: hallefirst1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!