[1489037060] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 55% (15 min average)
[1489037400] HOST DOWNTIME ALERT: PROCARE;STOPPED; Host has exited from a period of scheduled downtime
[1489037400] HOST DOWNTIME ALERT: OFFICEKMS;STOPPED; Host has exited from a period of scheduled downtime
[1489037400] HOST DOWNTIME ALERT: TABLEAU;STOPPED; Host has exited from a period of scheduled downtime
[1489037400] HOST DOWNTIME ALERT: R25-ALPHA2012;STOPPED; Host has exited from a period of scheduled downtime
[1489037400] HOST DOWNTIME ALERT: OPIE;STOPPED; Host has exited from a period of scheduled downtime
[1489037400] HOST DOWNTIME ALERT: PNCWEB-2012;STOPPED; Host has exited from a period of scheduled downtime
[1489037400] HOST DOWNTIME ALERT: XDM;STOPPED; Host has exited from a period of scheduled downtime
[1489037400] HOST DOWNTIME ALERT: PNCDB-2012;STOPPED; Host has exited from a period of scheduled downtime
[1489037400] HOST DOWNTIME ALERT: WINKMS;STOPPED; Host has exited from a period of scheduled downtime
[1489037400] HOST DOWNTIME ALERT: PNCAPP-2012;STOPPED; Host has exited from a period of scheduled downtime
[1489037400] HOST DOWNTIME ALERT: EMUSCRIPT;STOPPED; Host has exited from a period of scheduled downtime
[1489037400] HOST DOWNTIME ALERT: UHSPHARM-2012;STOPPED; Host has exited from a period of scheduled downtime
[1489037400] HOST DOWNTIME ALERT: TI22-PSYCLINIC;STOPPED; Host has exited from a period of scheduled downtime
[1489037400] HOST DOWNTIME ALERT: NETENG-PA;STOPPED; Host has exited from a period of scheduled downtime
[1489037400] HOST DOWNTIME ALERT: NESSUS;STOPPED; Host has exited from a period of scheduled downtime
[1489037400] HOST DOWNTIME ALERT: MBAMDB;STOPPED; Host has exited from a period of scheduled downtime
[1489037400] HOST DOWNTIME ALERT: HDBACKUPS;STOPPED; Host has exited from a period of scheduled downtime
[1489037400] HOST DOWNTIME ALERT: TI22-COE;STOPPED; Host has exited from a period of scheduled downtime
[1489037400] HOST DOWNTIME ALERT: MBAM;STOPPED; Host has exited from a period of scheduled downtime
[1489037400] HOST DOWNTIME ALERT: EMUSQLEXPRESS;STOPPED; Host has exited from a period of scheduled downtime
[1489037400] HOST DOWNTIME ALERT: FTK;STOPPED; Host has exited from a period of scheduled downtime
[1489037400] HOST DOWNTIME ALERT: R25-DELTA2012;STOPPED; Host has exited from a period of scheduled downtime
[1489037400] HOST DOWNTIME ALERT: TI22;STOPPED; Host has exited from a period of scheduled downtime
[1489037400] HOST DOWNTIME ALERT: DPS-SECURE02;STOPPED; Host has exited from a period of scheduled downtime
[1489037400] HOST DOWNTIME ALERT: COE-WEBDAV;STOPPED; Host has exited from a period of scheduled downtime
[1489037400] HOST DOWNTIME ALERT: SCCM-PRIM01;STOPPED; Host has exited from a period of scheduled downtime
[1489037400] HOST DOWNTIME ALERT: DIRECTACCESS;STOPPED; Host has exited from a period of scheduled downtime
[1489037400] HOST DOWNTIME ALERT: REFWIN;STOPPED; Host has exited from a period of scheduled downtime
[1489037400] HOST DOWNTIME ALERT: SANDBOX-DB;STOPPED; Host has exited from a period of scheduled downtime
[1489037400] HOST DOWNTIME ALERT: CEDARON;STOPPED; Host has exited from a period of scheduled downtime
[1489037400] HOST DOWNTIME ALERT: ADFS;STOPPED; Host has exited from a period of scheduled downtime
[1489037400] HOST DOWNTIME ALERT: CBORD;STOPPED; Host has exited from a period of scheduled downtime
[1489037400] HOST DOWNTIME ALERT: R25-ZETA2012;STOPPED; Host has exited from a period of scheduled downtime
[1489037400] HOST DOWNTIME ALERT: R25-INTERFACE;STOPPED; Host has exited from a period of scheduled downtime
[1489037400] HOST DOWNTIME ALERT: 4WINDSMANAGER;STOPPED; Host has exited from a period of scheduled downtime
[1489037400] HOST DOWNTIME ALERT: R25-WEBSERVICES;STOPPED; Host has exited from a period of scheduled downtime
[1489037400] HOST DOWNTIME ALERT: R25-WEBVIEWER;STOPPED; Host has exited from a period of scheduled downtime
[1489037401] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1489038261] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1489038320] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1489038380] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1489038979] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1489039968] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1489040029] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1489040179] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 37% (15 min average)
[1489040555] SERVICE FLAPPING ALERT: AD-RMAD;E:\ Drive Space;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1489041229] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1489041288] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;2;Warning:  idm1  are more than 120 seconds behind!
[1489041349] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1489041378] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1489041438] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1489042038] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1489042098] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 75% (15 min average)
[1489042698] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1489042758] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 91% (15 min average)
[1489043426] SERVICE ALERT: ADDC5;AD Group Changes;UNKNOWN;SOFT;1;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1489043503] SERVICE FLAPPING ALERT: hallefirst3;Black Print Cartridge HP C8543X;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1489043525] SERVICE FLAPPING ALERT: hallefirst2;Black Print Cartridge HP C8543X;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1489043657] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1489043703] SERVICE FLAPPING ALERT: hallefirst2;Paper Tray Status;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1489043704] SERVICE ALERT: ADDC5;AD Group Changes;OK;SOFT;2;WARNING:  0  groups have changed!
[1489043958] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1489044018] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1489044079] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 87% (15 min average)
[1489044348] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1489044678] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1489044738] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1489044798] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1489045548] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1489045905] HOST ALERT: owen101lab1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1489045995] HOST ALERT: owen101lab1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1489046019] SERVICE ALERT: owen101lab1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1489046049] HOST ALERT: owen101lab1;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1489046079] SERVICE ALERT: owen101lab1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1489046109] HOST ALERT: owen101lab1;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1489046139] SERVICE ALERT: owen101lab1;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1489046141] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1489046169] HOST ALERT: owen101lab1;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1489046201] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1489046231] HOST ALERT: owen101lab1;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1489046261] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1489046291] HOST ALERT: owen101lab1;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1489046381] HOST ALERT: owen101lab1;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1489046444] SERVICE ALERT: owen101lab1;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1489046471] HOST ALERT: owen101lab1;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1489046504] SERVICE ALERT: owen101lab1;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1489046534] HOST ALERT: owen101lab1;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1489046564] SERVICE ALERT: owen101lab1;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1489046780] SERVICE ALERT: BBTS_DB;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1489046840] SERVICE ALERT: BBTS_DB;CPU Load;OK;SOFT;2;CPU Load 83% (15 min average)
[1489047198] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 91% (15 min average)
[1489047798] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1489047858] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1489047918] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1489049118] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1489049718] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1489050318] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 66% (15 min average)
[1489050347] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1489050347] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1489051800] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service has exited from a period of scheduled downtime
[1489051803] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (37.0% change >= 20.0% threshold)
[1489051803] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1489051803] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1489051803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1489051803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1489051803] HOST DOWNTIME ALERT: BBTS-DEL;STARTED; Host has entered a period of scheduled downtime
[1489051803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1489051803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1489051803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1489051803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1489051803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1489051803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1489051803] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1489051803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1489051803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1489051803] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1489051803] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1489052718] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1489052778] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1489052837] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1489053437] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 51% (15 min average)
[1489053603] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (43.2% change >= 20.0% threshold)
[1489053603] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1489053603] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1489053603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1489053603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1489053603] HOST DOWNTIME ALERT: BBTS-DEL;STARTED; Host has entered a period of scheduled downtime
[1489053603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1489053603] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1489053603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1489053603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1489053603] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1489053603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1489053603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1489053603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1489053603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1489053603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1489053603] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1489055837] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1489055897] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1489055957] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 94% (15 min average)
[1489057757] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1489057817] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1489057877] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1489060204] SERVICE ALERT: markjeffersonpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1489060234] HOST ALERT: markjeffersonpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1489060264] SERVICE ALERT: markjeffersonpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1489060294] HOST ALERT: markjeffersonpk;UP;SOFT;2;PING WARNING - Packet loss = 90%, RTA = 0.74 ms
[1489060319] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;SOFT;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1489060919] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1489060979] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1489061039] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1489062077] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1489062153] SERVICE ALERT: EMUSCRIPT;Service: Group Policy;CRITICAL;SOFT;1;gpsvc: Stopped
[1489062157] SERVICE ALERT: EMUSCRIPT;VMWare Tools;WARNING;SOFT;1;VMTools: Error
[1489062213] SERVICE ALERT: EMUSCRIPT;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1489062217] SERVICE ALERT: EMUSCRIPT;VMWare Tools;OK;SOFT;2;VMTools: Started
[1489062260] SERVICE ALERT: strongpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1489062320] SERVICE ALERT: strongpk;Device Status;OK;SOFT;2;Open right  door.
[1489062677] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1489062737] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1489062797] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1489063397] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 49% (15 min average)
[1489063439] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1489064039] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1489064099] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1489064159] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1489065197] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1489065257] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 52% (15 min average)
[1489065377] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1489065436] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1489065495] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1489065857] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1489065917] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1489066517] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1489066520] SERVICE ALERT: strongpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1489066577] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1489066580] SERVICE ALERT: strongpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1489066637] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 94% (15 min average)
[1489066640] SERVICE ALERT: strongpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1489067240] SERVICE ALERT: strongpk;Device Status;OK;HARD;3;TRAY 2 OPEN  46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1489067562] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5269 found since last run. oneLevelSearchOps OK. searchOps is too high! 5499 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489067622] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489067840] SERVICE ALERT: strongpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1489067901] SERVICE ALERT: strongpk;Device Status;OK;SOFT;2;RETRIEVE JOB  NO STORED JOBS
[1489068463] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7395 found since last run. oneLevelSearchOps OK. searchOps is too high! 7598 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489068495] SERVICE ALERT: markjeffersonpk;Device Status;OK;HARD;3;Warming up
[1489068523] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489068939] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1489068999] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1489069058] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1489069096] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1489069123] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8813 found since last run. oneLevelSearchOps OK. searchOps is too high! 9022 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489069155] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1489069183] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489069215] SERVICE ALERT: markjeffersonpk;Device Status;OK;SOFT;3;50 72 65 73 73 20 17 20 74 6F 20 73 65 65 20 73  74 65 70  by step information
[1489069803] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1489069803] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1489069803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1489069803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1489069803] HOST DOWNTIME ALERT: BBTS-DEL;STARTED; Host has entered a period of scheduled downtime
[1489069803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1489069803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1489069803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1489069803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1489069803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1489069803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1489069803] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1489069803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1489069803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1489069803] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1489069803] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1489070022] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8633 found since last run. oneLevelSearchOps OK. searchOps is too high! 8966 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489070082] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489070158] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1489070158] SERVICE FLAPPING ALERT: markjeffersonpk;Paper Tray Status;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1489070837] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1489070897] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1489070920] SERVICE ALERT: owen101lab2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1489070924] HOST ALERT: owen101lab2;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.46 ms
[1489070933] SERVICE ALERT: owen101lab2;Device Status;OK;HARD;3;Warming up
[1489070957] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 87% (15 min average)
[1489071156] SERVICE ALERT: owen101lab1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1489071160] HOST ALERT: owen101lab1;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.54 ms
[1489071214] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at WARNING level!
[1489071282] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6253 found since last run. oneLevelSearchOps OK. searchOps is too high! 6532 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489071333] SERVICE ALERT: owen101lab1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1489071342] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489071454] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 78% - OK!
[1489071735] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1489071795] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1489071855] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1489071942] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7003 found since last run. oneLevelSearchOps OK. searchOps is too high! 7253 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489072002] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489072039] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1489072093] SERVICE ALERT: prayharroldpk;Device Status;OK;SOFT;2;Processing job from tray 3
[1489072133] SERVICE ALERT: owen101lab2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1489072183] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5000 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489072193] SERVICE ALERT: owen101lab2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1489072243] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489072254] SERVICE ALERT: owen101lab2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1489072450] SERVICE ALERT: prayharroldpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1489072602] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5588 found since last run. oneLevelSearchOps OK. searchOps is too high! 5751 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489072662] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489072843] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6102 found since last run. oneLevelSearchOps OK. searchOps is too high! 6321 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489072903] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489073123] SERVICE ALERT: porterpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1489073182] SERVICE ALERT: porterpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1489073189] SERVICE ALERT: porterpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1489073241] SERVICE ALERT: porterpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1489073249] SERVICE ALERT: porterpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1489073262] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6471 found since last run. oneLevelSearchOps OK. searchOps is too high! 6683 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489073293] SERVICE ALERT: studentcenterpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.150.252. Is the printer off?  Is the network down?
[1489073309] SERVICE ALERT: porterpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1489073322] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489073353] SERVICE ALERT: studentcenterpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.150.252. Is the printer off?  Is the network down?
[1489073413] SERVICE ALERT: studentcenterpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.150.252. Is the printer off?  Is the network down?
[1489073443] HOST ALERT: studentcenterpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1489073503] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7454 found since last run. oneLevelSearchOps OK. searchOps is too high! 7649 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489073507] HOST ALERT: studentcenterpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.45 ms
[1489073563] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489073649] SERVICE ALERT: porterpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1489073681] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1489073709] SERVICE ALERT: porterpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1489073769] SERVICE ALERT: porterpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1489073779] SERVICE ALERT: studentcenterpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.150.252. Is the printer off?  Is the network down?
[1489073795] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.150.252. Is the printer off?  Is the network down?
[1489073839] SERVICE ALERT: studentcenterpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.150.252. Is the printer off?  Is the network down?
[1489073855] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.150.252. Is the printer off?  Is the network down?
[1489073858] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1489073899] SERVICE ALERT: studentcenterpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.150.252. Is the printer off?  Is the network down?
[1489073915] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.150.252. Is the printer off?  Is the network down?
[1489073922] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7318 found since last run. oneLevelSearchOps OK. searchOps is too high! 7497 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489073982] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489074163] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9846 found since last run. oneLevelSearchOps OK. searchOps is too high! 10107 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489074223] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489074471] HOST ALERT: porterpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1489074504] SERVICE ALERT: porterpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1489074508] HOST ALERT: porterpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.56 ms
[1489074582] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9229 found since last run. oneLevelSearchOps OK. searchOps is too high! 9383 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489074642] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489074756] SERVICE ALERT: owen101lab1;Device Status;OK;HARD;3;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1489074814] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1489074823] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 12285 found since last run. oneLevelSearchOps OK. searchOps is too high! 12513 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489074883] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489074963] SERVICE ALERT: porterpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 74% - OK!
[1489075035] SERVICE ALERT: porterpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1489075157] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1489075217] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1489075242] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10162 found since last run. oneLevelSearchOps OK. searchOps is too high! 10377 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489075277] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1489075302] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489075369] HOST ALERT: porterpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1489075414] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at WARNING level!
[1489075433] HOST ALERT: porterpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.53 ms
[1489075433] HOST FLAPPING ALERT: porterpk;STARTED; Host appears to have started flapping (22.8% change > 20.0% threshold)
[1489075483] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10015 found since last run. oneLevelSearchOps OK. searchOps is too high! 10252 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489075543] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489075635] SERVICE ALERT: porterpk;Device Status;OK;HARD;3;Ready
[1489075663] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1489075724] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1489075784] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1489075877] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 87% (15 min average)
[1489075902] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5542 found since last run. oneLevelSearchOps OK. searchOps is too high! 5804 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489075962] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489076143] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8615 found since last run. oneLevelSearchOps OK. searchOps is too high! 8844 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489076159] SERVICE FLAPPING ALERT: markjeffersonpk;Paper Tray Status;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1489076203] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489076477] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1489076537] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1489076597] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1489076614] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 38% - OK!
[1489076614] SERVICE FLAPPING ALERT: owen101lab2;Black Print Cartridge HP C8543X;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1489076803] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9322 found since last run. oneLevelSearchOps OK. searchOps is too high! 9494 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489076863] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489076893] SERVICE ALERT: studentcenterpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1489076909] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 8% - OK!
[1489077007] SERVICE ALERT: studentcenterpk;Device Status;OK;HARD;3;ORDER CARTRIDGE LESS THAN 1800 PAGES 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1489077053] SERVICE ALERT: owen101lab2;Device Status;OK;HARD;3;There is no help available for this item. 54 6F 20 65 78 69 74 20 70 72 65 73 73 20 1F 2E
[1489077463] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7546 found since last run. oneLevelSearchOps OK. searchOps is too high! 7750 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489077523] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489077560] HOST FLAPPING ALERT: porterpk;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1489077652] SERVICE ALERT: owen101lab2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1489077663] SERVICE ALERT: sillpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1489077712] SERVICE ALERT: owen101lab2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1489077724] SERVICE ALERT: sillpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1489077772] SERVICE ALERT: owen101lab2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1489077772] SERVICE FLAPPING ALERT: owen101lab2;Device Status;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1489077783] SERVICE ALERT: sillpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1489078123] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8724 found since last run. oneLevelSearchOps OK. searchOps is too high! 8919 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489078183] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489078227] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1489078287] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1489078347] SERVICE ALERT: mckennypk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1489078349] SERVICE ALERT: sillpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1489078362] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6284 found since last run. oneLevelSearchOps OK. searchOps is too high! 6644 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489078390] SERVICE ALERT: sillpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1489078397] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 59% (15 min average)
[1489078397] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1489078409] SERVICE ALERT: sillpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1489078422] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489078469] SERVICE ALERT: sillpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1489078483] SERVICE ALERT: sillpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1489078543] SERVICE ALERT: sillpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1489078603] SERVICE ALERT: sillpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1489078783] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10113 found since last run. oneLevelSearchOps OK. searchOps is too high! 10328 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489078843] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489078947] SERVICE ALERT: mckennypk;Device Status;OK;HARD;3;CLOSE RIGHT DOOR  46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1489078997] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1489079022] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7418 found since last run. oneLevelSearchOps OK. searchOps is too high! 7581 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489079057] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 88% (15 min average)
[1489079082] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489079149] SERVICE ALERT: mckennypk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1489079209] SERVICE ALERT: mckennypk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1489079269] SERVICE ALERT: mckennypk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1489079443] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6785 found since last run. oneLevelSearchOps OK. searchOps is too high! 6992 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489079503] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489079657] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1489079682] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 11012 found since last run. oneLevelSearchOps OK. searchOps is too high! 11291 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489079695] HOST ALERT: crossroadspk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1489079717] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1489079742] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489079749] HOST ALERT: crossroadspk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1489079758] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1489079777] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1489079805] HOST ALERT: crossroadspk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1489079818] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1489079869] SERVICE ALERT: mckennypk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1489079875] SERVICE ALERT: eateriespk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1489079878] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1489079895] HOST ALERT: crossroadspk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1489079935] SERVICE ALERT: eateriespk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1489079946] SERVICE ALERT: marshallpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.137.49. Is the printer off?  Is the network down?
[1489079985] HOST ALERT: crossroadspk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1489079989] SERVICE ALERT: marshallpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.137.49. Is the printer off?  Is the network down?
[1489079995] SERVICE ALERT: eateriespk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1489080006] SERVICE ALERT: marshallpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.137.49. Is the printer off?  Is the network down?
[1489080009] SERVICE ALERT: marshallpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.137.49. Is the printer off?  Is the network down?
[1489080066] SERVICE ALERT: marshallpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.137.49. Is the printer off?  Is the network down?
[1489080069] SERVICE ALERT: marshallpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.137.49. Is the printer off?  Is the network down?
[1489080075] HOST ALERT: crossroadspk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1489080103] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9677 found since last run. oneLevelSearchOps OK. searchOps is too high! 9908 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489080129] SERVICE ALERT: marshallpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.137.49. Is the printer off?  Is the network down?
[1489080163] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489080165] HOST ALERT: crossroadspk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1489080255] HOST ALERT: crossroadspk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1489080266] HOST ALERT: honorspk;UP;HARD;10;PING OK - Packet loss = 44%, RTA = 6.33 ms
[1489080295] HOST ALERT: crossroadspk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1489080342] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 16214 found since last run. oneLevelSearchOps OK. searchOps is too high! 16344 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489080349] HOST ALERT: crossroadspk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1489080402] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489080443] SERVICE ALERT: honorspk;Device Status;OK;HARD;3;Ready
[1489080596] SERVICE ALERT: eateriespk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 49% - OK!
[1489080603] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (28.0% change >= 20.0% threshold)
[1489080603] SERVICE FLAPPING ALERT: owen101lab2;Black Print Cartridge HP C8543X;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1489080603] SERVICE FLAPPING ALERT: owen101lab2;Device Status;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1489080603] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1489080603] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1489080603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1489080603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1489080603] HOST DOWNTIME ALERT: BBTS-DEL;STARTED; Host has entered a period of scheduled downtime
[1489080603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1489080603] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1489080603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1489080603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1489080603] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1489080603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1489080603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1489080603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1489080603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1489080603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1489080603] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1489080633] SERVICE ALERT: honorspk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1489080763] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8464 found since last run. oneLevelSearchOps OK. searchOps is too high! 8670 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489080823] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489080824] SERVICE ALERT: BOEAPPTEST;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1489080852] SERVICE ALERT: honorspk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 51% - OK!
[1489080889] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1489080944] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1489080977] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1489081002] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 13000 found since last run. oneLevelSearchOps OK. searchOps is too high! 13074 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489081062] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489081078] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1489081399] SERVICE ALERT: hallegroundcolor;Toner Collection Unit HP CE265A;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Toner Collection Unit HP CE265A'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1489081454] SERVICE ALERT: hallegroundcolor;Toner Collection Unit HP CE265A;OK;SOFT;2;
[1489081544] SERVICE ALERT: BOEAPPTEST;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1489081662] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9052 found since last run. oneLevelSearchOps OK. searchOps is too high! 9173 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489081678] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1489081706] HOST ALERT: owen101lab2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1489081722] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489081724] SERVICE ALERT: owen101lab2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1489081737] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1489081754] HOST ALERT: owen101lab2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1489081784] SERVICE ALERT: owen101lab2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1489081797] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1489081797] SERVICE FLAPPING ALERT: markjeffersonpk;Paper Tray Status;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1489081815] HOST ALERT: owen101lab2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1489081844] SERVICE ALERT: owen101lab2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1489081874] HOST ALERT: owen101lab2;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1489081959] HOST ALERT: marshallpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1489081964] HOST ALERT: owen101lab2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1489081978] SERVICE ALERT: owen101lab2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1489082008] HOST ALERT: owen101lab2;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1489082018] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1489082023] HOST ALERT: marshallpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.55 ms
[1489082023] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7123 found since last run. oneLevelSearchOps OK. searchOps is too high! 7318 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489082048] HOST ALERT: owen101lab2;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1489082078] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1489082083] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489082108] HOST ALERT: owen101lab2;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1489082138] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1489082144] SERVICE ALERT: BOEAPPTEST;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1489082168] HOST ALERT: owen101lab2;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1489082258] HOST ALERT: owen101lab2;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1489082298] SERVICE ALERT: strongpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1489082322] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10559 found since last run. oneLevelSearchOps OK. searchOps is too high! 10706 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489082353] HOST ALERT: marshallpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1489082358] SERVICE ALERT: strongpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1489082382] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489082389] SERVICE ALERT: marshallpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1489082393] HOST ALERT: marshallpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.55 ms
[1489082393] HOST FLAPPING ALERT: marshallpk;STARTED; Host appears to have started flapping (23.8% change > 20.0% threshold)
[1489082418] SERVICE ALERT: strongpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1489082460] SERVICE ALERT: marshallpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 11% - OK!
[1489082523] SERVICE ALERT: marshallpk;Device Status;OK;HARD;3;Ready
[1489082683] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7406 found since last run. oneLevelSearchOps OK. searchOps is too high! 7602 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489082743] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489082911] SERVICE ALERT: eateriespk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1489082982] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7463 found since last run. oneLevelSearchOps OK. searchOps is too high! 7675 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489083042] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489083343] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5887 found since last run. oneLevelSearchOps OK. searchOps is too high! 6108 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489083403] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489083642] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5582 found since last run. oneLevelSearchOps OK. searchOps is too high! 5794 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489083702] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489084003] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8202 found since last run. oneLevelSearchOps OK. searchOps is too high! 8430 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489084063] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489084218] SERVICE ALERT: strongpk;Device Status;OK;HARD;3;Warming up
[1489084302] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8977 found since last run. oneLevelSearchOps OK. searchOps is too high! 9037 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489084362] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489084521] HOST FLAPPING ALERT: marshallpk;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1489084663] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6272 found since last run. oneLevelSearchOps OK. searchOps is too high! 6482 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489084723] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489084962] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8976 found since last run. oneLevelSearchOps OK. searchOps is too high! 9086 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489085022] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489085081] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 1 Trays Empty
[1489085141] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 1 Trays Empty
[1489085177] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1489085200] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 1 Trays Empty
[1489085237] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1489085297] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1489085323] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10224 found since last run. oneLevelSearchOps OK. searchOps is too high! 10392 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489085383] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489085622] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8809 found since last run. oneLevelSearchOps OK. searchOps is too high! 9015 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489085682] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489085899] SERVICE ALERT: strongpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.130.98. Is the printer off?  Is the network down?
[1489085953] SERVICE ALERT: strongpk;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! 2 of 3 Trays Still Full
[1489085983] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 11151 found since last run. oneLevelSearchOps OK. searchOps is too high! 11348 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489086043] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489086115] SERVICE ALERT: idmrbs.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 58307 found since last run. oneLevelSearchOps OK. searchOps is too high! 58308 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489086175] SERVICE ALERT: idmrbs.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489086282] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9481 found since last run. oneLevelSearchOps OK. searchOps is too high! 9648 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489086342] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489086596] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1489086643] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8043 found since last run. oneLevelSearchOps OK. searchOps is too high! 8275 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489086703] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489086942] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5757 found since last run. oneLevelSearchOps OK. searchOps is too high! 5873 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489087002] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489087097] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 84% (15 min average)
[1489087303] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8282 found since last run. oneLevelSearchOps OK. searchOps is too high! 8484 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489087363] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489087602] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5024 found since last run. oneLevelSearchOps OK. searchOps is too high! 5113 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489087662] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489087677] SERVICE ALERT: owen101lab3;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1489087697] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1489087737] SERVICE ALERT: owen101lab3;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1489087757] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1489087797] SERVICE ALERT: owen101lab3;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1489087963] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9251 found since last run. oneLevelSearchOps OK. searchOps is too high! 9437 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489087984] SERVICE ALERT: sillpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1489088023] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489088063] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Ready
[1489088197] SERVICE ALERT: sillpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 34% - OK!
[1489088200] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1489088257] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1489088318] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1489088378] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1489088623] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7322 found since last run. oneLevelSearchOps OK. searchOps is too high! 7519 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489088683] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489088737] SERVICE FLAPPING ALERT: owen101lab2;Black Print Cartridge HP C8543X;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1489088957] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1489088996] SERVICE ALERT: owen101lab3;Device Status;OK;HARD;3;50 72 65 73 73 20 17 20 74 6F 20 73 65 65 20 73  74 65 70  by step information
[1489088997] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1489089017] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1489089057] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1489089077] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1489089117] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1489089283] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8706 found since last run. oneLevelSearchOps OK. searchOps is too high! 8859 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489089343] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489089717] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1489089775] SERVICE ALERT: idmrbs.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8254 found since last run. oneLevelSearchOps OK. searchOps is too high! 8257 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489089778] SERVICE FLAPPING ALERT: owen101lab2;Device Status;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1489089835] SERVICE ALERT: idmrbs.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489089943] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 11807 found since last run. oneLevelSearchOps OK. searchOps is too high! 12011 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489090003] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489090218] SERVICE ALERT: strongpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1489090277] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1489090278] SERVICE ALERT: strongpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1489090338] SERVICE ALERT: strongpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1489090435] SERVICE ALERT: idmrbs.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8561 found since last run. oneLevelSearchOps OK. searchOps is too high! 8562 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489090495] SERVICE ALERT: idmrbs.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489090582] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 330 found since last run. oneLevelSearchOps OK. searchOps is too high! 331 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489090603] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8593 found since last run. oneLevelSearchOps OK. searchOps is too high! 8864 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489090642] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489090662] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8070 found since last run. oneLevelSearchOps OK. searchOps is too high! 8302 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489090663] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489090712] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1489090722] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489090772] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1489090831] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1489090939] SERVICE ALERT: strongpk;Device Status;OK;HARD;3;Open front door.
[1489090939] SERVICE FLAPPING ALERT: strongpk;Device Status;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1489091095] SERVICE ALERT: idmrbs.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8231 found since last run. oneLevelSearchOps OK. searchOps is too high! 8232 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489091155] SERVICE ALERT: idmrbs.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489091263] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9582 found since last run. oneLevelSearchOps OK. searchOps is too high! 9767 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489091322] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10301 found since last run. oneLevelSearchOps OK. searchOps is too high! 10327 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489091323] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489091382] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489091477] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1489091537] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1489091597] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 88% (15 min average)
[1489091755] SERVICE ALERT: idmrbs.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 16577 found since last run. oneLevelSearchOps OK. searchOps is too high! 16580 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489091815] SERVICE ALERT: idmrbs.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489091923] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5903 found since last run. oneLevelSearchOps OK. searchOps is too high! 6136 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489091982] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9258 found since last run. oneLevelSearchOps OK. searchOps is too high! 9437 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489091983] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489092042] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489092197] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1489092257] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 68% (15 min average)
[1489092415] SERVICE ALERT: idmrbs.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8231 found since last run. oneLevelSearchOps OK. searchOps is too high! 8236 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489092475] SERVICE ALERT: idmrbs.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489092638] SERVICE ALERT: owen101lab2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1489092642] HOST ALERT: owen101lab2;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1489092642] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7497 found since last run. oneLevelSearchOps OK. searchOps is too high! 7771 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489092702] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489092773] SERVICE ALERT: owen101lab2;Device Status;OK;HARD;3;Warming up
[1489092931] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at WARNING level!
[1489093075] SERVICE ALERT: idmrbs.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 16485 found since last run. oneLevelSearchOps OK. searchOps is too high! 16486 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489093135] SERVICE ALERT: idmrbs.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489093302] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6429 found since last run. oneLevelSearchOps OK. searchOps is too high! 6584 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489093362] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489093455] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1489093515] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1489093537] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1489093567] HOST ALERT: owen101lab2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1489093573] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 93% (15 min average)
[1489093657] HOST ALERT: owen101lab2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1489093747] HOST ALERT: owen101lab2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1489093783] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7021 found since last run. oneLevelSearchOps OK. searchOps is too high! 7193 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489093795] HOST ALERT: hallefirst1;UP;HARD;10;PING OK - Packet loss = 79%, RTA = 0.56 ms
[1489093837] HOST ALERT: owen101lab2;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1489093843] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489093843] SERVICE ALERT: owen101lab2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1489093873] HOST ALERT: owen101lab2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1489093903] SERVICE ALERT: owen101lab2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1489093933] HOST ALERT: owen101lab2;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1489093963] SERVICE ALERT: owen101lab2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1489093978] SERVICE ALERT: owen101lab2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1489093993] HOST ALERT: owen101lab2;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1489094038] SERVICE ALERT: owen101lab2;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1489094068] HOST ALERT: owen101lab2;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1489094098] SERVICE ALERT: owen101lab2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1489094128] HOST ALERT: owen101lab2;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1489094167] HOST ALERT: owen101lab2;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1489094173] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;Ready
[1489094259] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 9% - OK!
[1489094366] SERVICE ALERT: hallefirst1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1489094562] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7284 found since last run. oneLevelSearchOps OK. searchOps is too high! 7290 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489094622] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489095373] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1489095433] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1489095493] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1489095643] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6229 found since last run. oneLevelSearchOps OK. searchOps is too high! 6340 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489095703] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489095823] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7526 found since last run. oneLevelSearchOps OK. searchOps is too high! 7618 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489095882] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489096093] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 90% (15 min average)
[1489096303] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7691 found since last run. oneLevelSearchOps OK. searchOps is too high! 7805 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489096363] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489096481] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5742 found since last run. oneLevelSearchOps OK. searchOps is too high! 5758 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489096540] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489096693] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1489096753] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1489096813] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1489096963] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7627 found since last run. oneLevelSearchOps OK. searchOps is too high! 7699 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489097023] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489097377] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1489097623] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7587 found since last run. oneLevelSearchOps OK. searchOps is too high! 7662 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489097683] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489098283] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8137 found since last run. oneLevelSearchOps OK. searchOps is too high! 8205 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489098340] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6738 found since last run. oneLevelSearchOps OK. searchOps is too high! 7020 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489098343] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489098400] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489098578] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1489098613] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 87% (15 min average)
[1489098637] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1489098698] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1489098943] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6679 found since last run. oneLevelSearchOps OK. searchOps is too high! 6766 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489099003] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489099813] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1489099873] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1489099933] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1489100203] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5078 found since last run. oneLevelSearchOps OK. searchOps is too high! 5146 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489100263] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489100394] SERVICE ALERT: owen101lab3;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1489100453] SERVICE ALERT: owen101lab3;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1489100513] SERVICE ALERT: owen101lab3;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1489100533] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 79% (15 min average)
[1489100863] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5827 found since last run. oneLevelSearchOps OK. searchOps is too high! 5902 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489100923] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489101115] SERVICE FLAPPING ALERT: markjeffersonpk;Paper Tray Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1489101133] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1489101193] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1489101253] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 77% (15 min average)
[1489102000] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5987 found since last run. oneLevelSearchOps OK. searchOps is too high! 6128 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489102060] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489102313] SERVICE ALERT: owen101lab3;Device Status;OK;HARD;3;SIZE MISMATCH IN TRAY 4 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Processing job from tray 3
[1489102338] SERVICE FLAPPING ALERT: strongpk;Device Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1489102453] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1489102513] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1489102573] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1489104460] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5038 found since last run. oneLevelSearchOps OK. searchOps is too high! 5168 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1489104520] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1489104973] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 39% (15 min average)
[1489107631] SERVICE ALERT: eateriespk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1489107698] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1489108573] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1489108633] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1489108693] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 88% (15 min average)
[1489111093] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1489111153] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1489112142] SERVICE ALERT: EVENTSTEST-2012;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 39.66 Gb - used: 35.69 Gb (90%) - free 3.96 Gb (10%)
[1489112202] SERVICE ALERT: EVENTSTEST-2012;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 39.66 Gb - used: 35.69 Gb (90%) - free 3.96 Gb (10%)
[1489112261] SERVICE ALERT: EVENTSTEST-2012;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 39.66 Gb - used: 35.69 Gb (90%) - free 3.96 Gb (10%)
[1489112953] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1489112995] SERVICE ALERT: eateriespk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1489113013] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 83% (15 min average)
[1489113055] SERVICE ALERT: eateriespk;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 34% - OK!
[1489113461] SERVICE ALERT: EVENTSTEST-2012;C:\ Drive Space;OK;HARD;3;c:\ - total: 39.66 Gb - used: 35.08 Gb (88%) - free 4.58 Gb (12%)
[1489113613] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1489113673] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1489113733] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 94% (15 min average)
[1489114000] HOST ALERT: IMAGES;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1489114090] HOST ALERT: IMAGES;DOWN;HARD;2;PING CRITICAL - Packet loss = 100%
[1489114107] HOST ALERT: IMAGES;UP;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1489114107] SERVICE ALERT: IMAGES;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1489114147] SERVICE ALERT: IMAGES;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1489114933] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1489116733] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1489116793] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 91% (15 min average)
[1489117393] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1489117453] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1489117513] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1489118713] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1489119313] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1489120208] SERVICE ALERT: hallesecondsouth;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1489120268] SERVICE ALERT: hallesecondsouth;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1489120328] SERVICE ALERT: hallesecondsouth;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1489121527] SERVICE ALERT: hallesecondsouth;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full