[1478927275] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478927395] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1478927515] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 68% (15 min average)
[1478928115] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1478928235] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 72% (15 min average)
[1478928835] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1478928955] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 67% (15 min average)
[1478929555] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1478929675] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 38% (15 min average)
[1478930275] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1478930395] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1478930515] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 90% (15 min average)
[1478931115] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478931235] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1478931355] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1478931534] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1478931654] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1478932201] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1478932555] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 87% (15 min average)
[1478934053] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1478934173] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;2;Warning:  idm1  are more than 120 seconds behind!
[1478934293] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1478934955] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478935075] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1478935195] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1478935493] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1478936693] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1478937293] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1478937293] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1478937595] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1478938195] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1478938289] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1478938323] SERVICE ALERT: IT-2012;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1478938409] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1478938443] SERVICE ALERT: IT-2012;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1478938529] SERVICE ALERT: epprogsql2;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1478938563] SERVICE ALERT: IT-2012;CPU Load;OK;SOFT;3;CPU Load 4% (15 min average)
[1478938795] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1478938915] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 84% (15 min average)
[1478939093] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1478939129] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1478939515] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1478939635] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1478939693] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1478939755] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 93% (15 min average)
[1478940293] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1478940355] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1478940413] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;2;Warning:  idm1  are more than 120 seconds behind!
[1478940475] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1478940533] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1478940595] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1478940929] SERVICE ALERT: epprogsql2;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1478941529] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1478941529] SERVICE FLAPPING ALERT: epprogsql2;CPU Load;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1478941795] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 90% (15 min average)
[1478942129] SERVICE ALERT: epprogsql2;CPU Load;OK;HARD;3;CPU Load 67% (15 min average)
[1478942333] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1478943003] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (34.9% change >= 20.0% threshold)
[1478943003] SERVICE FLAPPING ALERT: epprogsql2;CPU Load;STARTED; Service appears to have started flapping (28.8% change >= 20.0% threshold)
[1478943003] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (42.0% change >= 20.0% threshold)
[1478943003] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1478943003] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1478943003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1478943003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1478943003] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1478943003] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1478943003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1478943003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1478943003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1478943003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1478943003] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1478943003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1478943003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1478943003] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1478943003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1478943003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1478943003] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1478944802] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (28.7% change >= 20.0% threshold)
[1478944802] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (38.8% change >= 20.0% threshold)
[1478944802] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1478944802] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1478944802] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1478944802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1478944802] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1478944802] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1478944802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1478944802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1478944802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1478944802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1478944802] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1478944802] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1478944802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1478944802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1478944802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1478944802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1478944802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1478948995] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478949115] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1478949235] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1478950201] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service has exited from a period of scheduled downtime
[1478952235] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1478952533] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1478954035] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1478954155] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1478954275] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1478955475] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1478955521] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 128 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478955641] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478961002] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1478961002] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1478961002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1478961002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1478961002] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1478961002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1478961002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1478961002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1478961002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1478961002] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1478961002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1478961002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1478961002] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1478961002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1478961002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1478961002] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1478961007] SERVICE ALERT: ADFS;Uptime;OK;HARD;3;System Uptime - 7 day(s) 23 hour(s) 38 minute(s)
[1478961007] SERVICE ALERT: ADFS;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1478961011] HOST ALERT: ADFS;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1478961170] SERVICE ALERT: ADFS;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1478961170] SERVICE ALERT: ADFS;C:\ Drive Space;OK;HARD;3;c:\ - total: 79.51 Gb - used: 18.70 Gb (24%) - free 60.81 Gb (76%)
[1478961257] SERVICE ALERT: ADFS;Service: RPC;OK;HARD;3;RpcSs: Started
[1478961257] SERVICE ALERT: ADFS;CPU Load;OK;HARD;3;CPU Load 0% (15 min average)
[1478961347] SERVICE ALERT: ADFS;Memory Usage;OK;HARD;3;Memory usage: total:10239.64 MB - used: 2218.55 MB (22%) - free: 8021.09 MB (78%)
[1478961348] SERVICE ALERT: ADFS;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1478961429] SERVICE ALERT: ADFS;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1478961432] SERVICE ALERT: ADFS;Service: WMI;OK;HARD;3;Winmgmt: Started
[1478961505] SERVICE ALERT: ADFS;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1478961521] SERVICE ALERT: ADFS;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1478961522] SERVICE ALERT: ADFS;Service: Windows Time;OK;HARD;3;W32Time: Started
[1478961522] SERVICE ALERT: ADFS;VMWare Tools;OK;HARD;3;VMTools: Started
[1478964779] SERVICE ALERT: DB01-2012;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478964899] SERVICE ALERT: DB01-2012;CPU Load;OK;SOFT;2;CPU Load 0% (15 min average)
[1478965499] SERVICE ALERT: DB01-2012;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478965619] SERVICE ALERT: DB01-2012;CPU Load;OK;SOFT;2;CPU Load 2% (15 min average)
[1478970174] SERVICE ALERT: CAVA2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1478970294] SERVICE ALERT: CAVA2;CPU Load;OK;SOFT;2;CPU Load 7% (15 min average)
[1478970475] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478970595] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1478970715] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 90% (15 min average)
[1478970894] SERVICE ALERT: CAVA2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1478971014] SERVICE ALERT: CAVA2;CPU Load;OK;SOFT;2;CPU Load 1% (15 min average)
[1478971803] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1478971803] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1478971803] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1478971803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1478971803] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1478971803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1478971803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1478971803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1478971803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1478971803] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1478971803] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1478971803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1478971803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1478971803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1478971803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1478971803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1478971810] SERVICE ALERT: ADFS;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478971810] SERVICE ALERT: ADFS;Uptime;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478971813] HOST ALERT: ADFS;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.57)
[1478971876] HOST ALERT: ADFS;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.4.57)
[1478971930] SERVICE ALERT: ADFS;Uptime;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478971930] SERVICE ALERT: ADFS;Service: Firewall;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478971933] HOST ALERT: ADFS;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.4.57)
[1478971973] SERVICE ALERT: ADFS;Service: NetLogin;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478971973] SERVICE ALERT: ADFS;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478971976] HOST ALERT: ADFS;DOWN;SOFT;4;CRITICAL - Host Unreachable (164.76.4.57)
[1478972039] HOST ALERT: ADFS;DOWN;SOFT;5;CRITICAL - Host Unreachable (164.76.4.57)
[1478972050] SERVICE ALERT: ADFS;Service: Firewall;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478972050] SERVICE ALERT: ADFS;Uptime;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478972060] SERVICE ALERT: ADFS;CPU Load;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478972060] SERVICE ALERT: ADFS;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478972063] HOST ALERT: ADFS;DOWN;SOFT;6;CRITICAL - Host Unreachable (164.76.4.57)
[1478972093] SERVICE ALERT: ADFS;C:\ Drive Space;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478972093] SERVICE ALERT: ADFS;Service: NetLogin;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478972096] HOST ALERT: ADFS;DOWN;SOFT;7;CRITICAL - Host Unreachable (164.76.4.57)
[1478972150] SERVICE ALERT: ADFS;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478972150] SERVICE ALERT: ADFS;Service: User Profile;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478972153] HOST ALERT: ADFS;DOWN;SOFT;8;CRITICAL - Host Unreachable (164.76.4.57)
[1478972180] SERVICE ALERT: ADFS;Service: RPC;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478972180] SERVICE ALERT: ADFS;CPU Load;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478972183] HOST ALERT: ADFS;DOWN;SOFT;9;CRITICAL - Host Unreachable (164.76.4.57)
[1478972213] SERVICE ALERT: ADFS;Service: NetLogin;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478972213] SERVICE ALERT: ADFS;C:\ Drive Space;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478972216] HOST ALERT: ADFS;DOWN;HARD;10;CRITICAL - Host Unreachable (164.76.4.57)
[1478972232] SERVICE ALERT: ADFS;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478972238] SERVICE ALERT: ADFS;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478972270] SERVICE ALERT: ADFS;Memory Usage;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478972270] SERVICE ALERT: ADFS;Service: User Profile;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478972300] SERVICE ALERT: ADFS;Service: RPC;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478972308] SERVICE ALERT: ADFS;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478972308] SERVICE ALERT: ADFS;CPU Load;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478972324] SERVICE ALERT: ADFS;VMWare Tools;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478972324] SERVICE ALERT: ADFS;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478972324] SERVICE ALERT: ADFS;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478972352] SERVICE ALERT: ADFS;NSClient++ Version;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478972358] SERVICE ALERT: ADFS;Service: WMI;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478972390] SERVICE ALERT: ADFS;Memory Usage;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478972390] SERVICE ALERT: ADFS;Service: User Profile;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478972428] SERVICE ALERT: ADFS;Service: Group Policy;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478972444] SERVICE ALERT: ADFS;Service: Windows Time;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478972444] SERVICE ALERT: ADFS;Service: DNS Client;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478972444] SERVICE ALERT: ADFS;VMWare Tools;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478972472] SERVICE ALERT: ADFS;NSClient++ Version;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478972478] SERVICE ALERT: ADFS;Service: WMI;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478972515] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1478972548] SERVICE ALERT: ADFS;Service: Group Policy;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478972564] SERVICE ALERT: ADFS;VMWare Tools;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478972564] SERVICE ALERT: ADFS;Service: DNS Client;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478972564] SERVICE ALERT: ADFS;Service: Windows Time;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478972635] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 25% (15 min average)
[1478973235] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1478973355] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1478973475] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 90% (15 min average)
[1478978275] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1478978395] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 82% (15 min average)
[1478978995] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1478979115] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1478979235] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 75% (15 min average)
[1478982094] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1478982153] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1478982213] SERVICE ALERT: crossroadspk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1478982813] SERVICE ALERT: crossroadspk;Device Status;OK;HARD;3;Open tray indicated.
[1478984013] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1478984073] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1478984133] SERVICE ALERT: crossroadspk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1478984734] SERVICE ALERT: crossroadspk;Device Status;OK;HARD;3;Open tray indicated.
[1478984734] SERVICE FLAPPING ALERT: crossroadspk;Device Status;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1478985333] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1478985393] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1478985453] SERVICE ALERT: crossroadspk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1478985834] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1478985954] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1478986053] SERVICE ALERT: crossroadspk;Device Status;OK;HARD;3;Open tray indicated.
[1478986074] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 61% (15 min average)
[1478987253] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1478987313] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1478987373] SERVICE ALERT: crossroadspk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1478987973] SERVICE ALERT: crossroadspk;Device Status;OK;HARD;3;Open tray indicated.
[1478989173] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1478989233] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1478989293] SERVICE ALERT: crossroadspk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1478989893] SERVICE ALERT: crossroadspk;Device Status;OK;HARD;3;Open tray indicated.
[1478990840] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1478990874] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478990891] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 124 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478990900] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1478990960] SERVICE ALERT: hallefirst2;Device Status;OK;SOFT;3;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1478990994] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1478991011] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478991093] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1478991153] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1478991213] SERVICE ALERT: crossroadspk;Device Status;OK;SOFT;3;Open tray indicated.
[1478991594] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1478991714] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1478991776] SERVICE ALERT: cava2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1478991834] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 80% (15 min average)
[1478991896] SERVICE ALERT: cava2;CPU Load;OK;SOFT;2;CPU Load 80% (15 min average)
[1478992413] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1478992473] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1478992496] SERVICE ALERT: cava2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1478992533] SERVICE ALERT: crossroadspk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1478992616] SERVICE ALERT: cava2;CPU Load;OK;SOFT;2;CPU Load 0% (15 min average)
[1478993134] SERVICE ALERT: crossroadspk;Device Status;OK;HARD;3;Open tray indicated.
[1478993634] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1478993754] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1478994333] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1478994354] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1478994393] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1478994453] SERVICE ALERT: crossroadspk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1478994474] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1478994594] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1478995053] SERVICE ALERT: crossroadspk;Device Status;OK;HARD;3;Open tray indicated.
[1478995194] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 70% (15 min average)
[1478996253] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1478996314] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1478996373] SERVICE ALERT: crossroadspk;Device Status;OK;SOFT;3;Open tray indicated.
[1478996394] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1478996514] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1478997114] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1478997234] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1478997354] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1478997573] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1478997633] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1478997693] SERVICE ALERT: crossroadspk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1478997954] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 80% (15 min average)
[1478997954] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1478998293] SERVICE ALERT: crossroadspk;Device Status;OK;HARD;3;Open tray indicated.
[1478999493] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1478999553] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1478999613] SERVICE ALERT: crossroadspk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1479000213] SERVICE ALERT: crossroadspk;Device Status;OK;HARD;3;Open tray indicated.
[1479000354] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1479000474] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 89% (15 min average)
[1479001413] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1479001473] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1479001533] SERVICE ALERT: crossroadspk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1479002134] SERVICE ALERT: crossroadspk;Device Status;OK;HARD;3;Open tray indicated.
[1479002733] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1479002793] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1479002853] SERVICE ALERT: crossroadspk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1479002874] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1479002994] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1479003114] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1479003453] SERVICE ALERT: crossroadspk;Device Status;OK;HARD;3;Open tray indicated.
[1479004314] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 52% (15 min average)
[1479004653] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1479004713] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1479004774] SERVICE ALERT: crossroadspk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1479004914] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1479005034] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1479005154] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1479005373] SERVICE ALERT: crossroadspk;Device Status;OK;HARD;3;Open tray indicated.
[1479005754] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1479006354] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1479006573] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1479006633] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1479006693] SERVICE ALERT: crossroadspk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1479007293] SERVICE ALERT: crossroadspk;Device Status;OK;HARD;3;Open tray indicated.
[1479008494] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1479008553] SERVICE ALERT: crossroadspk;Device Status;OK;SOFT;2;Open tray indicated.
[1479009611] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 100 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479009731] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479009753] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1479009813] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1479009873] SERVICE ALERT: crossroadspk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1479011073] SERVICE ALERT: crossroadspk;Device Status;OK;HARD;3;Open tray indicated.
[1479011675] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1479011733] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1479011793] SERVICE ALERT: crossroadspk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1479012393] SERVICE ALERT: crossroadspk;Device Status;OK;HARD;3;Open tray indicated.