[1492315608] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1492315668] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1492315728] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1492316328] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 44% (15 min average)
[1492317001] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1492317528] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1492317588] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 89% (15 min average)
[1492318302] SERVICE ALERT: ADWU1;Service: Group Policy;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1492318302] SERVICE ALERT: ADWU1;Service: Firewall;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1492318327] SERVICE ALERT: ADWU1;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1492318361] SERVICE ALERT: ADWU1;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1492318362] SERVICE ALERT: ADWU1;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1492318367] SERVICE ALERT: ADWU1;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1492318556] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;1;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1492318676] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;2;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1492318788] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1492318796] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;HARD;3;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1492318848] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1492318908] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1492319508] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1492319687] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1492319747] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1492320349] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1492320406] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1492320466] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1492321666] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1492321908] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1492322508] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1492322568] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 90% (15 min average)
[1492322866] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1492323168] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1492323228] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1492323288] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1492325267] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1492325267] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1492325433] SERVICE ALERT: ADDC5;AD Group Changes;UNKNOWN;SOFT;1;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1492325688] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1492325695] SERVICE ALERT: ADDC5;AD Group Changes;OK;SOFT;2;WARNING:  0  groups have changed!
[1492325866] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1492325926] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;2;Warning:  idm1  are more than 120 seconds behind!
[1492325986] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1492325996] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;OK;HARD;3;OK: No entries found
[1492326288] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1492326348] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1492326408] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1492327008] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 91% (15 min average)
[1492327186] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1492327608] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1492327668] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1492328180] SERVICE ALERT: MBAM;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1492328240] SERVICE ALERT: MBAM;CPU Load;OK;SOFT;2;CPU Load 6% (15 min average)
[1492328868] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1492328928] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1492328986] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1492328988] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1492329586] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1492329646] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1492329706] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1492330306] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1492330788] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 70% (15 min average)
[1492330906] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1492330966] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1492331400] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service has exited from a period of scheduled downtime
[1492331402] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (40.0% change >= 20.0% threshold)
[1492331402] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (46.5% change >= 20.0% threshold)
[1492331402] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1492331402] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1492331402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1492331402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1492331402] HOST DOWNTIME ALERT: INTLAPP;STARTED; Host has entered a period of scheduled downtime
[1492331402] HOST DOWNTIME ALERT: INTLDB;STARTED; Host has entered a period of scheduled downtime
[1492331402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1492331402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1492331402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1492331402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1492331402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1492331402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1492331402] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1492331402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1492331402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1492331402] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1492331402] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1492333203] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (33.5% change >= 20.0% threshold)
[1492333203] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (39.3% change >= 20.0% threshold)
[1492333203] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1492333203] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1492333203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1492333203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1492333203] HOST DOWNTIME ALERT: INTLAPP;STARTED; Host has entered a period of scheduled downtime
[1492333203] HOST DOWNTIME ALERT: INTLDB;STARTED; Host has entered a period of scheduled downtime
[1492333203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1492333203] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1492333203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1492333203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1492333203] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1492333203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1492333203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1492333203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1492333203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1492333203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1492333203] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1492336965] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1492337025] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1492337133] SERVICE ALERT: ADDC5;AD Group Changes;UNKNOWN;SOFT;1;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1492337388] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1492337405] SERVICE ALERT: ADDC5;AD Group Changes;OK;SOFT;2;WARNING:  0  groups have changed!
[1492337448] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1492337508] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1492341108] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1492341708] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1492341768] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1492341825] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1492341828] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1492343028] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 79% (15 min average)
[1492349403] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1492349403] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1492349403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1492349403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1492349403] HOST DOWNTIME ALERT: INTLAPP;STARTED; Host has entered a period of scheduled downtime
[1492349403] HOST DOWNTIME ALERT: INTLDB;STARTED; Host has entered a period of scheduled downtime
[1492349403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1492349403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1492349403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1492349403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1492349403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1492349403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1492349403] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1492349403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1492349403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1492349403] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1492349403] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1492351833] SERVICE ALERT: ADDC5;AD Group Changes;UNKNOWN;SOFT;1;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1492352099] SERVICE ALERT: ADDC5;AD Group Changes;OK;SOFT;2;WARNING:  0  groups have changed!
[1492358603] SERVICE ALERT: CAVA2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1492358663] SERVICE ALERT: CAVA2;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1492358723] SERVICE ALERT: CAVA2;CPU Load;OK;SOFT;3;CPU Load 57% (15 min average)
[1492359323] SERVICE ALERT: CAVA2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1492359383] SERVICE ALERT: CAVA2;CPU Load;OK;SOFT;2;CPU Load 5% (15 min average)
[1492360202] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1492360202] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1492360202] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1492360202] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1492360202] HOST DOWNTIME ALERT: INTLAPP;STARTED; Host has entered a period of scheduled downtime
[1492360202] HOST DOWNTIME ALERT: INTLDB;STARTED; Host has entered a period of scheduled downtime
[1492360202] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1492360202] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1492360202] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1492360202] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1492360202] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1492360202] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1492360202] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1492360202] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1492360202] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1492360202] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1492360202] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1492369201] SERVICE ALERT: ax-app2012;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.158.30 and port 12489: Connection refused
[1492369261] SERVICE ALERT: ax-app2012;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1492370300] SERVICE ALERT: hallefirst2;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1492370360] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! 1 of 3 Trays Still Full
[1492379027] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1492379087] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1492379147] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1492379747] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1492380332] SERVICE ALERT: ADDC5;AD Group Changes;UNKNOWN;SOFT;1;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1492380347] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1492380407] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1492380467] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1492380602] SERVICE ALERT: ADDC5;AD Group Changes;OK;SOFT;2;WARNING:  0  groups have changed!
[1492381066] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 88% (15 min average)
[1492381066] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1492386466] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1492386526] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 31% (15 min average)
[1492391926] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1492391986] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1492392046] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1492392463] SERVICE ALERT: TARGETX-2012;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1492392523] SERVICE ALERT: TARGETX-2012;CPU Load;OK;SOFT;2;CPU Load 70% (15 min average)
[1492392600] HOST DOWNTIME ALERT: IMAGES;STARTED; Host has entered a period of scheduled downtime
[1492393246] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 78% (15 min average)
[1492393576] SERVICE ALERT: IMAGES;MDT Monitor;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1492393606] HOST ALERT: IMAGES;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1492393696] HOST ALERT: IMAGES;DOWN;HARD;2;PING CRITICAL - Packet loss = 100%
[1492393703] HOST ALERT: IMAGES;UP;HARD;2;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1492393846] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1492393906] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1492393966] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1492394156] SERVICE ALERT: IMAGES;MDT Monitor;UNKNOWN;HARD;3;Failed to open service MDT_Monitor: 1060: The specified service does not exist as an installed service.
[1492394566] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 76% (15 min average)
[1492395166] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1492395226] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1492395286] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 75% (15 min average)
[1492396983] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1492397043] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 47% - OK!
[1492399800] HOST DOWNTIME ALERT: IMAGES;STOPPED; Host has exited from a period of scheduled downtime
[1492400086] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1492400146] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1492400206] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 84% (15 min average)
[1492400806] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1492400866] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1492400926] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)