[1460088083] SERVICE ALERT: INTLTESTAPP;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460088084] SERVICE ALERT: ADAUTHDC3;Service: Distributed File System Replication;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460088084] SERVICE ALERT: UHSPHARM-2012;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460088084] SERVICE ALERT: TI22-PSYCLINIC;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460088084] SERVICE ALERT: epprogsql2;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460088105] HOST ALERT: VUM;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460088109] SERVICE ALERT: CBORD;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460088109] SERVICE ALERT: CBORD;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460088113] HOST ALERT: INTLTESTAPP;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460088113] SERVICE ALERT: VUM;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460088116] SERVICE ALERT: INTLTESTAPP;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460088134] SERVICE ALERT: VUM;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460088139] HOST ALERT: CBORD;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460088143] HOST ALERT: VUM;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460088146] HOST ALERT: INTLTESTAPP;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460088155] SERVICE ALERT: CBORD;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460088163] SERVICE ALERT: INTLTESTAPP;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460088178] SERVICE ALERT: VUM;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460088184] HOST ALERT: CBORD;DOWN;SOFT;2;PING CRITICAL - Packet loss = 72%, RTA = 24904.88 ms
[1460088184] SERVICE ALERT: UHSPHARM-2012;CPU Load;OK;SOFT;2;CPU Load 3% (5 min average)
[1460088184] SERVICE ALERT: ADAUTHDC3;Service: Distributed File System Replication;OK;SOFT;2;DFSR: Started
[1460088184] HOST ALERT: INTLTESTAPP;DOWN;SOFT;3;PING CRITICAL - Packet loss = 63%, RTA = 18115.77 ms
[1460088184] SERVICE ALERT: TI22-PSYCLINIC;Memory Usage;OK;SOFT;2;Memory usage: total:10943.64 MB - used: 1964.82 MB (18%) - free: 8978.82 MB (82%)
[1460088184] SERVICE ALERT: INTLTESTAPP;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1460088184] SERVICE ALERT: epprogsql2;Service: RPC;OK;SOFT;2;RpcSs: Started
[1460088188] HOST ALERT: INTLTESTAPP;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1460088190] HOST ALERT: VUM;DOWN;SOFT;3;PING CRITICAL - Packet loss = 50%, RTA = 9086.21 ms
[1460088190] HOST ALERT: CBORD;UP;SOFT;3;PING OK - Packet loss = 14%, RTA = 2778.74 ms
[1460088209] SERVICE ALERT: CBORD;C:\ Drive Space;OK;SOFT;2;c:\ - total: 59.66 Gb - used: 36.74 Gb (62%) - free 22.92 Gb (38%)
[1460088209] SERVICE ALERT: CBORD;VMWare Tools;OK;SOFT;2;VMTools: Started
[1460088213] SERVICE ALERT: VUM;Uptime;OK;SOFT;2;System Uptime - 2 day(s) 1 hour(s) 29 minute(s)
[1460088216] SERVICE ALERT: INTLTESTAPP;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1460088217] HOST ALERT: VUM;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.41 ms
[1460088234] SERVICE ALERT: VUM;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1460088255] SERVICE ALERT: CBORD;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1460088263] SERVICE ALERT: INTLTESTAPP;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1460088278] SERVICE ALERT: VUM;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1460089090] SERVICE ALERT: halleground1;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Maintenance Kit HP 110V-C9152A, 220V-C9153A'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1460089150] SERVICE ALERT: halleground1;Maintenance Kit HP 110V-C9152A;OK;SOFT;2;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 30% - OK!
[1460093356] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 92% (5 min average)
[1460093956] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1460095156] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 87% (5 min average)
[1460095216] SERVICE ALERT: TI22-PSYCLINIC;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460095336] SERVICE ALERT: TI22-PSYCLINIC;CPU Load;OK;SOFT;2;CPU Load 2% (5 min average)
[1460095503] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1460095623] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 88% (5 min average)
[1460095756] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460095876] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1460095996] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1460096223] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1460096343] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 90% (5 min average)
[1460096463] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1460097063] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460098263] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 92% (5 min average)
[1460098863] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1460100796] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 75% (5 min average)
[1460101263] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 93% (5 min average)
[1460101396] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460101516] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1460101636] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 80% (5 min average)
[1460102236] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460102356] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1460102463] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 89% (5 min average)
[1460102476] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1460103076] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 59% (5 min average)
[1460103676] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 91% (5 min average)
[1460103796] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 28% (5 min average)
[1460105596] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1460105716] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460105836] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1460106436] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 80% (5 min average)
[1460107263] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1460107383] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 75% (5 min average)
[1460107983] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1460108103] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 72% (5 min average)
[1460110034] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460110154] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1460110274] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460112674] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 93% (5 min average)
[1460113274] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1460113503] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1460113874] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 54% (5 min average)
[1460115066] SERVICE ALERT: IMAGES;Service: MDT Monitor;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460115096] HOST ALERT: IMAGES;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460115114] SERVICE ALERT: IMAGES;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460115144] HOST ALERT: IMAGES;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460115158] SERVICE ALERT: IMAGES;Service: PXE Boot;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460115186] SERVICE ALERT: IMAGES;Service: MDT Monitor;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460115188] HOST ALERT: IMAGES;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460115209] SERVICE ALERT: IMAGES;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460115214] SERVICE ALERT: IMAGES;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1460115215] HOST ALERT: IMAGES;UP;SOFT;4;PING OK - Packet loss = 28%, RTA = 0.25 ms
[1460115258] SERVICE ALERT: IMAGES;Service: PXE Boot;UNKNOWN;SOFT;2;Failed to open service manager: 1115: A system shutdown is in progress.
[1460115286] SERVICE ALERT: IMAGES;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460115293] SERVICE ALERT: IMAGES;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460115293] SERVICE ALERT: IMAGES;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460115306] SERVICE ALERT: IMAGES;Service: MDT Monitor;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460115316] HOST ALERT: IMAGES;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460115329] SERVICE ALERT: IMAGES;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460115337] SERVICE ALERT: IMAGES;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460115339] SERVICE ALERT: IMAGES;MDT Monitor;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460115359] HOST ALERT: IMAGES;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460115383] SERVICE ALERT: IMAGES;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460115383] SERVICE ALERT: IMAGES;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460115398] SERVICE ALERT: IMAGES;Service: PXE Boot;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460115406] SERVICE ALERT: IMAGES;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460115413] SERVICE ALERT: IMAGES;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460115413] HOST ALERT: IMAGES;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460115413] SERVICE ALERT: IMAGES;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460115433] SERVICE ALERT: IMAGES;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460115443] HOST ALERT: IMAGES;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460115449] SERVICE ALERT: IMAGES;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460115457] SERVICE ALERT: IMAGES;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460115459] SERVICE ALERT: IMAGES;MDT Monitor;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460115479] HOST ALERT: IMAGES;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460115480] SERVICE ALERT: IMAGES;PXE Boot Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460115496] SERVICE ALERT: IMAGES;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460115503] SERVICE ALERT: IMAGES;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460115503] SERVICE ALERT: IMAGES;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460115510] HOST ALERT: IMAGES;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1460115526] SERVICE ALERT: IMAGES;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460115527] SERVICE ALERT: IMAGES;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460115533] SERVICE ALERT: IMAGES;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460115533] SERVICE ALERT: IMAGES;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460115553] SERVICE ALERT: IMAGES;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460115556] HOST ALERT: IMAGES;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1460115573] SERVICE ALERT: IMAGES;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460115577] SERVICE ALERT: IMAGES;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460115579] SERVICE ALERT: IMAGES;MDT Monitor;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460115583] SERVICE ALERT: IMAGES;PXE Boot Service;CRITICAL;SOFT;2;WDSServer: Starting
[1460115584] HOST ALERT: IMAGES;UP;SOFT;8;PING OK - Packet loss = 58%, RTA = 0.41 ms
[1460115596] SERVICE ALERT: IMAGES;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1460115603] SERVICE ALERT: IMAGES;Uptime;OK;SOFT;3;System Uptime - 0 day(s) 0 hour(s) 1 minute(s)
[1460115603] SERVICE ALERT: IMAGES;Memory Usage;OK;SOFT;3;Memory usage: total:37074.86 MB - used: 2384.92 MB (6%) - free: 34689.94 MB (94%)
[1460115627] SERVICE ALERT: IMAGES;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1460115653] SERVICE ALERT: IMAGES;NSClient++ Version;OK;SOFT;3;NSClient++ 0.4.4.15 2015-11-25
[1460115674] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460115693] SERVICE ALERT: IMAGES;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460115720] SERVICE ALERT: IMAGES;PXE Boot Service;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460115723] HOST ALERT: IMAGES;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460115794] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 39% (5 min average)
[1460115813] SERVICE ALERT: IMAGES;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460115813] HOST ALERT: IMAGES;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460115834] SERVICE ALERT: IMAGES;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460115864] HOST ALERT: IMAGES;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460115886] SERVICE ALERT: IMAGES;Service: MDT Monitor;OK;HARD;3;MDT_Monitor: Started
[1460115890] HOST ALERT: IMAGES;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1460115934] SERVICE ALERT: IMAGES;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1460116028] HOST ALERT: IMAGES;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1460116028] HOST FLAPPING ALERT: IMAGES;STARTED; Host appears to have started flapping (21.1% change > 20.0% threshold)
[1460116078] HOST ALERT: IMAGES;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460116151] HOST ALERT: IMAGES;UP;SOFT;3;PING WARNING - Packet loss = 80%, RTA = 0.29 ms
[1460116159] SERVICE ALERT: IMAGES;MDT Monitor;OK;HARD;3;MDT_Monitor: Started
[1460116251] HOST DOWNTIME ALERT: IMAGES;STARTED; Host has entered a period of scheduled downtime
[1460116393] SERVICE ALERT: IMAGES;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1460116394] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460116514] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 95% (5 min average)
[1460116534] HOST FLAPPING ALERT: IMAGES;STOPPED; Host appears to have stopped flapping (4.4% change < 5.0% threshold)
[1460116628] SERVICE ALERT: IMAGES;Service: RPC;OK;HARD;3;RpcSs: Started
[1460116634] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460116644] SERVICE ALERT: studentcenterpk;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Maintenance Kit HP 110V-C9152A, 220V-C9153A'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1460116704] SERVICE ALERT: studentcenterpk;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Maintenance Kit HP 110V-C9152A, 220V-C9153A'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1460116706] SERVICE ALERT: IMAGES;C:\ Drive Space;OK;HARD;3;c:\ - total: 1157.56 Gb - used: 537.55 Gb (46%) - free 620.02 Gb (54%)
[1460116713] SERVICE ALERT: IMAGES;CPU Load;OK;HARD;3;CPU Load 0% (5 min average)
[1460116713] SERVICE ALERT: IMAGES;Service: WMI;OK;HARD;3;Winmgmt: Started
[1460116757] SERVICE ALERT: IMAGES;Service: Windows Time;OK;HARD;3;W32Time: Started
[1460116764] SERVICE ALERT: studentcenterpk;Maintenance Kit HP 110V-C9152A;OK;SOFT;3;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 88% - OK!
[1460116796] SERVICE ALERT: IMAGES;Service: User Profile;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1460116798] SERVICE ALERT: IMAGES;Service: Group Policy;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1460116823] SERVICE ALERT: IMAGES;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460116823] SERVICE ALERT: IMAGES;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460116847] SERVICE ALERT: IMAGES;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460116873] SERVICE ALERT: IMAGES;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460116877] HOST ALERT: IMAGES;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460116936] SERVICE ALERT: IMAGES;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460116938] SERVICE ALERT: IMAGES;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460116943] SERVICE ALERT: IMAGES;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460116943] SERVICE ALERT: IMAGES;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460116950] HOST ALERT: IMAGES;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460116967] SERVICE ALERT: IMAGES;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460116969] SERVICE ALERT: marshallpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1460116993] SERVICE ALERT: IMAGES;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460116997] HOST ALERT: IMAGES;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460117013] SERVICE ALERT: IMAGES;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460117043] HOST ALERT: IMAGES;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460117056] SERVICE ALERT: IMAGES;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460117058] SERVICE ALERT: IMAGES;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460117063] SERVICE ALERT: IMAGES;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460117063] SERVICE ALERT: IMAGES;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460117086] HOST ALERT: IMAGES;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460117087] SERVICE ALERT: IMAGES;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460117106] SERVICE ALERT: IMAGES;Service: MDT Monitor;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460117113] SERVICE ALERT: IMAGES;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460117117] HOST ALERT: IMAGES;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1460117133] SERVICE ALERT: IMAGES;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460117154] SERVICE ALERT: IMAGES;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460117163] HOST ALERT: IMAGES;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1460117226] SERVICE ALERT: IMAGES;Service: MDT Monitor;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460117228] HOST ALERT: IMAGES;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1460117248] SERVICE ALERT: IMAGES;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460117253] SERVICE ALERT: IMAGES;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460117274] SERVICE ALERT: IMAGES;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460117278] HOST ALERT: IMAGES;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1460117326] SERVICE ALERT: IMAGES;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460117333] SERVICE ALERT: IMAGES;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460117333] SERVICE ALERT: IMAGES;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460117346] SERVICE ALERT: IMAGES;Service: MDT Monitor;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460117356] HOST ALERT: IMAGES;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1460117368] SERVICE ALERT: IMAGES;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460117377] SERVICE ALERT: IMAGES;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460117379] SERVICE ALERT: IMAGES;MDT Monitor;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460117394] SERVICE ALERT: IMAGES;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460117446] SERVICE ALERT: IMAGES;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460117453] SERVICE ALERT: IMAGES;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460117453] SERVICE ALERT: IMAGES;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460117488] SERVICE ALERT: IMAGES;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460117497] SERVICE ALERT: IMAGES;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460117499] SERVICE ALERT: IMAGES;MDT Monitor;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460117559] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1460117566] SERVICE ALERT: IMAGES;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460117573] SERVICE ALERT: IMAGES;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460117573] SERVICE ALERT: IMAGES;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460117617] SERVICE ALERT: IMAGES;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460117619] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1460117619] SERVICE ALERT: IMAGES;MDT Monitor;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460117679] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1460117700] SERVICE ALERT: ADRODC-WEMU;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460117713] SERVICE ALERT: studentcenterpk;Printer Status;OK;HARD;3;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1460117723] SERVICE ALERT: ADRODC-WEMU;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460117730] HOST ALERT: ADRODC-WEMU;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460117771] SERVICE ALERT: ADRODC-WEMU;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460117792] SERVICE ALERT: ADRODC-WEMU;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460117801] HOST ALERT: ADRODC-WEMU;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460117816] SERVICE ALERT: ADRODC-WEMU;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460117820] SERVICE ALERT: ADRODC-WEMU;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460117834] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 35% (5 min average)
[1460117843] SERVICE ALERT: ADRODC-WEMU;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460117846] HOST ALERT: ADRODC-WEMU;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1460117868] SERVICE ALERT: ADRODC-WEMU;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460117891] SERVICE ALERT: ADRODC-WEMU;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460117898] HOST ALERT: ADRODC-WEMU;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460117912] SERVICE ALERT: ADRODC-WEMU;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460117912] SERVICE ALERT: ADRODC-WEMU;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460117936] SERVICE ALERT: ADRODC-WEMU;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460117940] SERVICE ALERT: ADRODC-WEMU;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460117942] HOST ALERT: ADRODC-WEMU;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460117958] SERVICE ALERT: ADRODC-WEMU;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460117963] SERVICE ALERT: ADRODC-WEMU;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460117964] SERVICE ALERT: studentcenterpk;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Maintenance Kit HP 110V-C9152A, 220V-C9153A'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1460117988] SERVICE ALERT: ADRODC-WEMU;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460117988] HOST ALERT: ADRODC-WEMU;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1460117994] SERVICE ALERT: ADRODC-WEMU;C:\ Drive Space;OK;SOFT;3;c:\ - total: 97.31 Gb - used: 11.75 Gb (12%) - free 85.57 Gb (88%)
[1460117998] HOST ALERT: ADRODC-WEMU;UP;SOFT;7;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1460118012] SERVICE ALERT: ADRODC-WEMU;Service: WMI;OK;SOFT;3;Winmgmt: Started
[1460118012] SERVICE ALERT: ADRODC-WEMU;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1460118024] SERVICE ALERT: studentcenterpk;Maintenance Kit HP 110V-C9152A;OK;SOFT;2;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 88% - OK!
[1460118036] SERVICE ALERT: ADRODC-WEMU;CPU Load;OK;SOFT;3;CPU Load 0% (5 min average)
[1460118058] SERVICE ALERT: ADRODC-WEMU;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1460118088] SERVICE ALERT: ADRODC-WEMU;Memory Usage;OK;SOFT;3;Memory usage: total:21654.72 MB - used: 1169.95 MB (5%) - free: 20484.77 MB (95%)
[1460118143] HOST ALERT: IMAGES;UP;HARD;10;PING OK - Packet loss = 79%, RTA = 0.29 ms
[1460118520] SERVICE ALERT: ADRODC-WEMU;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1460118543] SERVICE ALERT: ADRODC-WEMU;Uptime;OK;HARD;3;System Uptime - 0 day(s) 19 hour(s) 43 minute(s)
[1460118671] SERVICE ALERT: IMAGES;Service: RPC;OK;HARD;3;RpcSs: Started
[1460118671] SERVICE FLAPPING ALERT: IMAGES;Service: RPC;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1460118700] SERVICE ALERT: IMAGES;PXE Boot Service;UNKNOWN;HARD;3;Failed to open service WDSServer: 1060: The specified service does not exist as an installed service.
[1460118746] SERVICE ALERT: IMAGES;C:\ Drive Space;OK;HARD;3;c:\ - total: 1157.56 Gb - used: 543.92 Gb (47%) - free 613.65 Gb (53%)
[1460118746] SERVICE FLAPPING ALERT: IMAGES;C:\ Drive Space;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1460118753] SERVICE ALERT: IMAGES;CPU Load;OK;HARD;3;CPU Load 3% (5 min average)
[1460118753] SERVICE FLAPPING ALERT: IMAGES;CPU Load;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1460118753] SERVICE ALERT: IMAGES;Service: WMI;OK;HARD;3;Winmgmt: Started
[1460118753] SERVICE FLAPPING ALERT: IMAGES;Service: WMI;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1460118797] SERVICE ALERT: IMAGES;Service: Windows Time;OK;HARD;3;W32Time: Started
[1460118797] SERVICE FLAPPING ALERT: IMAGES;Service: Windows Time;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1460118799] SERVICE ALERT: IMAGES;MDT Monitor;UNKNOWN;HARD;3;Failed to open service MDT_Monitor: 1060: The specified service does not exist as an installed service.
[1460118799] SERVICE FLAPPING ALERT: IMAGES;MDT Monitor;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1460118836] SERVICE ALERT: IMAGES;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1460118838] SERVICE ALERT: IMAGES;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1460118843] SERVICE ALERT: IMAGES;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 12 minute(s)
[1460118843] SERVICE ALERT: IMAGES;Memory Usage;OK;HARD;3;Memory usage: total:37842.86 MB - used: 2733.86 MB (7%) - free: 35109.01 MB (93%)
[1460118867] SERVICE ALERT: IMAGES;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1460118879] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 43% - OK!
[1460118893] SERVICE ALERT: IMAGES;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1460118913] SERVICE ALERT: studentcenterpk;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460118973] SERVICE ALERT: studentcenterpk;Printer Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460118978] SERVICE ALERT: IMAGES;Service: PXE Boot;UNKNOWN;HARD;3;Failed to open service WDSServer: 1060: The specified service does not exist as an installed service.
[1460119033] SERVICE ALERT: IMAGES;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1460119033] SERVICE FLAPPING ALERT: IMAGES;Service: Firewall;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1460119033] SERVICE ALERT: studentcenterpk;Printer Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460119107] SERVICE ALERT: INTLAPP-ISSMAP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460119126] SERVICE ALERT: IMAGES;Service: MDT Monitor;UNKNOWN;HARD;3;Failed to open service MDT_Monitor: 1060: The specified service does not exist as an installed service.
[1460119126] SERVICE FLAPPING ALERT: IMAGES;Service: MDT Monitor;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1460119174] SERVICE ALERT: IMAGES;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1460119227] SERVICE ALERT: INTLAPP-ISSMAP;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1460119991] SERVICE ALERT: ADRODC-WEMU;Service: NetLogin;CRITICAL;SOFT;1;Netlogon: Stopped
[1460120111] SERVICE ALERT: ADRODC-WEMU;Service: NetLogin;CRITICAL;SOFT;2;Netlogon: Stopped
[1460120231] SERVICE ALERT: ADRODC-WEMU;Service: NetLogin;CRITICAL;HARD;3;Netlogon: Stopped
[1460120268] HOST DOWNTIME ALERT: ADRODC-WEMU;STARTED; Host has entered a period of scheduled downtime
[1460120340] SERVICE ALERT: ADRODC-WEMU;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460120363] SERVICE ALERT: ADRODC-WEMU;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460120370] HOST ALERT: ADRODC-WEMU;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460120403] SERVICE FLAPPING ALERT: IMAGES;C:\ Drive Space;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1460120403] SERVICE FLAPPING ALERT: IMAGES;CPU Load;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1460120403] SERVICE FLAPPING ALERT: IMAGES;MDT Monitor;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1460120403] SERVICE FLAPPING ALERT: IMAGES;Service: Firewall;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1460120403] SERVICE FLAPPING ALERT: IMAGES;Service: RPC;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1460120403] SERVICE FLAPPING ALERT: IMAGES;Service: WMI;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1460120403] SERVICE FLAPPING ALERT: IMAGES;Service: Windows Time;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1460120403] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (29.7% change >= 20.0% threshold)
[1460120403] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1460120403] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1460120403] HOST DOWNTIME ALERT: ADRODC-WEMU;STARTED; Host has entered a period of scheduled downtime
[1460120403] HOST DOWNTIME ALERT: IMAGES;STARTED; Host has entered a period of scheduled downtime
[1460120415] SERVICE ALERT: ADRODC-WEMU;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.241.195 and port 12489: Connection refused
[1460120421] SERVICE ALERT: ADRODC-WEMU;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.241.195 and port 12489: Connection refused
[1460120430] SERVICE ALERT: ADRODC-WEMU;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.241.195 and port 12489: Connection refused
[1460120438] HOST ALERT: ADRODC-WEMU;UP;SOFT;2;PING OK - Packet loss = 79%, RTA = 0.52 ms
[1460120449] SERVICE ALERT: ADRODC-WEMU;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1460120456] SERVICE ALERT: ADRODC-WEMU;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460120463] SERVICE ALERT: ADRODC-WEMU;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1460120532] SERVICE ALERT: ADRODC-WEMU;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1460120532] SERVICE ALERT: ADRODC-WEMU;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1460120550] SERVICE ALERT: ADRODC-WEMU;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1460120556] SERVICE ALERT: ADRODC-WEMU;CPU Load;OK;SOFT;2;CPU Load 25% (5 min average)
[1460120831] SERVICE ALERT: ADRODC-WEMU;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1460122048] SERVICE ALERT: hallefirst1;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460122108] SERVICE ALERT: hallefirst1;Printer Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460122168] SERVICE ALERT: hallefirst1;Printer Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460122202] SERVICE FLAPPING ALERT: IMAGES;C:\ Drive Space;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1460122202] SERVICE FLAPPING ALERT: IMAGES;CPU Load;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1460122202] SERVICE FLAPPING ALERT: IMAGES;MDT Monitor;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1460122202] SERVICE FLAPPING ALERT: IMAGES;Service: Firewall;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1460122202] SERVICE FLAPPING ALERT: IMAGES;Service: RPC;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1460122202] SERVICE FLAPPING ALERT: IMAGES;Service: WMI;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1460122202] SERVICE FLAPPING ALERT: IMAGES;Service: Windows Time;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1460122202] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (27.4% change >= 20.0% threshold)
[1460122202] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1460122202] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1460122202] HOST DOWNTIME ALERT: ADRODC-WEMU;STARTED; Host has entered a period of scheduled downtime
[1460122202] HOST DOWNTIME ALERT: IMAGES;STARTED; Host has entered a period of scheduled downtime
[1460122295] SERVICE ALERT: ADRODC-WEMU;Service: AD Web Service;CRITICAL;SOFT;1;ADWS: Stopped
[1460122415] SERVICE ALERT: ADRODC-WEMU;Service: AD Web Service;CRITICAL;SOFT;2;ADWS: Stopped
[1460122535] SERVICE ALERT: ADRODC-WEMU;Service: AD Web Service;CRITICAL;HARD;3;ADWS: Stopped
[1460122634] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1460122754] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 71% (5 min average)
[1460122768] SERVICE ALERT: hallefirst1;Printer Status;OK;HARD;3;ORDER CARTRIDGE LESS THAN 1600 PAGES 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1460124003] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1460124003] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1460124003] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1460124003] HOST DOWNTIME ALERT: ADRODC-WEMU;STARTED; Host has entered a period of scheduled downtime
[1460124003] HOST DOWNTIME ALERT: IMAGES;STARTED; Host has entered a period of scheduled downtime
[1460124053] SERVICE ALERT: BBTS_DB;Size: IFDATA03.DBF;WARNING;SOFT;1;Unknown: File not found
[1460124096] SERVICE ALERT: BBTS_DB;Size: IFINDX03.DBF;WARNING;SOFT;1;Unknown: File not found
[1460124143] SERVICE ALERT: BBTS_DB;Size: TS_DATA02.DBF;WARNING;SOFT;1;Unknown: File not found
[1460124169] SERVICE ALERT: BBTS_DB;Size: IFDATA03.DBF;WARNING;SOFT;2;Unknown: File not found
[1460124189] SERVICE ALERT: BBTS_DB;Size: TS_INDEX02.DBF;WARNING;SOFT;1;Unknown: File not found
[1460124216] SERVICE ALERT: BBTS_DB;Size: IFINDX03.DBF;WARNING;SOFT;2;Unknown: File not found
[1460124264] SERVICE ALERT: BBTS_DB;Size: TS_DATA02.DBF;WARNING;SOFT;2;Unknown: File not found
[1460124289] SERVICE ALERT: BBTS_DB;Size: IFDATA03.DBF;WARNING;HARD;3;Unknown: File not found
[1460124310] SERVICE ALERT: BBTS_DB;Size: TS_INDEX02.DBF;WARNING;SOFT;2;Unknown: File not found
[1460124337] SERVICE ALERT: BBTS_DB;Size: IFINDX03.DBF;WARNING;HARD;3;Unknown: File not found
[1460124383] SERVICE ALERT: BBTS_DB;Size: TS_DATA02.DBF;WARNING;HARD;3;Unknown: File not found
[1460124429] SERVICE ALERT: BBTS_DB;Size: TS_INDEX02.DBF;WARNING;HARD;3;Unknown: File not found
[1460124554] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 91% (5 min average)
[1460124674] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1460124794] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 95% (5 min average)
[1460124966] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (27.2% change >= 20.0% threshold)
[1460124966] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1460124966] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1460124966] HOST DOWNTIME ALERT: ADRODC-WEMU;STARTED; Host has entered a period of scheduled downtime
[1460124966] HOST DOWNTIME ALERT: IMAGES;STARTED; Host has entered a period of scheduled downtime
[1460124984] SERVICE ALERT: BBTS_DB;Size: TS_DATA02.DBF;OK;HARD;3;OK: Size is  13924.0078125  Mb
[1460125206] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (27.2% change >= 20.0% threshold)
[1460125206] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1460125206] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1460125206] HOST DOWNTIME ALERT: ADRODC-WEMU;STARTED; Host has entered a period of scheduled downtime
[1460125206] HOST DOWNTIME ALERT: IMAGES;STARTED; Host has entered a period of scheduled downtime
[1460125320] SERVICE ALERT: IMAGES;PXE Boot Service;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460125350] HOST ALERT: IMAGES;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460125366] SERVICE ALERT: IMAGES;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460125373] SERVICE ALERT: IMAGES;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460125373] SERVICE ALERT: IMAGES;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460125396] HOST ALERT: IMAGES;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460125417] SERVICE ALERT: IMAGES;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460125419] SERVICE ALERT: IMAGES;MDT Monitor;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460125419] SERVICE FLAPPING ALERT: IMAGES;MDT Monitor;STARTED; Service appears to have started flapping (24.0% change >= 20.0% threshold)
[1460125429] SERVICE ALERT: ADDC5;GPO Replication;WARNING;SOFT;1;CRITICAL! 180 policies are not replicating properly!
[1460125447] HOST ALERT: IMAGES;DOWN;SOFT;3;(Host check timed out after 30.01 seconds)
[1460125456] SERVICE ALERT: IMAGES;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460125458] SERVICE ALERT: IMAGES;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460125463] SERVICE ALERT: IMAGES;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460125463] SERVICE ALERT: IMAGES;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460125486] SERVICE ALERT: IMAGES;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460125486] HOST ALERT: IMAGES;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460125487] SERVICE ALERT: IMAGES;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460125490] SERVICE ALERT: BBTS_DB;Size: IFDATA03.DBF;OK;HARD;3;OK: Size is  100.0078125  Mb
[1460125493] SERVICE ALERT: IMAGES;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460125493] SERVICE ALERT: IMAGES;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460125513] SERVICE ALERT: IMAGES;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460125517] HOST ALERT: IMAGES;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460125537] SERVICE ALERT: BBTS_DB;Size: IFINDX03.DBF;OK;HARD;3;OK: Size is  3300.0078125  Mb
[1460125537] SERVICE ALERT: IMAGES;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460125556] SERVICE ALERT: IMAGES;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1460125558] SERVICE ALERT: IMAGES;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1460125563] SERVICE ALERT: IMAGES;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 1 minute(s)
[1460125563] SERVICE ALERT: IMAGES;Memory Usage;OK;SOFT;2;Memory usage: total:37842.86 MB - used: 1443.86 MB (4%) - free: 36399.01 MB (96%)
[1460125567] HOST ALERT: IMAGES;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1460125586] SERVICE ALERT: IMAGES;C:\ Drive Space;OK;SOFT;3;c:\ - total: 1157.56 Gb - used: 545.67 Gb (47%) - free 611.90 Gb (53%)
[1460125587] SERVICE ALERT: IMAGES;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1460125593] SERVICE ALERT: IMAGES;CPU Load;OK;SOFT;3;CPU Load 6% (5 min average)
[1460125593] SERVICE ALERT: IMAGES;Service: WMI;OK;SOFT;3;Winmgmt: Started
[1460125613] SERVICE ALERT: IMAGES;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1460125616] HOST ALERT: IMAGES;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1460125630] SERVICE ALERT: BBTS_DB;Size: TS_INDEX02.DBF;OK;HARD;3;OK: Size is  100.0078125  Mb
[1460125637] SERVICE ALERT: IMAGES;Service: Windows Time;OK;SOFT;3;W32Time: Started
[1460125663] HOST ALERT: IMAGES;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1460125727] HOST ALERT: IMAGES;UP;SOFT;9;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1460125730] SERVICE ALERT: ADDC5;GPO Replication;WARNING;SOFT;2;CRITICAL! 180 policies are not replicating properly!
[1460125768] SERVICE ALERT: hallefirst1;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460125828] SERVICE ALERT: hallefirst1;Printer Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460125877] SERVICE ALERT: hallefirst1;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1460125889] SERVICE ALERT: hallefirst1;Printer Status;OK;SOFT;3;Clearing paper path
[1460125900] SERVICE ALERT: IMAGES;PXE Boot Service;UNKNOWN;HARD;3;Failed to open service WDSServer: 1060: The specified service does not exist as an installed service.
[1460125900] SERVICE FLAPPING ALERT: IMAGES;PXE Boot Service;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1460125936] SERVICE ALERT: hallefirst1;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! All Trays Still Full
[1460126035] SERVICE ALERT: ADDC5;GPO Replication;WARNING;HARD;3;CRITICAL! 180 policies are not replicating properly!
[1460126049] HOST ALERT: IMAGES;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460126139] HOST ALERT: IMAGES;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460126176] SERVICE ALERT: IMAGES;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460126178] SERVICE ALERT: IMAGES;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460126180] HOST ALERT: IMAGES;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1460126183] SERVICE ALERT: IMAGES;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460126183] SERVICE ALERT: IMAGES;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460126276] SERVICE ALERT: IMAGES;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1460126278] SERVICE ALERT: IMAGES;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1460126283] SERVICE ALERT: IMAGES;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 2 minute(s)
[1460126283] SERVICE ALERT: IMAGES;Memory Usage;OK;SOFT;2;Memory usage: total:37842.86 MB - used: 1933.07 MB (5%) - free: 35909.79 MB (95%)
[1460126594] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 72% (5 min average)
[1460126599] SERVICE ALERT: IMAGES;MDT Monitor;UNKNOWN;HARD;3;Failed to open service MDT_Monitor: 1060: The specified service does not exist as an installed service.
[1460127461] HOST DOWNTIME ALERT: ADRODC-WEMU;STOPPED; Host has exited from a period of scheduled downtime
[1460127603] SERVICE FLAPPING ALERT: IMAGES;MDT Monitor;STARTED; Service appears to have started flapping (28.2% change >= 20.0% threshold)
[1460127603] SERVICE FLAPPING ALERT: IMAGES;PXE Boot Service;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1460127603] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1460127603] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1460127603] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1460127603] HOST DOWNTIME ALERT: IMAGES;STARTED; Host has entered a period of scheduled downtime
[1460127650] SERVICE ALERT: IMAGES;Service: MDT Monitor;UNKNOWN;SOFT;1;Failed to open service MDT_Monitor: 1060: The specified service does not exist as an installed service.
[1460127697] SERVICE ALERT: IMAGES;Service: PXE Boot;CRITICAL;SOFT;1;WDSServer: Stopped
[1460127700] SERVICE ALERT: IMAGES;PXE Boot Service;CRITICAL;HARD;3;WDSServer: Stopped
[1460127770] SERVICE ALERT: IMAGES;Service: MDT Monitor;UNKNOWN;SOFT;2;Failed to open service MDT_Monitor: 1060: The specified service does not exist as an installed service.
[1460127817] SERVICE ALERT: IMAGES;Service: PXE Boot;CRITICAL;SOFT;2;WDSServer: Stopped
[1460127890] SERVICE ALERT: IMAGES;Service: MDT Monitor;UNKNOWN;HARD;3;Failed to open service MDT_Monitor: 1060: The specified service does not exist as an installed service.
[1460127937] SERVICE ALERT: IMAGES;Service: PXE Boot;CRITICAL;HARD;3;WDSServer: Stopped
[1460128188] SERVICE ALERT: VCOMPOSER;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460128218] HOST ALERT: VCOMPOSER;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460128235] SERVICE ALERT: VCOMPOSER;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460128265] HOST ALERT: VCOMPOSER;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460128269] SERVICE ALERT: SNAP;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1460128272] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.40)
[1460128284] SERVICE ALERT: VCOMPOSER;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460128297] SERVICE ALERT: VCOMPOSER;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460128308] SERVICE ALERT: VCOMPOSER;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460128309] HOST ALERT: SNAP;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.40 ms
[1460128314] HOST ALERT: VCOMPOSER;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460128330] SERVICE ALERT: VCOMPOSER;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460128355] SERVICE ALERT: VCOMPOSER;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460128360] HOST ALERT: VCOMPOSER;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460128377] SERVICE ALERT: SNAP;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1460128377] SERVICE ALERT: VCOMPOSER;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460128377] SERVICE ALERT: VCOMPOSER;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460128404] SERVICE ALERT: VCOMPOSER;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460128407] HOST ALERT: VCOMPOSER;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460128417] SERVICE ALERT: VCOMPOSER;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460128428] SERVICE ALERT: VCOMPOSER;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460128447] HOST ALERT: VCOMPOSER;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1460128450] SERVICE ALERT: VCOMPOSER;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460128475] SERVICE ALERT: VCOMPOSER;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460128480] HOST ALERT: VCOMPOSER;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1460128487] SERVICE ALERT: VCOMPOSER;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460128497] SERVICE ALERT: VCOMPOSER;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460128497] SERVICE ALERT: VCOMPOSER;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460128508] SERVICE ALERT: VCOMPOSER;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460128514] SERVICE ALERT: VCOMPOSER;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460128517] HOST ALERT: VCOMPOSER;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1460128524] SERVICE ALERT: VCOMPOSER;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460128537] SERVICE ALERT: VCOMPOSER;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460128554] HOST ALERT: VCOMPOSER;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1460128556] SERVICE ALERT: VCOMPOSER;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460128570] SERVICE ALERT: VCOMPOSER;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460128586] HOST ALERT: VCOMPOSER;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1460128607] SERVICE ALERT: VCOMPOSER;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460128607] SERVICE ALERT: VCOMPOSER;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460128617] SERVICE ALERT: VCOMPOSER;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460128617] SERVICE ALERT: VCOMPOSER;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460128628] SERVICE ALERT: VCOMPOSER;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460128634] SERVICE ALERT: VCOMPOSER;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460128649] SERVICE ALERT: VCOMPOSER;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460128676] SERVICE ALERT: VCOMPOSER;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460128695] SERVICE ALERT: VCOMPOSER;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460128727] SERVICE ALERT: VCOMPOSER;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460128727] SERVICE ALERT: VCOMPOSER;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460128748] SERVICE ALERT: VCOMPOSER;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460128754] SERVICE ALERT: VCOMPOSER;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460128769] SERVICE ALERT: VCOMPOSER;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460128796] SERVICE ALERT: VCOMPOSER;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460128815] SERVICE ALERT: VCOMPOSER;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460128847] SERVICE ALERT: VCOMPOSER;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460128862] HOST ALERT: VCOMPOSER;UP;HARD;10;PING OK - Packet loss = 68%, RTA = 0.35 ms
[1460128869] SERVICE ALERT: VCOMPOSER;Service: Group Policy;OK;SOFT;3;gpsvc: Started
[1460128915] SERVICE ALERT: VCOMPOSER;Service: NetLogin;OK;SOFT;3;Netlogon: Started
[1460128941] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.40)
[1460128956] SERVICE ALERT: SNAP;Uptime;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1460128959] HOST ALERT: SNAP;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.40)
[1460128982] HOST ALERT: SNAP;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.39 ms
[1460129008] SERVICE ALERT: VCOMPOSER;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1460129055] SERVICE ALERT: VCOMPOSER;Service: WMI;OK;HARD;3;Winmgmt: Started
[1460129073] SERVICE ALERT: SNAP;Uptime;OK;SOFT;2;System Uptime - 34 day(s) 15 hour(s) 35 minute(s)
[1460129104] SERVICE ALERT: VCOMPOSER;Service: Windows Time;OK;HARD;3;W32Time: Started
[1460129104] SERVICE ALERT: SNAP;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460129115] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.40)
[1460129117] SERVICE ALERT: VCOMPOSER;Memory Usage;OK;HARD;3;Memory usage: total:20477.93 MB - used: 2444.12 MB (12%) - free: 18033.81 MB (88%)
[1460129135] SERVICE ALERT: SNAP;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1460129138] HOST ALERT: SNAP;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.40)
[1460129150] SERVICE ALERT: VCOMPOSER;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 5 minute(s)
[1460129180] SERVICE ALERT: SNAP;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1460129183] HOST ALERT: SNAP;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.102.40)
[1460129197] SERVICE ALERT: VCOMPOSER;C:\ Drive Space;OK;HARD;3;c:\ - total: 199.90 Gb - used: 22.67 Gb (11%) - free 177.23 Gb (89%)
[1460129197] SERVICE ALERT: VCOMPOSER;VMWare Tools;OK;HARD;3;VMTools: Started
[1460129204] SERVICE ALERT: SNAP;Memory Usage;OK;SOFT;2;Memory usage: total:16362.77 MB - used: 1504.17 MB (9%) - free: 14858.60 MB (91%)
[1460129208] HOST ALERT: SNAP;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.39 ms
[1460129208] HOST FLAPPING ALERT: SNAP;STARTED; Host appears to have started flapping (21.1% change > 20.0% threshold)
[1460129272] SERVICE ALERT: SNAP;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460129275] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.40)
[1460129275] SERVICE ALERT: SNAP;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1460129300] SERVICE ALERT: SNAP;Service: DNS Client;CRITICAL;SOFT;2;connect to address 164.76.102.40 and port 12489: No route to host
[1460129303] HOST ALERT: SNAP;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.40)
[1460129307] SERVICE ALERT: VCOMPOSER;CPU Load;OK;HARD;3;CPU Load 0% (5 min average)
[1460129323] HOST ALERT: SNAP;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.41 ms
[1460129328] SERVICE ALERT: VCOMPOSER;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1460129334] SERVICE ALERT: VCOMPOSER;Service: RPC;OK;HARD;3;RpcSs: Started
[1460129372] SERVICE ALERT: SNAP;NSClient++ Version;OK;SOFT;3;NSClient++ 0.4.4.15 2015-11-25
[1460129376] SERVICE ALERT: VCOMPOSER;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1460129392] SERVICE ALERT: SNAP;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1460129417] SERVICE ALERT: SNAP;Service: DNS Client;OK;SOFT;3;Dnscache: Started
[1460129427] SERVICE ALERT: VCOMPOSER;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1460129931] HOST FLAPPING ALERT: SNAP;STOPPED; Host appears to have stopped flapping (4.3% change < 5.0% threshold)
[1460130100] SERVICE ALERT: IMAGES;PXE Boot Service;OK;HARD;3;WDSServer: Started
[1460130337] SERVICE ALERT: IMAGES;Service: PXE Boot;OK;HARD;3;WDSServer: Started
[1460130794] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460130914] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1460131034] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1460131634] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 81% (5 min average)
[1460132234] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1460132354] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1460132474] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 95% (5 min average)
[1460133074] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 73% (5 min average)
[1460133241] SERVICE ALERT: ADDC5;GPO Replication;OK;HARD;3;OK! All GPOs are properly replicated.
[1460133303] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1460133423] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 69% (5 min average)
[1460134033] SERVICE ALERT: studentcenterpk;Printer Status;OK;HARD;3;Clearing paper path
[1460134238] HOST DOWNTIME ALERT: IMAGES;STOPPED; Host has exited from a period of scheduled downtime
[1460135233] SERVICE ALERT: studentcenterpk;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460135293] SERVICE ALERT: studentcenterpk;Printer Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460135354] SERVICE ALERT: studentcenterpk;Printer Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460135473] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460135593] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460135713] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460136313] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 90% (5 min average)
[1460136913] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 67% (5 min average)
[1460137399] SERVICE FLAPPING ALERT: IMAGES;MDT Monitor;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1460137513] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460137623] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 91% (5 min average)
[1460137633] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 67% (5 min average)
[1460137743] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 74% (5 min average)
[1460138233] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460138353] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1460138426] SERVICE ALERT: INTLAPP-ISSMAP;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1460138473] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460138546] SERVICE ALERT: INTLAPP-ISSMAP;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1460138666] SERVICE ALERT: INTLAPP-ISSMAP;CPU Load;OK;SOFT;3;CPU Load 0% (5 min average)
[1460139072] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 94% (5 min average)
[1460139700] SERVICE FLAPPING ALERT: IMAGES;PXE Boot Service;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1460140110] SERVICE ALERT: NETENG-PA;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460140154] SERVICE ALERT: studentcenterpk;Printer Status;OK;HARD;3;Clearing paper path
[1460140223] SERVICE ALERT: studentcenterpk;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Maintenance Kit HP 110V-C9152A, 220V-C9153A'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1460140230] SERVICE ALERT: NETENG-PA;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1460140272] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460140284] SERVICE ALERT: studentcenterpk;Maintenance Kit HP 110V-C9152A;OK;SOFT;2;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 88% - OK!
[1460140759] SERVICE ALERT: studentcenterpk;Printer Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.150.252. Is the printer off?  Is the network down?
[1460140789] HOST ALERT: studentcenterpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460140819] SERVICE ALERT: studentcenterpk;Printer Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.150.252. Is the printer off?  Is the network down?
[1460140849] HOST ALERT: studentcenterpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460140872] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 84% (5 min average)
[1460140879] SERVICE ALERT: studentcenterpk;Printer Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.150.252. Is the printer off?  Is the network down?
[1460140879] SERVICE FLAPPING ALERT: studentcenterpk;Printer Status;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1460140889] SERVICE ALERT: studentcenterpk;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.150.252. Is the printer off?  Is the network down?
[1460140909] HOST ALERT: studentcenterpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460140949] SERVICE ALERT: studentcenterpk;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.150.252. Is the printer off?  Is the network down?
[1460140979] HOST ALERT: studentcenterpk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460141009] SERVICE ALERT: studentcenterpk;Maintenance Kit HP 110V-C9152A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.150.252. Is the printer off?  Is the network down?
[1460141039] HOST ALERT: studentcenterpk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460141051] SERVICE ALERT: studentcenterpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.150.252. Is the printer off?  Is the network down?
[1460141081] HOST ALERT: studentcenterpk;DOWN;SOFT;6;(Host check timed out after 30.01 seconds)
[1460141084] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.150.252. Is the printer off?  Is the network down?
[1460141111] SERVICE ALERT: studentcenterpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.150.252. Is the printer off?  Is the network down?
[1460141114] HOST ALERT: studentcenterpk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1460141144] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.150.252. Is the printer off?  Is the network down?
[1460141171] SERVICE ALERT: studentcenterpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.150.252. Is the printer off?  Is the network down?
[1460141174] HOST ALERT: studentcenterpk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1460141204] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.150.252. Is the printer off?  Is the network down?
[1460141234] HOST ALERT: studentcenterpk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1460141324] HOST ALERT: studentcenterpk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1460141472] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1460141494] SERVICE ALERT: hallefirst1;Printer Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1460141524] HOST ALERT: hallefirst1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460141541] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1460141554] SERVICE ALERT: hallefirst1;Printer Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1460141560] HOST ALERT: hallefirst1;UP;SOFT;2;PING OK - Packet loss = 75%, RTA = 0.65 ms
[1460141592] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 84% (5 min average)
[1460141595] SERVICE ALERT: hallefirst1;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! 1 of 3 Trays Still Full
[1460141608] SERVICE ALERT: hallefirst1;Printer Status;OK;SOFT;3;RETRIEVE JOB  NO STORED JOBS
[1460141943] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1460142063] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 76% (5 min average)
[1460142757] SERVICE ALERT: C4SAS-SE;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1460142877] SERVICE ALERT: C4SAS-SE;CPU Load;OK;SOFT;2;CPU Load 25% (5 min average)
[1460143392] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460143512] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460143632] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460143863] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1460143983] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 90% (5 min average)
[1460144103] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 83% (5 min average)
[1460144232] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 60% (5 min average)
[1460144832] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1460144952] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 94% (5 min average)
[1460145072] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1460146272] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 45% (5 min average)
[1460146872] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1460146992] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 93% (5 min average)
[1460147112] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 79% (5 min average)
[1460148312] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1460148432] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1460148552] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460148915] SERVICE ALERT: VCOMPOSER;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.2.61 and port 12489: Connection refused
[1460148923] SERVICE ALERT: VCOMPOSER;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460149023] SERVICE ALERT: VCOMPOSER;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1460149035] SERVICE ALERT: VCOMPOSER;Memory Usage;OK;SOFT;2;Memory usage: total:20477.93 MB - used: 1845.08 MB (9%) - free: 18632.85 MB (91%)
[1460149152] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 80% (5 min average)
[1460149752] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460149872] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460149992] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 94% (5 min average)
[1460150651] SERVICE ALERT: ADCA1;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460150771] SERVICE ALERT: ADCA1;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1460150849] SERVICE ALERT: VIEW-CS;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460150879] HOST ALERT: VIEW-CS;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460150891] SERVICE ALERT: ADCA1;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1460150894] SERVICE ALERT: VIEW-CS;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460150913] SERVICE ALERT: VIEW-CS;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460150924] HOST ALERT: VIEW-CS;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460150937] SERVICE ALERT: ADCA2;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1460150969] SERVICE ALERT: VIEW-CS;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460150985] SERVICE ALERT: VIEW-CS;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460150999] HOST ALERT: VIEW-CS;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460151014] SERVICE ALERT: VIEW-CS;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460151033] SERVICE ALERT: VIEW-CS;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460151033] SERVICE ALERT: VIEW-CS;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460151044] HOST ALERT: VIEW-CS;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460151057] SERVICE ALERT: ADCA2;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1460151069] HOST ALERT: VIEW-CS;UP;SOFT;5;PING OK - Packet loss = 28%, RTA = 0.33 ms
[1460151069] SERVICE ALERT: VIEW-CS;Service: Firewall;OK;SOFT;3;MpsSvc: Started
[1460151105] SERVICE ALERT: VIEW-CS;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460151128] SERVICE ALERT: VIEW-CS;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460151134] SERVICE ALERT: VIEW-CS;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460151135] HOST ALERT: VIEW-CS;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460151153] SERVICE ALERT: VIEW-CS;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460151153] SERVICE ALERT: VIEW-CS;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460151166] SERVICE ALERT: VIEW-CS;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460151173] SERVICE ALERT: VIEW-CS;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460151177] SERVICE ALERT: ADCA2;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1460151183] HOST ALERT: VIEW-CS;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460151192] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460151207] SERVICE ALERT: VIEW-CS;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460151225] SERVICE ALERT: VIEW-CS;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460151237] HOST ALERT: VIEW-CS;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460151248] SERVICE ALERT: VIEW-CS;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460151273] SERVICE ALERT: VIEW-CS;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460151278] HOST ALERT: VIEW-CS;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460151286] SERVICE ALERT: VIEW-CS;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460151293] SERVICE ALERT: VIEW-CS;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460151306] SERVICE ALERT: VIEW-CS;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460151316] HOST ALERT: VIEW-CS;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460151327] SERVICE ALERT: VIEW-CS;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460151352] SERVICE ALERT: VIEW-CS;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460151357] HOST ALERT: VIEW-CS;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1460151362] SERVICE ALERT: VIEW-CS;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460151368] SERVICE ALERT: VIEW-CS;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460151392] HOST ALERT: VIEW-CS;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1460151406] SERVICE ALERT: VIEW-CS;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460151406] SERVICE ALERT: VIEW-CS;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460151413] SERVICE ALERT: VIEW-CS;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460151426] SERVICE ALERT: VIEW-CS;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460151436] HOST ALERT: VIEW-CS;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1460151447] SERVICE ALERT: VIEW-CS;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460151472] SERVICE ALERT: VIEW-CS;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460151477] HOST ALERT: VIEW-CS;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1460151482] SERVICE ALERT: VIEW-CS;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460151512] HOST ALERT: VIEW-CS;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1460151526] SERVICE ALERT: VIEW-CS;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460151546] SERVICE ALERT: VIEW-CS;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460151587] SERVICE ALERT: SQL1;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460151592] SERVICE ALERT: VIEW-CS;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460151602] SERVICE ALERT: VIEW-CS;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460151604] HOST ALERT: SQL1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460151630] SERVICE ALERT: SQL1;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460151646] SERVICE ALERT: VIEW-CS;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460151660] HOST ALERT: SQL1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460151683] SERVICE ALERT: SQL1;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460151683] SERVICE ALERT: VIEW-CS;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460151689] SERVICE ALERT: VIEW-CS;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460151707] SERVICE ALERT: SQL1;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460151713] SERVICE ALERT: SQL1;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460151713] HOST ALERT: SQL1;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460151723] SERVICE ALERT: SQL1;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460151750] SERVICE ALERT: SQL1;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460151753] HOST ALERT: SQL1;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460151773] SERVICE ALERT: SQL1;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460151777] SERVICE ALERT: ADCA2;CPU Load;OK;HARD;3;CPU Load 0% (5 min average)
[1460151803] SERVICE ALERT: SQL1;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460151803] HOST ALERT: SQL1;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460151803] SERVICE ALERT: VIEW-CS;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460151809] SERVICE ALERT: VIEW-CS;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460151817] SERVICE ALERT: SQL1;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460151822] SERVICE ALERT: SQL1;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1460151823] SERVICE ALERT: SQL1;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1460151827] SERVICE ALERT: SQL1;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460151833] HOST ALERT: SQL1;UP;SOFT;6;PING WARNING - Packet loss = 86%, RTA = 0.35 ms
[1460151850] SERVICE ALERT: SQL1;Service: NetLogin;OK;SOFT;3;Netlogon: Started
[1460151873] SERVICE ALERT: SQL1;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1460151903] SERVICE ALERT: SQL1;Service: RPC;OK;SOFT;3;RpcSs: Started
[1460151917] SERVICE ALERT: SQL1;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1460151923] SERVICE ALERT: VIEW-CS;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460151929] SERVICE ALERT: VIEW-CS;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460152091] SERVICE ALERT: ADCA1;CPU Load;OK;HARD;3;CPU Load 0% (5 min average)
[1460152279] SERVICE FLAPPING ALERT: studentcenterpk;Printer Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1460152407] SERVICE ALERT: SQL1;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1460152890] SERVICE ALERT: ADAUTHDC1;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460152917] SERVICE ALERT: ADAUTHDC3;Service: DNS Service;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1460152937] SERVICE ALERT: ADAUTHDC2;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460152937] SERVICE ALERT: ADAUTHDC2;Service: DFS Namespace;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460152943] SERVICE ALERT: ADAUTHDC3;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460152945] SERVICE ALERT: CSCFILESERV;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.102.131 and port 12489: No route to host
[1460152948] HOST ALERT: CSCFILESERV;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.131)
[1460152963] SERVICE ALERT: ADAUTHDC3;Service: Dell Backup Agent;CRITICAL;SOFT;1;ErdAgent: Starting
[1460152964] SERVICE ALERT: ADAUTHDC2;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1460152968] HOST ALERT: markjeffersonpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460152978] SERVICE ALERT: CSCFILESERV;Uptime;CRITICAL;SOFT;1;connect to address 164.76.102.131 and port 12489: No route to host
[1460152981] HOST ALERT: CSCFILESERV;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.131)
[1460152990] SERVICE ALERT: ADAUTHDC1;Service: RPC;OK;SOFT;2;RpcSs: Started
[1460152992] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 50% (5 min average)
[1460152994] SERVICE ALERT: markjeffersonpk;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1460153024] HOST ALERT: markjeffersonpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460153029] SERVICE ALERT: CSCFILESERV;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.102.131 and port 12489: No route to host
[1460153032] HOST ALERT: CSCFILESERV;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.102.131)
[1460153037] SERVICE ALERT: ADAUTHDC2;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1460153037] SERVICE ALERT: ADAUTHDC3;Service: DNS Service;OK;SOFT;2;DNS: Started
[1460153040] SERVICE ALERT: ADAUTHDC2;Service: DFS Namespace;OK;SOFT;2;Dfs: Started
[1460153043] SERVICE ALERT: ADAUTHDC3;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 1 minute(s)
[1460153047] SERVICE ALERT: ADAUTHDC2;Service: Dell Backup Agent;CRITICAL;SOFT;1;ErdAgent: Starting
[1460153047] SERVICE ALERT: ADCA1;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.7.150 and port 12489: Connection refused
[1460153054] SERVICE ALERT: markjeffersonpk;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1460153056] SERVICE ALERT: CSCFILESERV;Service: Windows Time;CRITICAL;SOFT;2;connect to address 164.76.102.131 and port 12489: No route to host
[1460153059] HOST ALERT: CSCFILESERV;DOWN;SOFT;4;CRITICAL - Host Unreachable (164.76.102.131)
[1460153068] SERVICE ALERT: CSCFILESERV;CPU Load;CRITICAL;SOFT;1;connect to address 164.76.102.131 and port 12489: No route to host
[1460153071] SERVICE ALERT: markjeffersonpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1460153083] SERVICE ALERT: ADAUTHDC3;Service: Dell Backup Agent;CRITICAL;SOFT;2;ErdAgent: Starting
[1460153084] SERVICE ALERT: ADAUTHDC2;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1460153084] HOST ALERT: markjeffersonpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460153098] SERVICE ALERT: CSCFILESERV;Uptime;CRITICAL;SOFT;2;connect to address 164.76.102.131 and port 12489: No route to host
[1460153101] HOST ALERT: CSCFILESERV;DOWN;SOFT;5;CRITICAL - Host Unreachable (164.76.102.131)
[1460153114] SERVICE ALERT: markjeffersonpk;Maintenance Kit HP 110V-C9152A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1460153116] SERVICE ALERT: CSCFILESERV;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.102.131 and port 12489: No route to host
[1460153119] HOST ALERT: CSCFILESERV;DOWN;SOFT;6;CRITICAL - Host Unreachable (164.76.102.131)
[1460153131] SERVICE ALERT: markjeffersonpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1460153144] HOST ALERT: markjeffersonpk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460153148] SERVICE ALERT: CSCFILESERV;C:\ Drive Space;CRITICAL;SOFT;2;connect to address 164.76.102.131 and port 12489: No route to host
[1460153150] SERVICE ALERT: ADDC6;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1460153151] HOST ALERT: CSCFILESERV;DOWN;SOFT;7;CRITICAL - Host Unreachable (164.76.102.131)
[1460153160] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1460153160] SERVICE ALERT: CSCFILESERV;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.102.131 and port 12489: No route to host
[1460153167] SERVICE ALERT: ADAUTHDC2;Service: Dell Backup Agent;OK;SOFT;2;ErdAgent: Started
[1460153167] SERVICE ALERT: ADCA1;C:\ Drive Space;OK;SOFT;2;c:\ - total: 59.90 Gb - used: 31.59 Gb (53%) - free 28.31 Gb (47%)
[1460153176] SERVICE ALERT: CSCFILESERV;Service: Windows Time;CRITICAL;HARD;3;connect to address 164.76.102.131 and port 12489: No route to host
[1460153179] HOST ALERT: CSCFILESERV;DOWN;SOFT;8;CRITICAL - Host Unreachable (164.76.102.131)
[1460153188] SERVICE ALERT: CSCFILESERV;CPU Load;CRITICAL;SOFT;2;connect to address 164.76.102.131 and port 12489: No route to host
[1460153190] HOST ALERT: markjeffersonpk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460153191] SERVICE ALERT: markjeffersonpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1460153195] SERVICE ALERT: ADAUTHDC3;AD Replication;UNKNOWN;SOFT;1;CHECK_NRPE: Invalid packet type received from server.
[1460153203] SERVICE ALERT: ADAUTHDC3;Service: Dell Backup Agent;OK;SOFT;3;ErdAgent: Started
[1460153205] SERVICE ALERT: SP2010DB;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1460153207] SERVICE ALERT: CSCFILESERV;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.102.131 and port 12489: No route to host
[1460153210] HOST ALERT: CSCFILESERV;DOWN;SOFT;9;CRITICAL - Host Unreachable (164.76.102.131)
[1460153218] SERVICE ALERT: CSCFILESERV;Uptime;CRITICAL;HARD;3;connect to address 164.76.102.131 and port 12489: No route to host
[1460153220] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1460153221] HOST ALERT: markjeffersonpk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1460153226] SERVICE ALERT: TARGETX-2012;CPU Load;WARNING;SOFT;1;CPU Load 94% (5 min average)
[1460153236] SERVICE ALERT: CSCFILESERV;Memory Usage;CRITICAL;SOFT;2;connect to address 164.76.102.131 and port 12489: No route to host
[1460153237] SERVICE ALERT: ADAUTHDC1;AD Replication;WARNING;SOFT;1;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1460153239] HOST ALERT: CSCFILESERV;DOWN;HARD;10;CRITICAL - Host Unreachable (164.76.102.131)
[1460153246] SERVICE ALERT: ADDC6;Service: Dell Backup Agent;CRITICAL;SOFT;1;ErdAgent: Starting
[1460153255] SERVICE ALERT: CSCFILESERV;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.102.131 and port 12489: No route to host
[1460153267] SERVICE ALERT: ADDC6;Service: Windows Time;CRITICAL;SOFT;2;W32Time: Stopped
[1460153268] SERVICE ALERT: CSCFILESERV;C:\ Drive Space;CRITICAL;HARD;3;connect to address 164.76.102.131 and port 12489: No route to host
[1460153279] SERVICE ALERT: markjeffersonpk;Printer Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1460153280] SERVICE ALERT: CSCFILESERV;NSClient++ Version;CRITICAL;SOFT;2;connect to address 164.76.102.131 and port 12489: No route to host
[1460153280] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1460153300] SERVICE ALERT: CSCFILESERV;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.102.131 and port 12489: No route to host
[1460153308] SERVICE ALERT: CSCFILESERV;CPU Load;CRITICAL;HARD;3;connect to address 164.76.102.131 and port 12489: No route to host
[1460153309] HOST ALERT: markjeffersonpk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1460153314] SERVICE ALERT: ADAUTHDC3;AD Replication;UNKNOWN;SOFT;2;CHECK_NRPE: Invalid packet type received from server.
[1460153325] SERVICE ALERT: SP2010DB;CPU Load;OK;SOFT;2;CPU Load 72% (5 min average)
[1460153326] SERVICE ALERT: CSCFILESERV;Service: DNS Client;CRITICAL;SOFT;2;connect to address 164.76.102.131 and port 12489: No route to host
[1460153339] SERVICE ALERT: markjeffersonpk;Printer Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1460153346] SERVICE ALERT: TARGETX-2012;CPU Load;OK;SOFT;2;CPU Load 12% (5 min average)
[1460153348] SERVICE ALERT: CSCFILESERV;Service: NetLogin;CRITICAL;SOFT;1;connect to address 164.76.102.131 and port 12489: No route to host
[1460153356] SERVICE ALERT: CSCFILESERV;Memory Usage;CRITICAL;HARD;3;connect to address 164.76.102.131 and port 12489: No route to host
[1460153356] SERVICE ALERT: ADAUTHDC1;AD Replication;WARNING;SOFT;2;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1460153366] SERVICE ALERT: ADDC6;Service: Dell Backup Agent;OK;SOFT;2;ErdAgent: Started
[1460153369] HOST ALERT: markjeffersonpk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1460153373] SERVICE ALERT: SP2010;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460153375] SERVICE ALERT: CSCFILESERV;Service: Firewall;CRITICAL;SOFT;2;connect to address 164.76.102.131 and port 12489: No route to host
[1460153387] SERVICE ALERT: ADDC6;Service: Windows Time;OK;SOFT;3;W32Time: Started
[1460153394] SERVICE ALERT: CSCFILESERV;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.102.131 and port 12489: No route to host
[1460153397] SERVICE ALERT: CSCFILESERV;NSClient++ Version;CRITICAL;HARD;3;connect to address 164.76.102.131 and port 12489: No route to host
[1460153399] SERVICE ALERT: markjeffersonpk;Printer Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1460153420] SERVICE ALERT: CSCFILESERV;Service: Group Policy;CRITICAL;SOFT;2;connect to address 164.76.102.131 and port 12489: No route to host
[1460153429] HOST ALERT: markjeffersonpk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1460153434] SERVICE ALERT: ADAUTHDC3;AD Replication;UNKNOWN;HARD;3;CHECK_NRPE: Invalid packet type received from server.
[1460153443] SERVICE ALERT: CSCFILESERV;Service: User Profile;CRITICAL;SOFT;1;connect to address 164.76.102.131 and port 12489: No route to host
[1460153446] SERVICE ALERT: CSCFILESERV;Service: DNS Client;CRITICAL;HARD;3;connect to address 164.76.102.131 and port 12489: No route to host
[1460153463] SERVICE ALERT: TARGETX-2012;Service: Group Policy;CRITICAL;SOFT;1;gpsvc: Stopped
[1460153468] SERVICE ALERT: CSCFILESERV;Service: NetLogin;CRITICAL;SOFT;2;connect to address 164.76.102.131 and port 12489: No route to host
[1460153476] SERVICE ALERT: ADAUTHDC1;AD Replication;WARNING;HARD;3;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1460153489] SERVICE ALERT: CSCFILESERV;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.102.131 and port 12489: No route to host
[1460153491] SERVICE ALERT: EMUFILES;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460153493] SERVICE ALERT: SP2010;CPU Load;OK;SOFT;2;CPU Load 87% (5 min average)
[1460153495] SERVICE ALERT: CSCFILESERV;Service: Firewall;CRITICAL;HARD;3;connect to address 164.76.102.131 and port 12489: No route to host
[1460153514] SERVICE ALERT: CSCFILESERV;Service: RPC;CRITICAL;SOFT;2;connect to address 164.76.102.131 and port 12489: No route to host
[1460153519] HOST ALERT: markjeffersonpk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1460153540] SERVICE ALERT: CSCFILESERV;Service: Group Policy;CRITICAL;HARD;3;connect to address 164.76.102.131 and port 12489: No route to host
[1460153555] SERVICE ALERT: ADDC6;Remote Loader for IDM;CRITICAL;SOFT;1;dirxml_remote.exe: not running
[1460153563] SERVICE ALERT: CSCFILESERV;Service: User Profile;CRITICAL;SOFT;2;connect to address 164.76.102.131 and port 12489: No route to host
[1460153577] SERVICE ALERT: ADCA2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460153583] SERVICE ALERT: TARGETX-2012;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1460153588] SERVICE ALERT: CSCFILESERV;Service: NetLogin;CRITICAL;HARD;3;connect to address 164.76.102.131 and port 12489: No route to host
[1460153592] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1460153609] SERVICE ALERT: CSCFILESERV;Service: WMI;CRITICAL;SOFT;2;connect to address 164.76.102.131 and port 12489: No route to host
[1460153611] SERVICE ALERT: EMUFILES;CPU Load;OK;SOFT;2;CPU Load 85% (5 min average)
[1460153634] SERVICE ALERT: CSCFILESERV;Service: RPC;CRITICAL;HARD;3;connect to address 164.76.102.131 and port 12489: No route to host
[1460153675] SERVICE ALERT: ADDC6;Remote Loader for IDM;CRITICAL;SOFT;2;dirxml_remote.exe: not running
[1460153677] SERVICE ALERT: C4SAS-SE;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460153683] SERVICE ALERT: CSCFILESERV;Service: User Profile;CRITICAL;HARD;3;connect to address 164.76.102.131 and port 12489: No route to host
[1460153697] SERVICE ALERT: ADCA2;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1460153703] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1460153712] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460153723] SERVICE ALERT: AXIOM-2012;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460153729] SERVICE ALERT: CSCFILESERV;Service: WMI;CRITICAL;HARD;3;connect to address 164.76.102.131 and port 12489: No route to host
[1460153795] SERVICE ALERT: ADDC6;Remote Loader for IDM;CRITICAL;HARD;3;dirxml_remote.exe: not running
[1460153797] SERVICE ALERT: C4SAS-SE;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1460153823] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 44% (5 min average)
[1460153823] SERVICE ALERT: AXIOM-2012;Memory Usage;OK;SOFT;2;Memory usage: total:5503.60 MB - used: 1297.67 MB (24%) - free: 4205.93 MB (76%)
[1460153832] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460153911] SERVICE ALERT: ADSECURED01;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460153917] SERVICE ALERT: C4SAS-SE;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1460153925] SERVICE ALERT: SP2010DB;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460153941] HOST ALERT: ADSECURED01;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460153956] SERVICE ALERT: ADSECURED01;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460153973] SERVICE ALERT: IT-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460153986] HOST ALERT: ADSECURED01;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460154003] SERVICE ALERT: ADSECURED01;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460154025] SERVICE ALERT: ADAUTHDC1;Service: DFS Namespace;CRITICAL;SOFT;1;connect to address 164.76.102.171 and port 12489: Connection refused
[1460154025] SERVICE ALERT: ADAUTHDC2;Service: Remote Management;CRITICAL;SOFT;1;connect to address 164.76.102.172 and port 12489: Connection refused
[1460154025] SERVICE ALERT: ADAUTHDC1;Service: NetLogin;CRITICAL;SOFT;1;connect to address 164.76.102.171 and port 12489: Connection refused
[1460154027] SERVICE ALERT: ADAUTHDC2;Service: AD Domain Service;CRITICAL;SOFT;1;connect to address 164.76.102.172 and port 12489: Connection refused
[1460154031] SERVICE ALERT: ADSECURED01;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460154033] HOST ALERT: ADSECURED01;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460154045] SERVICE ALERT: SP2010DB;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1460154047] SERVICE ALERT: ADAUTHDC1;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460154050] SERVICE ALERT: ADSECURED01;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460154073] HOST ALERT: ADSECURED01;UP;SOFT;4;PING OK - Packet loss = 79%, RTA = 0.35 ms
[1460154076] SERVICE ALERT: ADSECURED01;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460154079] SERVICE ALERT: ADDFS01;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.99.199 and port 12489: Connection refused
[1460154093] SERVICE ALERT: IT-2012;CPU Load;WARNING;SOFT;2;CPU Load 93% (5 min average)
[1460154093] SERVICE ALERT: SP2010;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1460154103] SERVICE ALERT: ADSECURED01;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1460154131] SERVICE ALERT: ADSECURED01;NSClient++ Version;OK;SOFT;3;NSClient++ 0.4.4.15 2015-11-25
[1460154145] SERVICE ALERT: ADAUTHDC1;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1460154145] SERVICE ALERT: ADAUTHDC1;Service: DFS Namespace;OK;SOFT;2;Dfs: Started
[1460154145] SERVICE ALERT: ADAUTHDC2;Service: Remote Management;OK;SOFT;2;WinRM: Started
[1460154147] SERVICE ALERT: ADAUTHDC2;Service: AD Domain Service;OK;SOFT;2;NTDS: Started
[1460154147] SERVICE ALERT: ADAUTHDC1;Memory Usage;OK;SOFT;2;Memory usage: total:10943.64 MB - used: 829.13 MB (8%) - free: 10114.51 MB (92%)
[1460154150] SERVICE ALERT: ADSECURED01;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1460154165] SERVICE ALERT: SP2010DB;CPU Load;OK;SOFT;3;CPU Load 57% (5 min average)
[1460154176] SERVICE ALERT: ADSECURED01;Service: DNS Client;OK;SOFT;3;Dnscache: Started
[1460154196] SERVICE ALERT: ADDFS01;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1460154213] SERVICE ALERT: IT-2012;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1460154213] SERVICE ALERT: SP2010;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1460154333] SERVICE ALERT: SP2010;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1460154338] SERVICE ALERT: TARGETX-2012;Uptime;CRITICAL;SOFT;1;connect to address 164.76.99.154 and port 12489: Connection refused
[1460154368] HOST ALERT: TARGETX-2012;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460154382] HOST ALERT: TARGETX-2012;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1460154395] SERVICE ALERT: ADDC6;Remote Loader for IDM;OK;HARD;3;dirxml_remote.exe: Running
[1460154403] SERVICE ALERT: ADAUTHDC3;Service: Dell Backup Agent;CRITICAL;SOFT;1;ErdAgent: Starting
[1460154449] SERVICE ALERT: TARGETX-2012;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 2 minute(s)
[1460154523] SERVICE ALERT: ADAUTHDC3;Service: Dell Backup Agent;OK;SOFT;2;ErdAgent: Started
[1460154706] SERVICE ALERT: C4SAS-SE;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460154765] SERVICE ALERT: SP2010DB;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460154806] SERVICE ALERT: C4SAS-SE;Memory Usage;OK;SOFT;2;Memory usage: total:10237.85 MB - used: 1207.79 MB (12%) - free: 9030.06 MB (88%)
[1460154808] SERVICE ALERT: SP2010;VMWare Tools;CRITICAL;SOFT;1;connect to address 164.76.158.63 and port 12489: Connection refused
[1460154813] SERVICE ALERT: IT-2012;CPU Load;OK;HARD;3;CPU Load 5% (5 min average)
[1460154817] SERVICE ALERT: SP2010;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460154838] HOST ALERT: SP2010;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460154886] SERVICE ALERT: SP2010DB;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1460154886] HOST ALERT: SP2010;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1460154917] SERVICE ALERT: SP2010;C:\ Drive Space;OK;SOFT;2;c:\ - total: 79.90 Gb - used: 57.82 Gb (72%) - free 22.08 Gb (28%)
[1460154921] SERVICE ALERT: SP2010;VMWare Tools;OK;SOFT;2;VMTools: Started
[1460154993] SERVICE ALERT: NETIDAD;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460155005] SERVICE ALERT: SP2010DB;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1460155093] SERVICE ALERT: NETIDAD;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1460155102] SERVICE ALERT: ADRODC-DET;Service: DFS Namespace;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1460155102] SERVICE ALERT: ADRODC-DET;Service: User Profile;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1460155241] SERVICE ALERT: ADRODC-DET;Service: DFS Namespace;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460155242] SERVICE ALERT: ADRODC-DET;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460155341] SERVICE ALERT: ADRODC-DET;Service: DFS Namespace;OK;SOFT;3;Dfs: Started
[1460155342] SERVICE ALERT: ADRODC-DET;Service: User Profile;OK;SOFT;3;ProfSvc: Started
[1460155350] SERVICE ALERT: ADSECURED01;Service: Group Policy;CRITICAL;SOFT;1;gpsvc: Stopped
[1460155473] SERVICE ALERT: ADSECURED01;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460155490] SERVICE ALERT: ADSECURED01;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460155503] HOST ALERT: ADSECURED01;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460155523] SERVICE ALERT: ADSECURED01;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460155553] HOST ALERT: ADSECURED01;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460155571] SERVICE ALERT: ADSECURED01;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460155593] SERVICE ALERT: ADSECURED01;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460155601] HOST ALERT: ADSECURED01;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460155607] SERVICE ALERT: EMUPRINTPS;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1460155610] SERVICE ALERT: ADSECURED01;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460155615] SERVICE ALERT: SNAP;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1460155615] SERVICE ALERT: ADSECURED01;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460155618] SERVICE ALERT: SNAP;CPU Load;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1460155618] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.40)
[1460155623] SERVICE ALERT: ADSECURED01;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 1 minute(s)
[1460155627] SERVICE ALERT: SNAP;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1460155629] HOST ALERT: ADSECURED01;UP;SOFT;4;PING OK - Packet loss = 75%, RTA = 0.44 ms
[1460155671] SERVICE ALERT: ADSECURED01;C:\ Drive Space;OK;SOFT;2;c:\ - total: 278.86 Gb - used: 14.98 Gb (5%) - free 263.88 Gb (95%)
[1460155693] SERVICE ALERT: ADSECURED01;Service: Windows Time;OK;SOFT;3;W32Time: Started
[1460155715] SERVICE ALERT: ADSECURED01;CPU Load;OK;SOFT;2;CPU Load 6% (5 min average)
[1460155721] SERVICE ALERT: SNAP;Service: NetLogin;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1460155724] SERVICE ALERT: SNAP;Memory Usage;CRITICAL;SOFT;2;connect to address 164.76.102.40 and port 12489: No route to host
[1460155724] HOST ALERT: SNAP;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.40)
[1460155726] SERVICE ALERT: EMUPRINTPS;CPU Load;OK;SOFT;2;CPU Load 50% (5 min average)
[1460155730] SERVICE ALERT: SNAP;CPU Load;CRITICAL;SOFT;2;connect to address 164.76.102.40 and port 12489: No route to host
[1460155747] SERVICE ALERT: SNAP;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1460155751] HOST ALERT: SNAP;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1460155791] SERVICE ALERT: DB01-2012;Service: RPC;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1460155838] SERVICE ALERT: SNAP;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1460155843] SERVICE ALERT: SNAP;Memory Usage;OK;SOFT;3;Memory usage: total:16362.77 MB - used: 1228.64 MB (8%) - free: 15134.13 MB (92%)
[1460155846] SERVICE ALERT: SNAP;CPU Load;OK;SOFT;3;CPU Load 12% (5 min average)
[1460155880] SERVICE ALERT: ADAUTHDC1;AD Replication;OK;HARD;3;OK: replication is up and running
[1460155911] SERVICE ALERT: DB01-2012;Service: RPC;OK;SOFT;2;RpcSs: Started
[1460156067] SERVICE ALERT: DB01-2012;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1460156119] SERVICE ALERT: JSS9;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.102.98 and port 12489: Connection refused
[1460156187] SERVICE ALERT: DB01-2012;CPU Load;OK;SOFT;2;CPU Load 13% (5 min average)
[1460156190] SERVICE ALERT: ADSECURED01;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1460156205] SERVICE ALERT: SP2010DB;CPU Load;OK;HARD;3;CPU Load 2% (5 min average)
[1460156209] SERVICE ALERT: ADRODC-LIV;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460156223] SERVICE ALERT: ADRODC-LIV;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460156223] SERVICE ALERT: ADRODC-LIV;Service: DFS Namespace;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460156232] SERVICE ALERT: ADRODC-LIV;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460156239] HOST ALERT: ADRODC-LIV;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460156239] SERVICE ALERT: JSS9;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1460156243] SERVICE ALERT: ADRODC-LIV;Service: Distributed File System Replication;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460156244] SERVICE ALERT: ADRODC-LIV;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460156263] SERVICE ALERT: ADRODC-LIV;Service: Intersite Messaging;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460156273] HOST ALERT: ADRODC-LIV;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460156285] SERVICE ALERT: ADRODC-LIV;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460156296] SERVICE ALERT: DB01-2012;Service: Group Policy;CRITICAL;SOFT;1;gpsvc: Stopped
[1460156303] SERVICE ALERT: ADRODC-LIV;Service: Key Distribution Center;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460156303] SERVICE ALERT: ADRODC-LIV;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460156309] SERVICE ALERT: ADRODC-LIV;Service: Group Policy;CRITICAL;SOFT;2;connect to address 207.74.163.140 and port 12489: Connection refused
[1460156313] SERVICE ALERT: ADRODC-LIV;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460156313] SERVICE ALERT: ADRODC-LIV;Service: DNS Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460156315] HOST ALERT: ADRODC-LIV;UP;SOFT;3;PING WARNING - Packet loss = 90%, RTA = 7.03 ms
[1460156317] SERVICE ALERT: C4SAS-SE;CPU Load;OK;HARD;3;CPU Load 0% (5 min average)
[1460156332] SERVICE ALERT: ADRODC-LIV;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1460156332] SERVICE ALERT: ADRODC-LIV;Service: DFS Namespace;OK;SOFT;2;Dfs: Started
[1460156333] SERVICE ALERT: ADRODC-LIV;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1460156343] SERVICE ALERT: ADRODC-LIV;Service: Distributed File System Replication;OK;SOFT;2;DFSR: Started
[1460156344] SERVICE ALERT: ADRODC-LIV;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1460156345] SERVICE ALERT: DB01-2012;Service: NetLogin;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1460156363] SERVICE ALERT: ADRODC-LIV;Service: Intersite Messaging;OK;SOFT;2;IsmServ: Started
[1460156385] SERVICE ALERT: ADRODC-LIV;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1460156395] SERVICE ALERT: EVENTS-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1460156403] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (46.8% change >= 20.0% threshold)
[1460156403] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1460156403] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1460156413] SERVICE ALERT: ADRODC-LIV;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1460156413] SERVICE ALERT: ADRODC-LIV;Service: DNS Service;OK;SOFT;2;DNS: Started
[1460156416] SERVICE ALERT: DB01-2012;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1460156429] SERVICE ALERT: ADRODC-LIV;Service: Group Policy;OK;SOFT;3;gpsvc: Started
[1460156449] SERVICE ALERT: ADRODC-LIV;C:\ Drive Space;OK;SOFT;2;c:\ - total: 1115.16 Gb - used: 557.52 Gb (50%) - free 557.64 Gb (50%)
[1460156451] SERVICE ALERT: MAPS01;Service: User Profile;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1460156465] SERVICE ALERT: DB01-2012;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1460156495] SERVICE ALERT: NETIDADPROXY;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1460156496] SERVICE ALERT: ADRODC-LIV;Service: Key Distribution Center;OK;SOFT;2;KDC: Started
[1460156515] SERVICE ALERT: EVENTS-2012;CPU Load;OK;SOFT;2;CPU Load 49% (5 min average)
[1460156535] SERVICE ALERT: EVENTS-2012;Service: DNS Client;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1460156571] SERVICE ALERT: MAPS01;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1460156593] SERVICE ALERT: INTLDB;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1460156615] SERVICE ALERT: NETIDADPROXY;CPU Load;OK;SOFT;2;CPU Load 85% (5 min average)
[1460156655] SERVICE ALERT: EVENTS-2012;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1460156713] SERVICE ALERT: INTLDB;CPU Load;OK;SOFT;2;CPU Load 88% (5 min average)
[1460156733] SERVICE ALERT: SP2010;CPU Load;OK;HARD;3;CPU Load 0% (5 min average)
[1460156787] SERVICE ALERT: DB01-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460156832] SERVICE ALERT: NETIDADPROXY;Service: Group Policy;CRITICAL;SOFT;1;gpsvc: Stopped
[1460156871] SERVICE ALERT: NETIDAD;Uptime;CRITICAL;SOFT;1;connect to address 164.76.99.247 and port 12489: Connection refused
[1460156907] SERVICE ALERT: DB01-2012;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460156917] SERVICE ALERT: C4SAS-SE;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460156952] SERVICE ALERT: NETIDADPROXY;Service: Group Policy;CRITICAL;SOFT;2;gpsvc: Stopped
[1460156955] SERVICE ALERT: NETIDAD;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460156982] SERVICE ALERT: NETIDAD;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 2 minute(s)
[1460157027] SERVICE ALERT: DB01-2012;CPU Load;WARNING;HARD;3;CPU Load 93% (5 min average)
[1460157034] SERVICE ALERT: ADAUTHDC3;AD Replication;WARNING;HARD;3;Send-MailMessage : The remote name could not be resolved: 'meritsmtp.emich.edu'
[1460157037] SERVICE ALERT: C4SAS-SE;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1460157072] SERVICE ALERT: NETIDADPROXY;Service: Group Policy;OK;SOFT;3;gpsvc: Started
[1460157075] SERVICE ALERT: NETIDAD;CPU Load;OK;SOFT;2;CPU Load 60% (5 min average)
[1460157157] SERVICE ALERT: C4SAS-SE;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1460157161] SERVICE ALERT: ADRODC-DET;Service: DFS Namespace;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460157162] SERVICE ALERT: ADRODC-DET;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460157175] SERVICE ALERT: ADRODC-DET;Service: Remote Management;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460157183] SERVICE ALERT: ADRODC-DET;Service: Intersite Messaging;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460157191] HOST ALERT: ADRODC-DET;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460157204] SERVICE ALERT: ADRODC-DET;Service: AD Web Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460157207] SERVICE ALERT: ITRF;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460157215] SERVICE ALERT: NETIDADPROXY;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460157215] SERVICE ALERT: ADRODC-DET;Service: Key Distribution Center;CRITICAL;SOFT;1;connect to address 198.109.113.140 and port 12489: Connection refused
[1460157224] SERVICE ALERT: ITRF;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460157226] SERVICE ALERT: ADRODC-DET;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 198.109.113.140 and port 12489: Connection refused
[1460157234] HOST ALERT: ADRODC-DET;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460157237] HOST ALERT: ITRF;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460157243] HOST DOWNTIME ALERT: ADAUTHDC1;STARTED; Host has entered a period of scheduled downtime
[1460157243] HOST DOWNTIME ALERT: ADAUTHDC2;STARTED; Host has entered a period of scheduled downtime
[1460157243] HOST DOWNTIME ALERT: ADAUTHDC3;STARTED; Host has entered a period of scheduled downtime
[1460157243] HOST DOWNTIME ALERT: ADCA1;STARTED; Host has entered a period of scheduled downtime
[1460157243] HOST DOWNTIME ALERT: ADCA2;STARTED; Host has entered a period of scheduled downtime
[1460157243] HOST DOWNTIME ALERT: ADDC1;STARTED; Host has entered a period of scheduled downtime
[1460157243] HOST DOWNTIME ALERT: ADDC6;STARTED; Host has entered a period of scheduled downtime
[1460157243] HOST DOWNTIME ALERT: ADDFS01;STARTED; Host has entered a period of scheduled downtime
[1460157243] HOST DOWNTIME ALERT: ADRODC-DET;STARTED; Host has entered a period of scheduled downtime
[1460157243] HOST DOWNTIME ALERT: ADRODC-LIV;STARTED; Host has entered a period of scheduled downtime
[1460157243] HOST DOWNTIME ALERT: ADRODC-WEMU;STARTED; Host has entered a period of scheduled downtime
[1460157243] HOST DOWNTIME ALERT: ADSECURED01;STARTED; Host has entered a period of scheduled downtime
[1460157243] HOST DOWNTIME ALERT: AX-PROD;STARTED; Host has entered a period of scheduled downtime
[1460157243] HOST DOWNTIME ALERT: AXIOM-2012;STARTED; Host has entered a period of scheduled downtime
[1460157243] HOST DOWNTIME ALERT: AXLICENSE;STARTED; Host has entered a period of scheduled downtime
[1460157243] HOST DOWNTIME ALERT: BOEAPP;STARTED; Host has entered a period of scheduled downtime
[1460157243] HOST DOWNTIME ALERT: C4SAS-SE;STARTED; Host has entered a period of scheduled downtime
[1460157243] HOST DOWNTIME ALERT: CASPER;STARTED; Host has entered a period of scheduled downtime
[1460157243] HOST DOWNTIME ALERT: CSCFILESERV;STARTED; Host has entered a period of scheduled downtime
[1460157243] HOST DOWNTIME ALERT: DB01-2012;STARTED; Host has entered a period of scheduled downtime
[1460157243] HOST DOWNTIME ALERT: EMUFILES;STARTED; Host has entered a period of scheduled downtime
[1460157243] HOST DOWNTIME ALERT: EMUPRINTPS;STARTED; Host has entered a period of scheduled downtime
[1460157243] HOST DOWNTIME ALERT: EVENTS-2012;STARTED; Host has entered a period of scheduled downtime
[1460157243] HOST DOWNTIME ALERT: FITNESSGRAM;STARTED; Host has entered a period of scheduled downtime
[1460157243] HOST DOWNTIME ALERT: GOVGIS;STARTED; Host has entered a period of scheduled downtime
[1460157243] HOST DOWNTIME ALERT: INTLAPP;STARTED; Host has entered a period of scheduled downtime
[1460157243] HOST DOWNTIME ALERT: INTLDB;STARTED; Host has entered a period of scheduled downtime
[1460157243] HOST DOWNTIME ALERT: ITRF;STARTED; Host has entered a period of scheduled downtime
[1460157243] HOST DOWNTIME ALERT: JSS9;STARTED; Host has entered a period of scheduled downtime
[1460157243] HOST DOWNTIME ALERT: MAPS01;STARTED; Host has entered a period of scheduled downtime
[1460157243] HOST DOWNTIME ALERT: NETIDAD;STARTED; Host has entered a period of scheduled downtime
[1460157243] HOST DOWNTIME ALERT: NETIDADPROXY;STARTED; Host has entered a period of scheduled downtime
[1460157243] HOST DOWNTIME ALERT: PLANT-METASYS;STARTED; Host has entered a period of scheduled downtime
[1460157243] HOST DOWNTIME ALERT: R25-BETA2012;STARTED; Host has entered a period of scheduled downtime
[1460157243] HOST DOWNTIME ALERT: R25-EPSILON2012;STARTED; Host has entered a period of scheduled downtime
[1460157243] HOST DOWNTIME ALERT: R25-GAMMA2012;STARTED; Host has entered a period of scheduled downtime
[1460157243] HOST DOWNTIME ALERT: SCCM-PRIM01;STARTED; Host has entered a period of scheduled downtime
[1460157243] HOST DOWNTIME ALERT: SNAP;STARTED; Host has entered a period of scheduled downtime
[1460157243] HOST DOWNTIME ALERT: SP2010;STARTED; Host has entered a period of scheduled downtime
[1460157243] HOST DOWNTIME ALERT: SP2010DB;STARTED; Host has entered a period of scheduled downtime
[1460157243] HOST DOWNTIME ALERT: TARGETX-2012;STARTED; Host has entered a period of scheduled downtime
[1460157243] HOST DOWNTIME ALERT: emu-fndn-fe7;STARTED; Host has entered a period of scheduled downtime
[1460157243] HOST DOWNTIME ALERT: epprogsql2;STARTED; Host has entered a period of scheduled downtime
[1460157274] SERVICE ALERT: ADRODC-DET;Service: MDT Monitor;CRITICAL;SOFT;1;MDT_Monitor: Starting
[1460157275] SERVICE ALERT: ADRODC-DET;Service: Remote Management;OK;SOFT;2;WinRM: Started
[1460157276] SERVICE ALERT: ITRF;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.2.78 and port 12489: Connection refused
[1460157278] HOST ALERT: ADRODC-DET;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 8.62 ms
[1460157281] SERVICE ALERT: ADRODC-DET;Service: DFS Namespace;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460157282] SERVICE ALERT: ADRODC-DET;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460157283] SERVICE ALERT: ADRODC-DET;Service: Intersite Messaging;OK;SOFT;2;IsmServ: Started
[1460157285] HOST ALERT: ITRF;UP;SOFT;2;PING OK - Packet loss = 50%, RTA = 0.28 ms
[1460157303] SERVICE ALERT: ITRF;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460157304] SERVICE ALERT: ADRODC-DET;Service: AD Web Service;OK;SOFT;2;ADWS: Started
[1460157307] SERVICE ALERT: ITRF;Memory Usage;OK;SOFT;2;Memory usage: total:8188.82 MB - used: 1048.79 MB (13%) - free: 7140.04 MB (87%)
[1460157313] SERVICE ALERT: INTLDB;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460157324] SERVICE ALERT: ITRF;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1460157333] SERVICE ALERT: SP2010;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460157335] SERVICE ALERT: NETIDADPROXY;CPU Load;CRITICAL;SOFT;2;CPU Load 97% (5 min average)
[1460157335] SERVICE ALERT: ADRODC-DET;Service: Key Distribution Center;OK;SOFT;2;KDC: Started
[1460157346] SERVICE ALERT: ADRODC-DET;C:\ Drive Space;OK;SOFT;2;c:\ - total: 1115.16 Gb - used: 268.29 Gb (24%) - free 846.87 Gb (76%)
[1460157381] SERVICE ALERT: ADRODC-DET;Service: DFS Namespace;OK;SOFT;3;Dfs: Started
[1460157382] SERVICE ALERT: ADRODC-DET;Service: User Profile;OK;SOFT;3;ProfSvc: Started
[1460157385] SERVICE ALERT: ADRODC-DET;Service: MDT Monitor;OK;SOFT;2;MDT_Monitor: Started
[1460157393] SERVICE ALERT: ITRF;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1460157403] SERVICE ALERT: ITRF;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1460157403] SERVICE ALERT: DB01-2012;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.99.225 and port 12489: Connection refused
[1460157405] SERVICE ALERT: SP2010DB;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460157433] SERVICE ALERT: INTLDB;CPU Load;OK;SOFT;2;CPU Load 68% (5 min average)
[1460157453] SERVICE ALERT: SP2010;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1460157455] SERVICE ALERT: NETIDADPROXY;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1460157462] SERVICE ALERT: NETIDADPROXY;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1460157523] SERVICE ALERT: DB01-2012;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1460157525] SERVICE ALERT: SP2010DB;CPU Load;OK;SOFT;2;CPU Load 4% (5 min average)
[1460157582] SERVICE ALERT: NETIDADPROXY;Service: Windows Time;CRITICAL;SOFT;2;W32Time: Stopped
[1460157627] SERVICE ALERT: DB01-2012;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1460157702] SERVICE ALERT: NETIDADPROXY;Service: Windows Time;CRITICAL;HARD;3;W32Time: Stopped
[1460157713] SERVICE ALERT: epprogsql2;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460157757] SERVICE ALERT: C4SAS-SE;CPU Load;OK;HARD;3;CPU Load 3% (5 min average)
[1460157757] SERVICE FLAPPING ALERT: C4SAS-SE;CPU Load;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1460157811] SERVICE ALERT: EMUFILES;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460157813] SERVICE ALERT: epprogsql2;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1460157931] SERVICE ALERT: EMUFILES;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1460157963] SERVICE ALERT: INTLDB;Service: WMI;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1460158023] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460158023] SERVICE ALERT: SCCM-PRIM01;Service: Firewall;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1460158032] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 28% (5 min average)
[1460158033] SERVICE ALERT: INTLDB;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460158055] SERVICE ALERT: NETIDADPROXY;CPU Load;OK;HARD;3;CPU Load 49% (5 min average)
[1460158056] SERVICE ALERT: INTLAPP;CPU Load;WARNING;SOFT;1;CPU Load 94% (5 min average)
[1460158073] SERVICE ALERT: SCCM-PRIM01;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.102.80 and port 12489: No route to host
[1460158076] HOST ALERT: SCCM-PRIM01;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.80)
[1460158083] SERVICE ALERT: INTLDB;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1460158118] SERVICE ALERT: SCCM-PRIM01;Service: NetLogin;CRITICAL;SOFT;1;connect to address 164.76.102.80 and port 12489: No route to host
[1460158121] HOST ALERT: SCCM-PRIM01;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.80)
[1460158143] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 13% (5 min average)
[1460158146] SERVICE ALERT: SCCM-PRIM01;Service: Firewall;CRITICAL;SOFT;2;connect to address 164.76.102.80 and port 12489: No route to host
[1460158148] SERVICE ALERT: EVENTS-2012;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.99.156 and port 12489: Connection refused
[1460158148] SERVICE ALERT: EVENTS-2012;VMWare Tools;CRITICAL;SOFT;1;connect to address 164.76.99.156 and port 12489: Connection refused
[1460158149] HOST ALERT: SCCM-PRIM01;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.102.80)
[1460158153] SERVICE ALERT: INTLDB;CPU Load;OK;SOFT;2;CPU Load 46% (5 min average)
[1460158172] SERVICE ALERT: SCCM-PRIM01;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.102.80 and port 12489: No route to host
[1460158175] HOST ALERT: SCCM-PRIM01;DOWN;SOFT;4;CRITICAL - Host Unreachable (164.76.102.80)
[1460158176] SERVICE ALERT: INTLAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1460158178] HOST ALERT: EVENTS-2012;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460158182] SERVICE ALERT: SCCM-PRIM01;Service: Group Policy;CRITICAL;SOFT;2;connect to address 164.76.102.80 and port 12489: Connection refused
[1460158235] SERVICE ALERT: SCCM-PRIM01;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1460158239] HOST ALERT: SCCM-PRIM01;UP;SOFT;5;PING OK - Packet loss = 0%, RTA = 0.56 ms
[1460158242] HOST ALERT: EVENTS-2012;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1460158263] SERVICE ALERT: SCCM-PRIM01;Service: Firewall;OK;SOFT;3;MpsSvc: Started
[1460158265] SERVICE ALERT: EVENTS-2012;VMWare Tools;OK;SOFT;2;VMTools: Started
[1460158267] SERVICE ALERT: EVENTS-2012;C:\ Drive Space;OK;SOFT;2;c:\ - total: 82.66 Gb - used: 36.68 Gb (44%) - free 45.98 Gb (56%)
[1460158289] SERVICE ALERT: SCCM-PRIM01;Service: RPC;OK;SOFT;2;RpcSs: Started
[1460158296] SERVICE ALERT: INTLAPP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1460158301] SERVICE ALERT: SCCM-PRIM01;Service: Group Policy;OK;SOFT;3;gpsvc: Started
[1460158315] SERVICE ALERT: EVENTS-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460158435] SERVICE ALERT: EVENTS-2012;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1460158555] SERVICE ALERT: EVENTS-2012;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460158632] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460158664] SERVICE ALERT: NETIDADPROXY;CPU Load;CRITICAL;SOFT;1;connect to address 164.76.99.248 and port 12489: Connection refused
[1460158752] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1460158753] SERVICE ALERT: INTLDB;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1460158775] SERVICE ALERT: NETIDADPROXY;CPU Load;OK;SOFT;2;CPU Load 85% (5 min average)
[1460158827] SERVICE ALERT: DB01-2012;CPU Load;OK;HARD;3;CPU Load 0% (5 min average)
[1460158872] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1460158873] SERVICE ALERT: INTLDB;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1460158982] HOST ALERT: ADDFS01;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460158989] SERVICE ALERT: ADDFS01;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460158993] SERVICE ALERT: INTLDB;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1460159016] SERVICE ALERT: ADDFS01;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159019] HOST ALERT: ADDFS01;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460159076] SERVICE ALERT: ADDFS01;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159106] HOST ALERT: ADDFS01;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460159109] SERVICE ALERT: ADDFS01;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460159123] SERVICE ALERT: ADDFS01;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159136] SERVICE ALERT: ADDFS01;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460159139] HOST ALERT: ADDFS01;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460159155] SERVICE ALERT: EVENTS-2012;CPU Load;OK;HARD;3;CPU Load 14% (5 min average)
[1460159165] SERVICE ALERT: ADDFS01;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159195] HOST ALERT: ADDFS01;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460159196] SERVICE ALERT: ADDFS01;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460159215] SERVICE ALERT: ADDFS01;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159216] SERVICE ALERT: ADWU1;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159216] SERVICE ALERT: ADWU1;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159226] HOST ALERT: ADDFS01;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1460159229] SERVICE ALERT: ADDFS01;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460159243] SERVICE ALERT: ADDFS01;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460159246] HOST ALERT: ADWU1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460159256] SERVICE ALERT: ADDFS01;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460159259] HOST ALERT: ADDFS01;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1460159267] SERVICE ALERT: ADDFS01;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159267] SERVICE ALERT: ADWU1;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159284] HOST ALERT: ADWU1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460159285] SERVICE ALERT: ADDFS01;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460159285] SERVICE ALERT: ADWU1;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159297] HOST ALERT: ADDFS01;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1460159307] SERVICE ALERT: ADDFS01;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159311] SERVICE ALERT: ADDFS01;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159313] SERVICE ALERT: ADWU1;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159315] HOST ALERT: ADWU1;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460159316] SERVICE ALERT: ADDFS01;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460159335] SERVICE ALERT: ADDFS01;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460159336] SERVICE ALERT: ADWU1;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460159336] SERVICE ALERT: ADWU1;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460159337] HOST ALERT: ADDFS01;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1460159353] SERVICE ALERT: ADDFS01;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159358] HOST ALERT: ADWU1;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460159363] SERVICE ALERT: ADDFS01;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460159363] SERVICE ALERT: ADWU1;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159363] SERVICE ALERT: ADDFS01;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159383] HOST ALERT: ADDFS01;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1460159387] SERVICE ALERT: ADDFS01;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460159387] SERVICE ALERT: ADWU1;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460159393] HOST ALERT: ADWU1;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460159402] SERVICE ALERT: ADDFS01;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159402] SERVICE ALERT: ADWU1;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159405] SERVICE ALERT: ADDFS01;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460159405] SERVICE ALERT: ADWU1;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460159427] SERVICE ALERT: ADDFS01;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460159431] SERVICE ALERT: ADDFS01;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460159432] HOST ALERT: ADWU1;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1460159433] SERVICE ALERT: OFFICEKMS;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159433] SERVICE ALERT: ADWU1;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460159450] SERVICE ALERT: ADDFS01;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159451] SERVICE ALERT: ADWU1;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159455] SERVICE ALERT: ADDFS01;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460159456] SERVICE ALERT: ADWU1;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460159456] SERVICE ALERT: ADWU1;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460159463] HOST ALERT: OFFICEKMS;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460159463] HOST ALERT: ADWU1;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1460159471] SERVICE ALERT: OFFICEKMS;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159472] SERVICE ALERT: OFFICEKMS;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159473] SERVICE ALERT: ADDFS01;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460159483] SERVICE ALERT: ADWU1;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460159483] SERVICE ALERT: ADDFS01;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460159496] SERVICE ALERT: INTLAPP;CPU Load;OK;HARD;3;CPU Load 0% (5 min average)
[1460159501] HOST ALERT: OFFICEKMS;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460159502] SERVICE ALERT: ADWU1;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159507] SERVICE ALERT: ADDFS01;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460159507] SERVICE ALERT: ADWU1;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460159513] HOST ALERT: ADWU1;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1460159517] SERVICE ALERT: OFFICEKMS;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159518] SERVICE ALERT: OFFICEKMS;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159522] SERVICE ALERT: ADDFS01;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460159522] SERVICE ALERT: ADWU1;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460159525] SERVICE ALERT: ADWU1;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460159544] SERVICE ALERT: ADDFS01;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159547] SERVICE ALERT: ADDFS01;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460159547] HOST ALERT: OFFICEKMS;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460159547] SERVICE ALERT: ADWU1;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159551] SERVICE ALERT: ADDFS01;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460159552] HOST ALERT: ADWU1;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1460159553] SERVICE ALERT: OFFICEKMS;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460159553] SERVICE ALERT: ADWU1;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460159567] SERVICE ALERT: OFFICEKMS;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159570] SERVICE ALERT: ADDFS01;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460159571] SERVICE ALERT: ADWU1;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460159583] HOST ALERT: OFFICEKMS;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460159583] HOST ALERT: ADWU1;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1460159591] SERVICE ALERT: OFFICEKMS;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460159591] SERVICE ALERT: ADWU1;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159592] SERVICE ALERT: OFFICEKMS;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460159593] SERVICE ALERT: INTLDB;CPU Load;OK;HARD;3;CPU Load 0% (5 min average)
[1460159593] SERVICE ALERT: ADDFS01;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460159603] SERVICE ALERT: ADWU1;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460159603] SERVICE ALERT: ADDFS01;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460159613] SERVICE ALERT: OFFICEKMS;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159621] HOST ALERT: OFFICEKMS;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460159622] SERVICE ALERT: ADWU1;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460159637] SERVICE ALERT: OFFICEKMS;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460159638] SERVICE ALERT: OFFICEKMS;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460159642] SERVICE ALERT: ADDFS01;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460159642] SERVICE ALERT: ADWU1;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460159651] SERVICE ALERT: HDBACKUPS;CPU Load;CRITICAL;SOFT;1;connect to address 164.76.102.104 and port 12489: No route to host
[1460159654] HOST ALERT: HDBACKUPS;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.104)
[1460159664] SERVICE ALERT: ADDFS01;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460159667] HOST ALERT: OFFICEKMS;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1460159667] SERVICE ALERT: ADWU1;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460159669] SERVICE ALERT: HDBACKUPS;F:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.102.104 and port 12489: No route to host
[1460159672] HOST ALERT: HDBACKUPS;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.104)
[1460159673] SERVICE ALERT: OFFICEKMS;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460159678] SERVICE ALERT: ADWU1;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159687] SERVICE ALERT: OFFICEKMS;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460159690] SERVICE ALERT: ADDFS01;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460159691] SERVICE ALERT: ADWU1;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460159697] SERVICE ALERT: OFFICEKMS;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159701] SERVICE ALERT: SP2010;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159703] HOST ALERT: OFFICEKMS;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1460159711] SERVICE ALERT: OFFICEKMS;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460159711] SERVICE ALERT: ADWU1;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460159712] SERVICE ALERT: OFFICEKMS;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460159719] SERVICE ALERT: HDBACKUPS;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.102.104 and port 12489: No route to host
[1460159722] HOST ALERT: HDBACKUPS;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.102.104)
[1460159723] SERVICE ALERT: ADWU1;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159731] HOST ALERT: SP2010;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460159733] SERVICE ALERT: OFFICEKMS;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460159737] SERVICE ALERT: SP2010;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159739] SERVICE ALERT: SP2010;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159741] HOST ALERT: OFFICEKMS;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1460159742] SERVICE ALERT: ADWU1;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460159747] SERVICE ALERT: OFFICEKMS;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159750] SERVICE ALERT: SP2010;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159757] SERVICE ALERT: OFFICEKMS;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460159758] SERVICE ALERT: OFFICEKMS;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460159762] SERVICE ALERT: HDBACKUPS;CPU Load;CRITICAL;SOFT;2;connect to address 164.76.102.104 and port 12489: No route to host
[1460159765] HOST ALERT: HDBACKUPS;DOWN;SOFT;4;CRITICAL - Host Unreachable (164.76.102.104)
[1460159767] HOST ALERT: SP2010;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460159772] SERVICE ALERT: ADWU1;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159777] HOST ALERT: OFFICEKMS;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1460159784] SERVICE ALERT: ADDFS01;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460159787] SERVICE ALERT: ADWU1;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460159789] SERVICE ALERT: HDBACKUPS;F:\ Drive Space;CRITICAL;SOFT;2;connect to address 164.76.102.104 and port 12489: No route to host
[1460159790] SERVICE ALERT: OFFICEKMS;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159792] HOST ALERT: HDBACKUPS;DOWN;SOFT;5;CRITICAL - Host Unreachable (164.76.102.104)
[1460159796] SERVICE ALERT: SP2010;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159798] SERVICE ALERT: ADWU1;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460159806] SERVICE ALERT: HDBACKUPS;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.102.104 and port 12489: No route to host
[1460159807] SERVICE ALERT: OFFICEKMS;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460159809] HOST ALERT: HDBACKUPS;DOWN;SOFT;6;CRITICAL - Host Unreachable (164.76.102.104)
[1460159817] SERVICE ALERT: OFFICEKMS;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460159820] HOST ALERT: OFFICEKMS;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1460159821] SERVICE ALERT: SP2010;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460159826] HOST ALERT: SP2010;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460159831] SERVICE ALERT: ADWU1;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460159836] SERVICE ALERT: OFFICEKMS;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159839] SERVICE ALERT: HDBACKUPS;Memory Usage;CRITICAL;SOFT;2;connect to address 164.76.102.104 and port 12489: No route to host
[1460159842] HOST ALERT: HDBACKUPS;DOWN;SOFT;7;CRITICAL - Host Unreachable (164.76.102.104)
[1460159843] SERVICE ALERT: ADWU1;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460159843] SERVICE ALERT: SP2010DB;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159847] SERVICE ALERT: SP2010;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159853] SERVICE ALERT: OFFICEKMS;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460159857] SERVICE ALERT: SP2010;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460159859] SERVICE ALERT: SP2010;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460159859] SERVICE ALERT: HDBACKUPS;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.102.104 and port 12489: No route to host
[1460159859] SERVICE ALERT: HDBACKUPS;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.102.104 and port 12489: No route to host
[1460159862] HOST ALERT: HDBACKUPS;DOWN;SOFT;8;CRITICAL - Host Unreachable (164.76.102.104)
[1460159867] HOST ALERT: SP2010DB;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460159867] SERVICE ALERT: OFFICEKMS;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460159870] SERVICE ALERT: SP2010;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460159873] SERVICE ALERT: SP2010;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159877] HOST ALERT: SP2010;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460159882] SERVICE ALERT: OFFICEKMS;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159882] SERVICE ALERT: HDBACKUPS;CPU Load;CRITICAL;HARD;3;connect to address 164.76.102.104 and port 12489: No route to host
[1460159885] HOST ALERT: HDBACKUPS;DOWN;SOFT;9;CRITICAL - Host Unreachable (164.76.102.104)
[1460159890] SERVICE ALERT: SP2010DB;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159892] SERVICE ALERT: ADWU1;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460159893] SERVICE ALERT: SP2010;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159908] SERVICE ALERT: HDBACKUPS;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.102.104 and port 12489: No route to host
[1460159908] SERVICE ALERT: HDBACKUPS;F:\ Drive Space;CRITICAL;HARD;3;connect to address 164.76.102.104 and port 12489: No route to host
[1460159910] SERVICE ALERT: OFFICEKMS;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460159911] HOST ALERT: HDBACKUPS;DOWN;HARD;10;CRITICAL - Host Unreachable (164.76.102.104)
[1460159916] SERVICE ALERT: SP2010;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460159918] SERVICE ALERT: ADWU1;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460159920] HOST ALERT: SP2010DB;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460159923] HOST ALERT: SP2010;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460159926] SERVICE ALERT: HDBACKUPS;Service: DNS Client;CRITICAL;SOFT;2;connect to address 164.76.102.104 and port 12489: No route to host
[1460159933] SERVICE ALERT: OFFICEKMS;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159933] SERVICE ALERT: SP2010;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159934] SERVICE ALERT: SP2010DB;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159937] SERVICE ALERT: OFFICEKMS;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460159941] SERVICE ALERT: SP2010;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460159945] SERVICE ALERT: SP2010DB;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159954] SERVICE ALERT: HDBACKUPS;Service: NetLogin;CRITICAL;SOFT;1;connect to address 164.76.102.104 and port 12489: No route to host
[1460159956] SERVICE ALERT: OFFICEKMS;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460159957] SERVICE ALERT: HDBACKUPS;Memory Usage;CRITICAL;HARD;3;connect to address 164.76.102.104 and port 12489: No route to host
[1460159963] HOST ALERT: SP2010;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1460159963] SERVICE ALERT: ADWU1;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460159963] SERVICE ALERT: SP2010DB;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159963] SERVICE ALERT: SP2010DB;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460159964] HOST ALERT: SP2010DB;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460159967] SERVICE ALERT: SP2010;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460159976] SERVICE ALERT: OFFICEKMS;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159977] SERVICE ALERT: SP2010;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460159979] SERVICE ALERT: SP2010;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460159979] SERVICE ALERT: HDBACKUPS;NSClient++ Version;CRITICAL;SOFT;2;connect to address 164.76.102.104 and port 12489: No route to host
[1460159979] SERVICE ALERT: HDBACKUPS;Service: Firewall;CRITICAL;SOFT;2;connect to address 164.76.102.104 and port 12489: No route to host
[1460159986] SERVICE ALERT: SP2010;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159987] SERVICE ALERT: OFFICEKMS;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460159988] SERVICE ALERT: SP2010DB;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460159990] SERVICE ALERT: SP2010;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460159993] SERVICE ALERT: SP2010;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460159997] HOST ALERT: SP2010;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1460160000] SERVICE ALERT: HDBACKUPS;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.102.104 and port 12489: No route to host
[1460160002] SERVICE FLAPPING ALERT: C4SAS-SE;CPU Load;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1460160002] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (39.1% change >= 20.0% threshold)
[1460160002] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1460160002] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1460160002] HOST DOWNTIME ALERT: ADAUTHDC2;STARTED; Host has entered a period of scheduled downtime
[1460160002] HOST DOWNTIME ALERT: ADDC1;STARTED; Host has entered a period of scheduled downtime
[1460160002] HOST DOWNTIME ALERT: AXIOM-2012;STARTED; Host has entered a period of scheduled downtime
[1460160002] HOST DOWNTIME ALERT: MAPS01;STARTED; Host has entered a period of scheduled downtime
[1460160002] HOST DOWNTIME ALERT: NETIDAD;STARTED; Host has entered a period of scheduled downtime
[1460160002] HOST DOWNTIME ALERT: AXLICENSE;STARTED; Host has entered a period of scheduled downtime
[1460160002] HOST DOWNTIME ALERT: NETIDADPROXY;STARTED; Host has entered a period of scheduled downtime
[1460160002] HOST DOWNTIME ALERT: PLANT-METASYS;STARTED; Host has entered a period of scheduled downtime
[1460160002] HOST DOWNTIME ALERT: ADDC6;STARTED; Host has entered a period of scheduled downtime
[1460160002] HOST DOWNTIME ALERT: BOEAPP;STARTED; Host has entered a period of scheduled downtime
[1460160002] HOST DOWNTIME ALERT: R25-BETA2012;STARTED; Host has entered a period of scheduled downtime
[1460160002] HOST DOWNTIME ALERT: R25-EPSILON2012;STARTED; Host has entered a period of scheduled downtime
[1460160002] HOST DOWNTIME ALERT: C4SAS-SE;STARTED; Host has entered a period of scheduled downtime
[1460160002] HOST DOWNTIME ALERT: R25-GAMMA2012;STARTED; Host has entered a period of scheduled downtime
[1460160002] HOST DOWNTIME ALERT: SCCM-PRIM01;STARTED; Host has entered a period of scheduled downtime
[1460160002] HOST DOWNTIME ALERT: ADAUTHDC3;STARTED; Host has entered a period of scheduled downtime
[1460160002] HOST DOWNTIME ALERT: ADDFS01;STARTED; Host has entered a period of scheduled downtime
[1460160002] HOST DOWNTIME ALERT: CASPER;STARTED; Host has entered a period of scheduled downtime
[1460160002] HOST DOWNTIME ALERT: SNAP;STARTED; Host has entered a period of scheduled downtime
[1460160002] HOST DOWNTIME ALERT: SP2010;STARTED; Host has entered a period of scheduled downtime
[1460160002] HOST DOWNTIME ALERT: CSCFILESERV;STARTED; Host has entered a period of scheduled downtime
[1460160002] HOST DOWNTIME ALERT: SP2010DB;STARTED; Host has entered a period of scheduled downtime
[1460160002] HOST DOWNTIME ALERT: TARGETX-2012;STARTED; Host has entered a period of scheduled downtime
[1460160002] HOST DOWNTIME ALERT: ADRODC-DET;STARTED; Host has entered a period of scheduled downtime
[1460160002] HOST DOWNTIME ALERT: DB01-2012;STARTED; Host has entered a period of scheduled downtime
[1460160002] HOST DOWNTIME ALERT: emu-fndn-fe7;STARTED; Host has entered a period of scheduled downtime
[1460160002] HOST DOWNTIME ALERT: epprogsql2;STARTED; Host has entered a period of scheduled downtime
[1460160002] HOST DOWNTIME ALERT: EMUFILES;STARTED; Host has entered a period of scheduled downtime
[1460160002] HOST DOWNTIME ALERT: ADAUTHDC1;STARTED; Host has entered a period of scheduled downtime
[1460160002] HOST DOWNTIME ALERT: ADCA1;STARTED; Host has entered a period of scheduled downtime
[1460160002] HOST DOWNTIME ALERT: ADRODC-LIV;STARTED; Host has entered a period of scheduled downtime
[1460160002] HOST DOWNTIME ALERT: EMUPRINTPS;STARTED; Host has entered a period of scheduled downtime
[1460160002] HOST DOWNTIME ALERT: EVENTS-2012;STARTED; Host has entered a period of scheduled downtime
[1460160002] HOST DOWNTIME ALERT: ADRODC-WEMU;STARTED; Host has entered a period of scheduled downtime
[1460160002] HOST DOWNTIME ALERT: FITNESSGRAM;STARTED; Host has entered a period of scheduled downtime
[1460160002] HOST DOWNTIME ALERT: GOVGIS;STARTED; Host has entered a period of scheduled downtime
[1460160002] HOST DOWNTIME ALERT: ADCA2;STARTED; Host has entered a period of scheduled downtime
[1460160002] HOST DOWNTIME ALERT: ADSECURED01;STARTED; Host has entered a period of scheduled downtime
[1460160002] HOST DOWNTIME ALERT: INTLAPP;STARTED; Host has entered a period of scheduled downtime
[1460160002] HOST DOWNTIME ALERT: INTLDB;STARTED; Host has entered a period of scheduled downtime
[1460160002] HOST DOWNTIME ALERT: AX-PROD;STARTED; Host has entered a period of scheduled downtime
[1460160002] HOST DOWNTIME ALERT: ITRF;STARTED; Host has entered a period of scheduled downtime
[1460160002] HOST DOWNTIME ALERT: JSS9;STARTED; Host has entered a period of scheduled downtime
[1460160028] SERVICE ALERT: HDBACKUPS;Service: Group Policy;CRITICAL;SOFT;2;connect to address 164.76.102.104 and port 12489: No route to host
[1460160029] SERVICE ALERT: SP2010;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460160030] SERVICE ALERT: OFFICEKMS;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460160034] SERVICE ALERT: SP2010DB;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460160035] SERVICE ALERT: ADDFS01;Service: Windows Time;OK;HARD;3;W32Time: Started
[1460160036] SERVICE ALERT: SP2010;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460160038] SERVICE ALERT: ADAUTHDC3;AD Replication;OK;HARD;3;OK: replication is up and running
[1460160039] HOST ALERT: ADDFS01;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1460160044] HOST ALERT: SP2010DB;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460160046] SERVICE ALERT: HDBACKUPS;Service: DNS Client;CRITICAL;HARD;3;connect to address 164.76.102.104 and port 12489: No route to host
[1460160046] SERVICE ALERT: HDBACKUPS;Service: User Profile;CRITICAL;SOFT;1;connect to address 164.76.102.104 and port 12489: No route to host
[1460160048] SERVICE ALERT: ADDFS01;Service: WMI;OK;HARD;3;Winmgmt: Started
[1460160053] SERVICE ALERT: OFFICEKMS;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460160053] SERVICE ALERT: SP2010;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460160054] SERVICE ALERT: SP2010DB;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460160058] SERVICE ALERT: SP2010DB;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460160059] HOST ALERT: SP2010;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460160065] SERVICE ALERT: SP2010DB;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460160074] SERVICE ALERT: HDBACKUPS;Service: NetLogin;CRITICAL;SOFT;2;connect to address 164.76.102.104 and port 12489: No route to host
[1460160075] SERVICE ALERT: SP2010DB;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460160076] SERVICE ALERT: OFFICEKMS;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460160076] SERVICE ALERT: SP2010;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460160081] SERVICE ALERT: OFFICEKMS;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460160083] SERVICE ALERT: SP2010DB;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460160083] SERVICE ALERT: SP2010DB;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460160084] HOST ALERT: SP2010DB;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460160087] SERVICE ALERT: SP2010;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460160087] SERVICE ALERT: ADDFS01;Uptime;OK;HARD;3;System Uptime - 0 day(s) 1 hour(s) 40 minute(s)
[1460160096] SERVICE ALERT: OFFICEKMS;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460160096] SERVICE ALERT: HDBACKUPS;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.102.104 and port 12489: No route to host
[1460160096] SERVICE ALERT: HDBACKUPS;Service: Firewall;CRITICAL;HARD;3;connect to address 164.76.102.104 and port 12489: No route to host
[1460160096] SERVICE ALERT: INTLAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460160097] SERVICE ALERT: SP2010;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460160099] SERVICE ALERT: HDBACKUPS;NSClient++ Version;CRITICAL;HARD;3;connect to address 164.76.102.104 and port 12489: No route to host
[1460160106] SERVICE ALERT: SP2010;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460160106] HOST ALERT: SP2010;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460160108] SERVICE ALERT: SP2010DB;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460160113] SERVICE ALERT: SP2010;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460160115] SERVICE ALERT: ADWU1;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460160119] SERVICE ALERT: SP2010;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460160120] SERVICE ALERT: HDBACKUPS;Service: RPC;CRITICAL;SOFT;2;connect to address 164.76.102.104 and port 12489: No route to host
[1460160123] SERVICE ALERT: SP2010DB;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460160127] SERVICE ALERT: ADDFS01;C:\ Drive Space;OK;HARD;3;c:\ - total: 79.66 Gb - used: 14.40 Gb (18%) - free 65.25 Gb (82%)
[1460160131] SERVICE ALERT: ADDFS01;VMWare Tools;OK;HARD;3;VMTools: Started
[1460160133] SERVICE ALERT: SP2010DB;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460160138] SERVICE ALERT: HDBACKUPS;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.102.104 and port 12489: No route to host
[1460160138] HOST ALERT: SP2010DB;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460160143] HOST ALERT: SP2010;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1460160148] SERVICE ALERT: HDBACKUPS;Service: Group Policy;CRITICAL;HARD;3;connect to address 164.76.102.104 and port 12489: No route to host
[1460160149] SERVICE ALERT: SP2010;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460160154] SERVICE ALERT: SP2010DB;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460160168] SERVICE ALERT: HDBACKUPS;Service: User Profile;CRITICAL;SOFT;2;connect to address 164.76.102.104 and port 12489: No route to host
[1460160173] SERVICE ALERT: OFFICEKMS;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460160173] SERVICE ALERT: SP2010;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460160173] SERVICE ALERT: ADDFS01;CPU Load;OK;HARD;3;CPU Load 2% (5 min average)
[1460160173] SERVICE ALERT: SP2010DB;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460160173] SERVICE ALERT: SP2010;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460160174] SERVICE ALERT: SP2010DB;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460160178] SERVICE ALERT: SP2010DB;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460160179] HOST ALERT: SP2010;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460160183] SERVICE ALERT: ADDFS01;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1460160184] HOST ALERT: SP2010DB;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460160185] SERVICE ALERT: SP2010DB;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460160186] SERVICE ALERT: HDBACKUPS;Uptime;CRITICAL;SOFT;1;connect to address 164.76.102.104 and port 12489: No route to host
[1460160186] SERVICE ALERT: HDBACKUPS;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.102.104 and port 12489: No route to host
[1460160194] SERVICE ALERT: HDBACKUPS;Service: NetLogin;CRITICAL;HARD;3;connect to address 164.76.102.104 and port 12489: No route to host
[1460160195] SERVICE ALERT: SP2010DB;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460160196] SERVICE ALERT: SP2010;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460160201] SERVICE ALERT: OFFICEKMS;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460160203] SERVICE ALERT: SP2010DB;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460160215] HOST ALERT: SP2010DB;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1460160216] SERVICE ALERT: OFFICEKMS;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460160216] SERVICE ALERT: INTLAPP;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1460160216] SERVICE ALERT: HDBACKUPS;Service: WMI;CRITICAL;SOFT;2;connect to address 164.76.102.104 and port 12489: No route to host
[1460160217] SERVICE ALERT: SP2010;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460160221] SERVICE ALERT: SP2010DB;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460160222] SERVICE ALERT: ADDFS01;Memory Usage;OK;HARD;3;Memory usage: total:10943.64 MB - used: 872.98 MB (8%) - free: 10070.66 MB (92%)
[1460160226] SERVICE ALERT: SP2010;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460160226] HOST ALERT: SP2010;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1460160228] SERVICE ALERT: SP2010DB;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460160239] SERVICE ALERT: SP2010;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460160240] SERVICE ALERT: HDBACKUPS;Service: RPC;CRITICAL;HARD;3;connect to address 164.76.102.104 and port 12489: No route to host
[1460160243] SERVICE ALERT: SP2010DB;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460160251] HOST ALERT: SP2010DB;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1460160253] SERVICE ALERT: SP2010DB;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460160258] SERVICE ALERT: HDBACKUPS;Service: Windows Time;CRITICAL;SOFT;2;connect to address 164.76.102.104 and port 12489: No route to host
[1460160269] SERVICE ALERT: SP2010;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460160269] HOST ALERT: SP2010;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1460160270] SERVICE ALERT: ADDFS01;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1460160274] SERVICE ALERT: SP2010DB;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460160283] HOST ALERT: SP2010DB;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1460160288] SERVICE ALERT: HDBACKUPS;Service: User Profile;CRITICAL;HARD;3;connect to address 164.76.102.104 and port 12489: No route to host
[1460160293] SERVICE ALERT: SP2010;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460160293] SERVICE ALERT: SP2010DB;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460160306] SERVICE ALERT: HDBACKUPS;C:\ Drive Space;CRITICAL;SOFT;2;connect to address 164.76.102.104 and port 12489: No route to host
[1460160306] SERVICE ALERT: HDBACKUPS;Uptime;CRITICAL;SOFT;2;connect to address 164.76.102.104 and port 12489: No route to host
[1460160315] SERVICE ALERT: SP2010DB;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460160316] SERVICE ALERT: SP2010;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460160323] HOST ALERT: SP2010DB;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1460160323] HOST ALERT: SP2010;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1460160336] SERVICE ALERT: HDBACKUPS;Service: WMI;CRITICAL;HARD;3;connect to address 164.76.102.104 and port 12489: No route to host
[1460160341] SERVICE ALERT: SP2010DB;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460160353] SERVICE ALERT: SP2010DB;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460160359] SERVICE ALERT: SP2010;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460160363] SERVICE ALERT: SP2010DB;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460160364] SERVICE ALERT: ADDFS01;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1460160371] HOST ALERT: SP2010DB;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1460160373] SERVICE ALERT: SP2010DB;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460160378] SERVICE ALERT: HDBACKUPS;Service: Windows Time;CRITICAL;HARD;3;connect to address 164.76.102.104 and port 12489: No route to host
[1460160389] HOST ALERT: SP2010;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1460160398] SERVICE ALERT: SP2010DB;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460160409] SERVICE ALERT: ADDFS01;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1460160413] SERVICE ALERT: SP2010;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460160413] SERVICE ALERT: SP2010DB;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460160426] SERVICE ALERT: HDBACKUPS;C:\ Drive Space;CRITICAL;HARD;3;connect to address 164.76.102.104 and port 12489: No route to host
[1460160426] SERVICE ALERT: HDBACKUPS;Uptime;CRITICAL;HARD;3;connect to address 164.76.102.104 and port 12489: No route to host
[1460160436] SERVICE ALERT: ADDFS01;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1460160443] HOST ALERT: SP2010;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1460160461] SERVICE ALERT: SP2010DB;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460160473] SERVICE ALERT: SP2010DB;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460160496] SERVICE ALERT: ADDFS01;Service: RPC;OK;HARD;3;RpcSs: Started
[1460160518] SERVICE ALERT: SP2010DB;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460160543] SERVICE ALERT: ADDFS01;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1460160593] SERVICE ALERT: SP2010DB;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460160638] SERVICE ALERT: SP2010DB;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460160833] SERVICE ALERT: BOEAPPTEST;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1460160953] SERVICE ALERT: BOEAPPTEST;CPU Load;OK;SOFT;2;CPU Load 7% (5 min average)
[1460161803] SERVICE FLAPPING ALERT: C4SAS-SE;CPU Load;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1460161803] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (32.3% change >= 20.0% threshold)
[1460161803] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1460161803] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1460161803] HOST DOWNTIME ALERT: emu-fndn-fe7;STARTED; Host has entered a period of scheduled downtime
[1460161803] HOST DOWNTIME ALERT: SNAP;STARTED; Host has entered a period of scheduled downtime
[1460161803] HOST DOWNTIME ALERT: MAPS01;STARTED; Host has entered a period of scheduled downtime
[1460161803] HOST DOWNTIME ALERT: AXIOM-2012;STARTED; Host has entered a period of scheduled downtime
[1460161803] HOST DOWNTIME ALERT: AX-PROD;STARTED; Host has entered a period of scheduled downtime
[1460161803] HOST DOWNTIME ALERT: JSS9;STARTED; Host has entered a period of scheduled downtime
[1460161803] HOST DOWNTIME ALERT: ADSECURED01;STARTED; Host has entered a period of scheduled downtime
[1460161803] HOST DOWNTIME ALERT: ADRODC-WEMU;STARTED; Host has entered a period of scheduled downtime
[1460161803] HOST DOWNTIME ALERT: SCCM-PRIM01;STARTED; Host has entered a period of scheduled downtime
[1460161803] HOST DOWNTIME ALERT: ITRF;STARTED; Host has entered a period of scheduled downtime
[1460161803] HOST DOWNTIME ALERT: ADRODC-LIV;STARTED; Host has entered a period of scheduled downtime
[1460161803] HOST DOWNTIME ALERT: ADRODC-DET;STARTED; Host has entered a period of scheduled downtime
[1460161803] HOST DOWNTIME ALERT: INTLDB;STARTED; Host has entered a period of scheduled downtime
[1460161803] HOST DOWNTIME ALERT: ADDFS01;STARTED; Host has entered a period of scheduled downtime
[1460161803] HOST DOWNTIME ALERT: ADDC6;STARTED; Host has entered a period of scheduled downtime
[1460161803] HOST DOWNTIME ALERT: TARGETX-2012;STARTED; Host has entered a period of scheduled downtime
[1460161803] HOST DOWNTIME ALERT: R25-GAMMA2012;STARTED; Host has entered a period of scheduled downtime
[1460161803] HOST DOWNTIME ALERT: INTLAPP;STARTED; Host has entered a period of scheduled downtime
[1460161803] HOST DOWNTIME ALERT: ADDC1;STARTED; Host has entered a period of scheduled downtime
[1460161803] HOST DOWNTIME ALERT: ADCA2;STARTED; Host has entered a period of scheduled downtime
[1460161803] HOST DOWNTIME ALERT: GOVGIS;STARTED; Host has entered a period of scheduled downtime
[1460161803] HOST DOWNTIME ALERT: ADCA1;STARTED; Host has entered a period of scheduled downtime
[1460161803] HOST DOWNTIME ALERT: ADAUTHDC3;STARTED; Host has entered a period of scheduled downtime
[1460161803] HOST DOWNTIME ALERT: R25-EPSILON2012;STARTED; Host has entered a period of scheduled downtime
[1460161803] HOST DOWNTIME ALERT: FITNESSGRAM;STARTED; Host has entered a period of scheduled downtime
[1460161803] HOST DOWNTIME ALERT: ADAUTHDC2;STARTED; Host has entered a period of scheduled downtime
[1460161803] HOST DOWNTIME ALERT: ADAUTHDC1;STARTED; Host has entered a period of scheduled downtime
[1460161803] HOST DOWNTIME ALERT: EVENTS-2012;STARTED; Host has entered a period of scheduled downtime
[1460161803] HOST DOWNTIME ALERT: epprogsql2;STARTED; Host has entered a period of scheduled downtime
[1460161803] HOST DOWNTIME ALERT: SP2010DB;STARTED; Host has entered a period of scheduled downtime
[1460161803] HOST DOWNTIME ALERT: R25-BETA2012;STARTED; Host has entered a period of scheduled downtime
[1460161803] HOST DOWNTIME ALERT: EMUPRINTPS;STARTED; Host has entered a period of scheduled downtime
[1460161803] HOST DOWNTIME ALERT: EMUFILES;STARTED; Host has entered a period of scheduled downtime
[1460161803] HOST DOWNTIME ALERT: PLANT-METASYS;STARTED; Host has entered a period of scheduled downtime
[1460161803] HOST DOWNTIME ALERT: DB01-2012;STARTED; Host has entered a period of scheduled downtime
[1460161803] HOST DOWNTIME ALERT: CSCFILESERV;STARTED; Host has entered a period of scheduled downtime
[1460161803] HOST DOWNTIME ALERT: SP2010;STARTED; Host has entered a period of scheduled downtime
[1460161803] HOST DOWNTIME ALERT: NETIDADPROXY;STARTED; Host has entered a period of scheduled downtime
[1460161803] HOST DOWNTIME ALERT: CASPER;STARTED; Host has entered a period of scheduled downtime
[1460161803] HOST DOWNTIME ALERT: C4SAS-SE;STARTED; Host has entered a period of scheduled downtime
[1460161803] HOST DOWNTIME ALERT: NETIDAD;STARTED; Host has entered a period of scheduled downtime
[1460161803] HOST DOWNTIME ALERT: BOEAPP;STARTED; Host has entered a period of scheduled downtime
[1460161803] HOST DOWNTIME ALERT: AXLICENSE;STARTED; Host has entered a period of scheduled downtime
[1460161810] SERVICE ALERT: OFFICEKMS;Service: RPC;OK;HARD;3;RpcSs: Started
[1460161812] HOST ALERT: CSCFILESERV;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1460161814] HOST ALERT: OFFICEKMS;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1460161816] SERVICE ALERT: SP2010;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1460161818] SERVICE ALERT: SP2010DB;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1460161820] HOST ALERT: SP2010;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.40 ms
[1460161822] HOST ALERT: SP2010DB;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1460161829] HOST ALERT: ADWU1;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.39 ms
[1460161836] SERVICE ALERT: ADWU1;VMWare Tools;OK;HARD;3;VMTools: Started
[1460161836] SERVICE ALERT: ADWU1;C:\ Drive Space;OK;HARD;3;c:\ - total: 419.66 Gb - used: 333.38 Gb (79%) - free 86.28 Gb (21%)
[1460161843] SERVICE ALERT: HDBACKUPS;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1460161843] SERVICE ALERT: CSCFILESERV;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1460161847] HOST ALERT: HDBACKUPS;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1460161856] SERVICE ALERT: OFFICEKMS;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1460161863] SERVICE ALERT: SP2010DB;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1460161867] SERVICE ALERT: SP2010;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1460161887] SERVICE ALERT: ADWU1;CPU Load;OK;HARD;3;CPU Load 0% (5 min average)
[1460161892] SERVICE ALERT: CSCFILESERV;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1460161893] SERVICE ALERT: SP2010;CPU Load;OK;HARD;3;CPU Load 15% (5 min average)
[1460161893] SERVICE FLAPPING ALERT: SP2010;CPU Load;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1460161895] SERVICE ALERT: ADWU1;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 44 minute(s)
[1460161896] SERVICE ALERT: HDBACKUPS;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1460161897] SERVICE ALERT: HDBACKUPS;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1460161905] SERVICE ALERT: ADWU1;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1460161933] SERVICE ALERT: ADWU1;Memory Usage;OK;HARD;3;Memory usage: total:22471.68 MB - used: 3860.47 MB (17%) - free: 18611.21 MB (83%)
[1460161936] SERVICE ALERT: CSCFILESERV;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1460161945] SERVICE ALERT: HDBACKUPS;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1460161953] SERVICE ALERT: OFFICEKMS;Service: Windows Time;OK;HARD;3;W32Time: Started
[1460161953] SERVICE ALERT: SP2010;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1460161954] SERVICE ALERT: SP2010DB;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1460161958] SERVICE ALERT: SP2010DB;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1460161965] SERVICE ALERT: SP2010DB;CPU Load;OK;HARD;3;CPU Load 1% (5 min average)
[1460161965] SERVICE FLAPPING ALERT: SP2010DB;CPU Load;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1460161981] SERVICE ALERT: OFFICEKMS;Service: WMI;OK;HARD;3;Winmgmt: Started
[1460161983] SERVICE ALERT: ADWU1;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1460161983] SERVICE ALERT: SP2010DB;C:\ Drive Space;OK;HARD;3;c:\ - total: 69.90 Gb - used: 54.75 Gb (78%) - free 15.15 Gb (22%)
[1460161985] SERVICE ALERT: CSCFILESERV;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1460161991] SERVICE ALERT: HDBACKUPS;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1460161996] SERVICE ALERT: OFFICEKMS;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 42 minute(s)
[1460161997] SERVICE ALERT: SP2010;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1460162006] SERVICE ALERT: SP2010;Service: RPC;OK;HARD;3;RpcSs: Started
[1460162008] SERVICE ALERT: SP2010DB;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1460162022] SERVICE ALERT: ADWU1;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1460162031] SERVICE ALERT: CSCFILESERV;Service: RPC;OK;HARD;3;RpcSs: Started
[1460162037] SERVICE ALERT: HDBACKUPS;Service: RPC;OK;HARD;3;RpcSs: Started
[1460162049] SERVICE ALERT: SP2010;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1460162053] SERVICE ALERT: OFFICEKMS;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1460162054] SERVICE ALERT: SP2010DB;Service: RPC;OK;HARD;3;RpcSs: Started
[1460162071] SERVICE ALERT: ADWU1;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1460162078] SERVICE ALERT: CSCFILESERV;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1460162085] SERVICE ALERT: HDBACKUPS;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1460162091] SERVICE ALERT: OFFICEKMS;C:\ Drive Space;OK;HARD;3;c:\ - total: 49.66 Gb - used: 13.42 Gb (27%) - free 36.23 Gb (73%)
[1460162092] SERVICE ALERT: OFFICEKMS;CPU Load;OK;HARD;3;CPU Load 0% (5 min average)
[1460162095] SERVICE ALERT: SP2010DB;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1460162096] SERVICE ALERT: SP2010;Service: WMI;OK;HARD;3;Winmgmt: Started
[1460162122] SERVICE ALERT: ADWU1;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1460162126] SERVICE ALERT: CSCFILESERV;Service: WMI;OK;HARD;3;Winmgmt: Started
[1460162133] SERVICE ALERT: HDBACKUPS;Service: WMI;OK;HARD;3;Winmgmt: Started
[1460162137] SERVICE ALERT: OFFICEKMS;Memory Usage;OK;HARD;3;Memory usage: total:10943.64 MB - used: 981.90 MB (9%) - free: 9961.75 MB (91%)
[1460162138] SERVICE ALERT: OFFICEKMS;VMWare Tools;OK;HARD;3;VMTools: Started
[1460162139] SERVICE ALERT: SP2010;Service: Windows Time;OK;HARD;3;W32Time: Started
[1460162143] SERVICE ALERT: SP2010DB;Service: WMI;OK;HARD;3;Winmgmt: Started
[1460162153] SERVICE ALERT: SP2010DB;VMWare Tools;OK;HARD;3;VMTools: Started
[1460162167] SERVICE ALERT: ADWU1;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1460162173] SERVICE ALERT: CSCFILESERV;Service: Windows Time;OK;HARD;3;W32Time: Started
[1460162175] SERVICE ALERT: HDBACKUPS;Service: Windows Time;OK;HARD;3;W32Time: Started
[1460162187] SERVICE ALERT: OFFICEKMS;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1460162193] SERVICE ALERT: SP2010;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 33 minute(s)
[1460162193] SERVICE ALERT: SP2010DB;Service: Windows Time;OK;HARD;3;W32Time: Started
[1460162211] SERVICE ALERT: ADWU1;Service: RPC;OK;HARD;3;RpcSs: Started
[1460162214] SERVICE ALERT: CSCFILESERV;Uptime;OK;HARD;3;System Uptime - 0 day(s) 2 hour(s) 0 minute(s)
[1460162223] SERVICE ALERT: HDBACKUPS;Uptime;OK;HARD;3;System Uptime - 15 day(s) 11 hour(s) 7 minute(s)
[1460162226] SERVICE ALERT: HDBACKUPS;C:\ Drive Space;OK;HARD;3;c:\ - total: 409.26 Gb - used: 141.03 Gb (34%) - free 268.22 Gb (66%)
[1460162233] SERVICE ALERT: OFFICEKMS;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1460162241] SERVICE ALERT: SP2010DB;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 6 minute(s)
[1460162265] SERVICE ALERT: CSCFILESERV;C:\ Drive Space;OK;HARD;3;c:\ - total: 58.50 Gb - used: 24.38 Gb (42%) - free 34.12 Gb (58%)
[1460162279] SERVICE ALERT: HDBACKUPS;CPU Load;OK;HARD;3;CPU Load 0% (5 min average)
[1460162298] SERVICE ALERT: ADWU1;Service: WMI;OK;HARD;3;Winmgmt: Started
[1460162305] SERVICE ALERT: CSCFILESERV;CPU Load;OK;HARD;3;CPU Load 0% (5 min average)
[1460162306] SERVICE ALERT: HDBACKUPS;F:\ Drive Space;OK;HARD;3;f:\ - total: 3315.22 Gb - used: 1984.13 Gb (60%) - free 1331.09 Gb (40%)
[1460162317] SERVICE ALERT: OFFICEKMS;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1460162321] SERVICE ALERT: SP2010;Memory Usage;OK;HARD;3;Memory usage: total:14333.44 MB - used: 1505.80 MB (11%) - free: 12827.64 MB (89%)
[1460162343] SERVICE ALERT: ADWU1;Service: Windows Time;OK;HARD;3;W32Time: Started
[1460162353] SERVICE ALERT: CSCFILESERV;Memory Usage;OK;HARD;3;Memory usage: total:8188.82 MB - used: 871.39 MB (11%) - free: 7317.44 MB (89%)
[1460162356] SERVICE ALERT: HDBACKUPS;Memory Usage;OK;HARD;3;Memory usage: total:37532.83 MB - used: 4229.79 MB (11%) - free: 33303.04 MB (89%)
[1460162357] SERVICE ALERT: SP2010;C:\ Drive Space;OK;HARD;3;c:\ - total: 79.90 Gb - used: 51.47 Gb (64%) - free 28.43 Gb (36%)
[1460162359] SERVICE ALERT: SP2010;VMWare Tools;OK;HARD;3;VMTools: Started
[1460162367] SERVICE ALERT: OFFICEKMS;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1460162370] SERVICE ALERT: SP2010;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1460162373] SERVICE ALERT: SP2010DB;Memory Usage;OK;HARD;3;Memory usage: total:14333.44 MB - used: 1399.36 MB (10%) - free: 12934.08 MB (90%)
[1460162397] SERVICE ALERT: CSCFILESERV;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1460163693] SERVICE ALERT: SP2010;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460163813] SERVICE ALERT: SP2010;CPU Load;WARNING;SOFT;2;CPU Load 94% (5 min average)
[1460163933] SERVICE ALERT: SP2010;CPU Load;OK;SOFT;3;CPU Load 54% (5 min average)
[1460164182] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;SOFT;1;e:\ - total: 499.87 Gb - used: 405.09 Gb (81%) - free 94.78 Gb (19%)
[1460164301] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;SOFT;2;e:\ - total: 499.87 Gb - used: 405.09 Gb (81%) - free 94.78 Gb (19%)
[1460164421] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;HARD;3;e:\ - total: 499.87 Gb - used: 405.09 Gb (81%) - free 94.78 Gb (19%)
[1460164871] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 84% (5 min average)
[1460165021] SERVICE ALERT: AD-RMAD;E:\ Drive Space;OK;HARD;3;e:\ - total: 499.87 Gb - used: 399.23 Gb (80%) - free 100.64 Gb (20%)
[1460165133] SERVICE ALERT: SP2010;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1460165253] SERVICE ALERT: SP2010;CPU Load;OK;SOFT;2;CPU Load 61% (5 min average)
[1460165471] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460165591] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1460165711] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1460167143] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 94% (5 min average)
[1460167263] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 63% (5 min average)
[1460168033] HOST DOWNTIME ALERT: emu-fndn-fe7;STOPPED; Host has exited from a period of scheduled downtime
[1460168033] HOST DOWNTIME ALERT: INTLDB;STOPPED; Host has exited from a period of scheduled downtime
[1460168033] HOST DOWNTIME ALERT: AXIOM-2012;STOPPED; Host has exited from a period of scheduled downtime
[1460168033] HOST DOWNTIME ALERT: AX-PROD;STOPPED; Host has exited from a period of scheduled downtime
[1460168033] HOST DOWNTIME ALERT: SNAP;STOPPED; Host has exited from a period of scheduled downtime
[1460168033] HOST DOWNTIME ALERT: MAPS01;STOPPED; Host has exited from a period of scheduled downtime
[1460168033] HOST DOWNTIME ALERT: JSS9;STOPPED; Host has exited from a period of scheduled downtime
[1460168033] HOST DOWNTIME ALERT: ADRODC-DET;STOPPED; Host has exited from a period of scheduled downtime
[1460168033] HOST DOWNTIME ALERT: ADAUTHDC1;STOPPED; Host has exited from a period of scheduled downtime
[1460168033] HOST DOWNTIME ALERT: ADSECURED01;STOPPED; Host has exited from a period of scheduled downtime
[1460168033] HOST DOWNTIME ALERT: ADRODC-WEMU;STOPPED; Host has exited from a period of scheduled downtime
[1460168033] HOST DOWNTIME ALERT: SCCM-PRIM01;STOPPED; Host has exited from a period of scheduled downtime
[1460168033] HOST DOWNTIME ALERT: ADRODC-LIV;STOPPED; Host has exited from a period of scheduled downtime
[1460168033] HOST DOWNTIME ALERT: ITRF;STOPPED; Host has exited from a period of scheduled downtime
[1460168033] HOST DOWNTIME ALERT: ADDFS01;STOPPED; Host has exited from a period of scheduled downtime
[1460168033] HOST DOWNTIME ALERT: ADCA2;STOPPED; Host has exited from a period of scheduled downtime
[1460168033] HOST DOWNTIME ALERT: TARGETX-2012;STOPPED; Host has exited from a period of scheduled downtime
[1460168033] HOST DOWNTIME ALERT: ADDC6;STOPPED; Host has exited from a period of scheduled downtime
[1460168033] HOST DOWNTIME ALERT: R25-GAMMA2012;STOPPED; Host has exited from a period of scheduled downtime
[1460168033] HOST DOWNTIME ALERT: ADDC1;STOPPED; Host has exited from a period of scheduled downtime
[1460168033] HOST DOWNTIME ALERT: ADAUTHDC2;STOPPED; Host has exited from a period of scheduled downtime
[1460168033] HOST DOWNTIME ALERT: INTLAPP;STOPPED; Host has exited from a period of scheduled downtime
[1460168033] HOST DOWNTIME ALERT: FITNESSGRAM;STOPPED; Host has exited from a period of scheduled downtime
[1460168033] HOST DOWNTIME ALERT: GOVGIS;STOPPED; Host has exited from a period of scheduled downtime
[1460168033] HOST DOWNTIME ALERT: ADAUTHDC3;STOPPED; Host has exited from a period of scheduled downtime
[1460168033] HOST DOWNTIME ALERT: ADCA1;STOPPED; Host has exited from a period of scheduled downtime
[1460168033] HOST DOWNTIME ALERT: R25-EPSILON2012;STOPPED; Host has exited from a period of scheduled downtime
[1460168033] HOST DOWNTIME ALERT: PLANT-METASYS;STOPPED; Host has exited from a period of scheduled downtime
[1460168033] HOST DOWNTIME ALERT: EVENTS-2012;STOPPED; Host has exited from a period of scheduled downtime
[1460168033] HOST DOWNTIME ALERT: SP2010DB;STOPPED; Host has exited from a period of scheduled downtime
[1460168033] HOST DOWNTIME ALERT: epprogsql2;STOPPED; Host has exited from a period of scheduled downtime
[1460168033] HOST DOWNTIME ALERT: R25-BETA2012;STOPPED; Host has exited from a period of scheduled downtime
[1460168033] HOST DOWNTIME ALERT: EMUPRINTPS;STOPPED; Host has exited from a period of scheduled downtime
[1460168033] HOST DOWNTIME ALERT: EMUFILES;STOPPED; Host has exited from a period of scheduled downtime
[1460168033] HOST DOWNTIME ALERT: DB01-2012;STOPPED; Host has exited from a period of scheduled downtime
[1460168033] HOST DOWNTIME ALERT: SP2010;STOPPED; Host has exited from a period of scheduled downtime
[1460168033] HOST DOWNTIME ALERT: CSCFILESERV;STOPPED; Host has exited from a period of scheduled downtime
[1460168033] HOST DOWNTIME ALERT: BOEAPP;STOPPED; Host has exited from a period of scheduled downtime
[1460168033] HOST DOWNTIME ALERT: NETIDADPROXY;STOPPED; Host has exited from a period of scheduled downtime
[1460168033] HOST DOWNTIME ALERT: NETIDAD;STOPPED; Host has exited from a period of scheduled downtime
[1460168033] HOST DOWNTIME ALERT: C4SAS-SE;STOPPED; Host has exited from a period of scheduled downtime
[1460168033] HOST DOWNTIME ALERT: CASPER;STOPPED; Host has exited from a period of scheduled downtime
[1460168033] HOST DOWNTIME ALERT: AXLICENSE;STOPPED; Host has exited from a period of scheduled downtime
[1460168090] SERVICE ALERT: INTLTESTAPP;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460168190] SERVICE ALERT: INTLTESTAPP;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1460169156] SERVICE FLAPPING ALERT: C4SAS-SE;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1460169193] SERVICE ALERT: CBORD;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460169293] SERVICE ALERT: CBORD;Memory Usage;OK;SOFT;2;Memory usage: total:10943.58 MB - used: 3348.18 MB (31%) - free: 7595.40 MB (69%)
[1460170863] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460170983] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 83% (5 min average)
[1460172164] SERVICE FLAPPING ALERT: SP2010DB;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1460172453] SERVICE FLAPPING ALERT: SP2010;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)