[1492488308] SERVICE ALERT: cava2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1492488367] SERVICE ALERT: cava2;CPU Load;OK;SOFT;2;CPU Load 1% (15 min average)
[1492488967] SERVICE ALERT: cava2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1492489027] SERVICE ALERT: cava2;CPU Load;OK;SOFT;2;CPU Load 1% (15 min average)
[1492489801] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1492490739] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1492490799] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1492490859] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1492491230] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;1;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1492491350] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;2;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1492491459] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1492491470] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;HARD;3;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1492492770] SERVICE FLAPPING ALERT: hallesecondsouth;Paper Tray Status;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1492492904] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1492492962] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1492493022] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1492493259] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1492493859] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1492494459] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1492494822] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1492495059] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1492496022] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1492496087] SERVICE ALERT: PROCARE;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1492496147] SERVICE ALERT: PROCARE;CPU Load;OK;SOFT;2;CPU Load 0% (15 min average)
[1492496622] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1492496682] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;2;Warning:  idm1  are more than 120 seconds behind!
[1492496742] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1492496742] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1492497342] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1492497942] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1492498670] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;OK;HARD;3;OK: No entries found
[1492499859] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 54% (15 min average)
[1492500342] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1492501543] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1492502259] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1492502319] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1492503342] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1492504200] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service has exited from a period of scheduled downtime
[1492504202] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1492504202] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (44.4% change >= 20.0% threshold)
[1492504202] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1492504202] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1492504202] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1492504202] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1492504202] HOST DOWNTIME ALERT: INTLAPP;STARTED; Host has entered a period of scheduled downtime
[1492504202] HOST DOWNTIME ALERT: INTLDB;STARTED; Host has entered a period of scheduled downtime
[1492504202] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1492504202] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1492504202] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1492504202] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1492504202] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1492504202] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1492504202] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1492504202] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1492504202] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1492504202] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1492504202] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1492504542] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1492504719] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1492504779] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1492505142] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1492505379] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1492505439] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1492505499] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1492506002] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1492506002] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (49.5% change >= 20.0% threshold)
[1492506002] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1492506002] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1492506002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1492506002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1492506002] HOST DOWNTIME ALERT: INTLAPP;STARTED; Host has entered a period of scheduled downtime
[1492506002] HOST DOWNTIME ALERT: INTLDB;STARTED; Host has entered a period of scheduled downtime
[1492506002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1492506002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1492506002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1492506002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1492506002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1492506002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1492506002] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1492506002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1492506002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1492506002] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1492506002] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1492506699] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 46% (15 min average)
[1492506942] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1492507002] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1492507299] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1492507359] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1492507419] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1492507602] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1492507662] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1492508019] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 34% (15 min average)
[1492508263] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1492508322] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;2;Warning:  idm1  are more than 120 seconds behind!
[1492508382] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;3;OK. All replicas within 120 seconds.
[1492510419] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1492510479] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1492510539] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1492511739] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1492512939] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1492513539] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1492514139] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1492515339] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1492515343] SERVICE ALERT: marshallpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1492515939] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1492516782] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1492518939] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1492519542] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5785 found since last run. oneLevelSearchOps OK. searchOps is too high! 6039 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492519602] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492519636] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1492519696] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1492519755] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1492520202] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8224 found since last run. oneLevelSearchOps OK. searchOps is too high! 8520 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492520262] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492520739] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1492520862] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8661 found since last run. oneLevelSearchOps OK. searchOps is too high! 8870 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492520922] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492521291] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5730 found since last run. oneLevelSearchOps OK. searchOps is too high! 6026 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492521339] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1492521351] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492521522] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5067 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492521555] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1492521582] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492521939] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1492521951] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 13204 found since last run. oneLevelSearchOps OK. searchOps is too high! 13490 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492522012] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492522156] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1492522182] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8475 found since last run. oneLevelSearchOps OK. searchOps is too high! 8660 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492522202] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (54.7% change >= 20.0% threshold)
[1492522202] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1492522202] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1492522202] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1492522202] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1492522202] HOST DOWNTIME ALERT: INTLAPP;STARTED; Host has entered a period of scheduled downtime
[1492522202] HOST DOWNTIME ALERT: INTLDB;STARTED; Host has entered a period of scheduled downtime
[1492522202] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1492522202] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1492522202] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1492522202] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1492522202] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1492522202] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1492522202] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1492522202] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1492522202] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1492522202] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1492522202] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1492522216] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1492522242] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492522276] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1492522539] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1492522548] SERVICE ALERT: porterpk;Device Status;OK;HARD;3;Checking printer
[1492522552] HOST ALERT: porterpk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1492522611] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8354 found since last run. oneLevelSearchOps OK. searchOps is too high! 8657 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492522671] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492522773] SERVICE ALERT: porterpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1492522842] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6895 found since last run. oneLevelSearchOps OK. searchOps is too high! 7018 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492522895] SERVICE ALERT: porterpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 52% - OK!
[1492522902] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492523103] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1492523162] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1492523222] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1492523271] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6374 found since last run. oneLevelSearchOps OK. searchOps is too high! 6665 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492523331] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492523502] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5723 found since last run. oneLevelSearchOps OK. searchOps is too high! 5915 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492523562] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492523739] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1492523931] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9220 found since last run. oneLevelSearchOps OK. searchOps is too high! 9463 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492523991] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492524591] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5766 found since last run. oneLevelSearchOps OK. searchOps is too high! 6046 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492524651] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492524675] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1492524675] SERVICE FLAPPING ALERT: alexanderpk;Paper Tray Status;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1492524762] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5413 found since last run. oneLevelSearchOps OK. searchOps is too high! 5527 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492524822] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492525203] SERVICE ALERT: markjeffersonpk;Device Status;OK;HARD;3;Clearing paper path
[1492525422] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6213 found since last run. oneLevelSearchOps OK. searchOps is too high! 6283 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492525482] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492525622] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;CRITICAL;HARD;3;Black Print Cartridge HP C8543X is at 1% - CRITICAL!
[1492525803] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1492525851] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6018 found since last run. oneLevelSearchOps OK. searchOps is too high! 6296 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492525863] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1492525875] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1492525876] HOST ALERT: hallefirst1;UP;HARD;10;PING OK - Packet loss = 44%, RTA = 8.00 ms
[1492525911] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492525923] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1492525936] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1492525956] HOST ALERT: hallefirst2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1492525996] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1492526020] HOST ALERT: hallefirst2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1492526042] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;MENUS  PAPER HANDLING  CONFIGURE DEVICE  DIAGNOSTICS
[1492526082] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8371 found since last run. oneLevelSearchOps OK. searchOps is too high! 8417 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492526139] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 76% (15 min average)
[1492526142] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492526222] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1492526235] SERVICE ALERT: hallefirst1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1492526463] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at WARNING level!
[1492526511] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5537 found since last run. oneLevelSearchOps OK. searchOps is too high! 5811 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492526571] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492526739] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1492526742] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10837 found since last run. oneLevelSearchOps OK. searchOps is too high! 11181 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492526799] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1492526802] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492526859] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1492527062] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 42% - OK!
[1492527066] HOST ALERT: halleground1;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1492527171] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7226 found since last run. oneLevelSearchOps OK. searchOps is too high! 7510 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492527173] SERVICE ALERT: halleground1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1492527183] SERVICE ALERT: halleground2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1492527187] HOST ALERT: halleground2;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.58 ms
[1492527231] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492527402] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 12325 found since last run. oneLevelSearchOps OK. searchOps is too high! 12673 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492527462] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492527492] SERVICE ALERT: halleground1;Device Status;OK;HARD;3;Processing job from tray 4
[1492527517] HOST ALERT: halleground2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1492527563] HOST ALERT: halleground2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1492527624] HOST ALERT: halleground2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1492527714] HOST ALERT: halleground2;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1492527767] SERVICE ALERT: strongpk;Device Status;OK;HARD;3;Processing job
[1492527771] HOST ALERT: strongpk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.68 ms
[1492527788] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1492527804] HOST ALERT: halleground2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1492527831] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7559 found since last run. oneLevelSearchOps OK. searchOps is too high! 7728 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492527848] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1492527866] HOST ALERT: halleground2;UP;SOFT;6;PING OK - Packet loss = 73%, RTA = 0.50 ms
[1492527891] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492527902] SERVICE ALERT: halleground2;Paper Tray Status;OK;SOFT;3;Printer Paper Is OK! All Trays Still Full
[1492528095] SERVICE ALERT: strongpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 60% - OK!
[1492528098] SERVICE ALERT: strongpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1492528163] HOST ALERT: halleground2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1492528163] HOST FLAPPING ALERT: halleground2;STARTED; Host appears to have started flapping (21.3% change > 20.0% threshold)
[1492528224] HOST ALERT: halleground2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1492528314] HOST ALERT: halleground2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1492528404] HOST ALERT: halleground2;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1492528491] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10594 found since last run. oneLevelSearchOps OK. searchOps is too high! 10736 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492528494] HOST ALERT: halleground2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1492528508] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1492528538] HOST ALERT: halleground2;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1492528551] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492528568] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1492528598] HOST ALERT: halleground2;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1492528628] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1492528658] HOST ALERT: halleground2;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1492528659] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1492528748] HOST ALERT: halleground2;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1492528824] HOST ALERT: halleground2;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1492529128] HOST ALERT: halleground2;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1492529151] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7271 found since last run. oneLevelSearchOps OK. searchOps is too high! 7412 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492529211] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492529223] SERVICE ALERT: halleground2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1492529259] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1492529262] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5104 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492529322] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492529327] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1492529389] SERVICE ALERT: halleground2;Device Status;OK;HARD;3;Ready 164.76.106.52
[1492529811] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8386 found since last run. oneLevelSearchOps OK. searchOps is too high! 8459 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492529859] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1492529871] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492529922] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5370 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492529982] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492530195] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1492530302] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1492530344] HOST FLAPPING ALERT: halleground2;STOPPED; Host appears to have stopped flapping (4.9% change < 5.0% threshold)
[1492530362] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1492530422] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1492530471] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7432 found since last run. oneLevelSearchOps OK. searchOps is too high! 7440 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492530532] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492530582] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5432 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492530642] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492530796] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1492530855] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1492530916] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1492531131] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8734 found since last run. oneLevelSearchOps OK. searchOps is too high! 8740 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492531190] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492531242] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6829 found since last run. oneLevelSearchOps OK. searchOps is too high! 7135 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492531302] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492531403] SERVICE ALERT: CAVA2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1492531462] SERVICE ALERT: CAVA2;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1492531516] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1492531522] SERVICE ALERT: CAVA2;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1492531789] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9664 found since last run. oneLevelSearchOps OK. searchOps is too high! 9668 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492531849] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492531902] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6992 found since last run. oneLevelSearchOps OK. searchOps is too high! 7495 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492531962] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492532115] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1492532123] SERVICE ALERT: CAVA2;CPU Load;OK;HARD;3;CPU Load 78% (15 min average)
[1492532176] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1492532236] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1492532448] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 15046 found since last run. oneLevelSearchOps OK. searchOps is too high! 15054 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492532507] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492532562] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5932 found since last run. oneLevelSearchOps OK. searchOps is too high! 6554 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492532622] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492533003] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (42.0% change >= 20.0% threshold)
[1492533003] SERVICE FLAPPING ALERT: alexanderpk;Paper Tray Status;STARTED; Service appears to have started flapping (40.9% change >= 20.0% threshold)
[1492533003] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1492533003] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1492533003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1492533003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1492533003] HOST DOWNTIME ALERT: INTLAPP;STARTED; Host has entered a period of scheduled downtime
[1492533003] HOST DOWNTIME ALERT: INTLDB;STARTED; Host has entered a period of scheduled downtime
[1492533003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1492533003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1492533003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1492533003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1492533003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1492533003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1492533003] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1492533003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1492533003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1492533003] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1492533003] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1492533070] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 99.90 Gb - used: 89.93 Gb (90%) - free 9.97 Gb (10%)
[1492533107] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 13001 found since last run. oneLevelSearchOps OK. searchOps is too high! 13009 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492533130] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 99.90 Gb - used: 89.93 Gb (90%) - free 9.97 Gb (10%)
[1492533167] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492533190] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;OK;SOFT;3;c:\ - total: 99.90 Gb - used: 89.90 Gb (90%) - free 10.00 Gb (10%)
[1492533222] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8015 found since last run. oneLevelSearchOps OK. searchOps is too high! 8769 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492533282] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492533422] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;PERFORM PRINTER MAINTENANCE 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1492533662] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1492533767] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6097 found since last run. oneLevelSearchOps OK. searchOps is too high! 6102 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492533827] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492533882] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9490 found since last run. oneLevelSearchOps OK. searchOps is too high! 9900 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492533942] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492534442] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1492534503] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1492534542] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8365 found since last run. oneLevelSearchOps OK. searchOps is too high! 8731 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492534562] SERVICE ALERT: hallefirst1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1492534602] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492534659] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1492535027] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6072 found since last run. oneLevelSearchOps OK. searchOps is too high! 6081 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492535087] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492535162] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;Processing job
[1492535202] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9226 found since last run. oneLevelSearchOps OK. searchOps is too high! 9833 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492535259] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1492535262] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492535295] SERVICE ALERT: hallefirstcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1492535354] SERVICE ALERT: hallefirstcolor;Device Status;OK;SOFT;2;Tray 1 PLAIN LETTER To change size or type press OK Tray 1 PLAIN LETTER To accept settings 70 72 65 73 73 20 12
[1492535841] SERVICE ALERT: alexanderpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.142.94. Is the printer off?  Is the network down?
[1492535859] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1492535862] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9013 found since last run. oneLevelSearchOps OK. searchOps is too high! 9661 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492535871] HOST ALERT: alexanderpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1492535922] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492535935] HOST ALERT: alexanderpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.63 ms
[1492535988] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1492536047] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1492536107] SERVICE ALERT: halleground2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1492536287] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6580 found since last run. oneLevelSearchOps OK. searchOps is too high! 6592 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492536347] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492536434] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1492536459] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1492536522] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9513 found since last run. oneLevelSearchOps OK. searchOps is too high! 9921 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492536582] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492536658] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;WARNING;SOFT;1;Cyan Cartridge HP CE261A is at 2% - WARNING!
[1492536718] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;WARNING;SOFT;2;Cyan Cartridge HP CE261A is at 2% - WARNING!
[1492536779] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;WARNING;HARD;3;Cyan Cartridge HP CE261A is at 2% - WARNING!
[1492536947] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9259 found since last run. oneLevelSearchOps OK. searchOps is too high! 9263 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492537007] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492537040] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1492537100] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1492537126] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1492537160] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1492537182] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9912 found since last run. oneLevelSearchOps OK. searchOps is too high! 10290 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492537187] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 94% - OK!
[1492537242] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492537268] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1492537321] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 99% - OK!
[1492537378] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;CRITICAL;HARD;3;Cyan Cartridge HP CE261A is at 1% - CRITICAL!
[1492537607] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8179 found since last run. oneLevelSearchOps OK. searchOps is too high! 8188 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492537659] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1492537667] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492537754] SERVICE ALERT: hallefirst1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1492537842] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8930 found since last run. oneLevelSearchOps OK. searchOps is too high! 9319 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492537902] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492538259] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1492538267] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 13325 found since last run. oneLevelSearchOps OK. searchOps is too high! 13330 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492538327] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492538502] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 12636 found since last run. oneLevelSearchOps OK. searchOps is too high! 13241 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492538562] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492538927] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10025 found since last run. oneLevelSearchOps OK. searchOps is too high! 10037 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492538987] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492539086] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1492539116] HOST ALERT: prayharroldpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1492539146] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1492539162] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8352 found since last run. oneLevelSearchOps OK. searchOps is too high! 8694 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492539176] HOST ALERT: prayharroldpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1492539187] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1492539195] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1492539206] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1492539217] HOST ALERT: prayharroldpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1492539222] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492539247] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1492539255] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1492539277] HOST ALERT: prayharroldpk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1492539307] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1492539315] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1492539337] HOST ALERT: prayharroldpk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1492539427] HOST ALERT: prayharroldpk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1492539457] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1492539517] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! 1 of 3 Trays Still Full
[1492539517] HOST ALERT: prayharroldpk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1492539581] HOST ALERT: prayharroldpk;UP;SOFT;8;PING OK - Packet loss = 0%, RTA = 0.58 ms
[1492539800] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 32% - OK!
[1492539822] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 12910 found since last run. oneLevelSearchOps OK. searchOps is too high! 13510 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492539882] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492539902] SERVICE ALERT: prayharroldpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1492539910] SERVICE ALERT: prayharroldpk;Device Status;OK;HARD;3;Ready
[1492540059] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 91% (15 min average)
[1492540482] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10408 found since last run. oneLevelSearchOps OK. searchOps is too high! 10993 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492540542] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492540659] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1492540719] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1492540779] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1492540787] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7057 found since last run. oneLevelSearchOps OK. searchOps is too high! 7073 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492540847] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492541138] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1492541142] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5629 found since last run. oneLevelSearchOps OK. searchOps is too high! 5916 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492541199] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1492541202] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492541259] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1492541447] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6362 found since last run. oneLevelSearchOps OK. searchOps is too high! 6377 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492541507] SERVICE ALERT: halleground2;Device Status;OK;HARD;3;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1492541507] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492541858] SERVICE ALERT: marshallpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1492541901] SERVICE ALERT: sillpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1492541961] SERVICE ALERT: sillpk;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 6% - OK!
[1492541979] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1492541991] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1492542021] HOST ALERT: halleground2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1492542047] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 91% - OK!
[1492542051] HOST ALERT: halleground2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.58 ms
[1492542107] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7060 found since last run. oneLevelSearchOps OK. searchOps is too high! 7085 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492542113] SERVICE ALERT: halleground2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1492542143] HOST ALERT: halleground2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1492542167] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492542173] SERVICE ALERT: halleground2;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1492542203] HOST ALERT: halleground2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1492542233] SERVICE ALERT: halleground2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1492542263] HOST ALERT: halleground2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1492542353] HOST ALERT: halleground2;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1492542402] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6148 found since last run. oneLevelSearchOps OK. searchOps is too high! 6635 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492542427] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1492542428] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1492542428] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1492542443] HOST ALERT: halleground2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1492542462] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492542482] SERVICE ALERT: hallefirst2;Device Status;OK;SOFT;2;Warming up
[1492542482] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 80% - OK!
[1492542487] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1492542517] HOST ALERT: halleground2;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1492542547] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1492542577] HOST ALERT: halleground2;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1492542579] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1492542651] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1492542667] HOST ALERT: halleground2;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1492542711] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1492542741] HOST ALERT: halleground2;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1492542767] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7258 found since last run. oneLevelSearchOps OK. searchOps is too high! 7280 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492542771] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1492542777] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;OK;HARD;3;Cyan Cartridge HP CE261A is at 100% - OK!
[1492542801] HOST ALERT: halleground2;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1492542827] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492543062] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 15258 found since last run. oneLevelSearchOps OK. searchOps is too high! 15635 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492543122] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492543141] SERVICE ALERT: halleground2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1492543145] HOST ALERT: halleground2;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 1.20 ms
[1492543365] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 91% - OK!
[1492543427] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8286 found since last run. oneLevelSearchOps OK. searchOps is too high! 8305 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492543427] SERVICE ALERT: halleground2;Device Status;OK;HARD;3;Ready 164.76.106.52
[1492543427] SERVICE FLAPPING ALERT: halleground2;Device Status;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1492543487] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492543722] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10499 found since last run. oneLevelSearchOps OK. searchOps is too high! 10897 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492543779] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1492543782] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492543924] SERVICE ALERT: markjeffersonpk;Device Status;OK;HARD;3;MENUS  RETRIEVE JOB  INFORMATION  PAPER HANDLING
[1492544087] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10266 found since last run. oneLevelSearchOps OK. searchOps is too high! 10457 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492544147] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492544379] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1492544382] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6025 found since last run. oneLevelSearchOps OK. searchOps is too high! 6412 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492544442] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492544523] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1492544566] SERVICE ALERT: strongpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1492544583] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1492544625] SERVICE ALERT: strongpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1492544643] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1492544685] SERVICE ALERT: strongpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1492544747] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10013 found since last run. oneLevelSearchOps OK. searchOps is too high! 10168 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492544807] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492545042] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6780 found since last run. oneLevelSearchOps OK. searchOps is too high! 7026 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492545102] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492545407] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8840 found since last run. oneLevelSearchOps OK. searchOps is too high! 8966 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492545419] SERVICE ALERT: hallegroundcolor;Toner Collection Unit HP CE265A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1492545467] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492545474] SERVICE ALERT: hallegroundcolor;Toner Collection Unit HP CE265A;OK;SOFT;2;
[1492545754] SERVICE ALERT: hallegroundcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Image Fuser Kit HP 110V-CE246A, 220V-CE247A'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1492545809] SERVICE ALERT: hallegroundcolor;Image Fuser Kit HP 110V-CE246A;OK;SOFT;2;Image Fuser Kit HP 110V-CE246A, 220V-CE247A is at 42% - OK!
[1492546067] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9074 found since last run. oneLevelSearchOps OK. searchOps is too high! 9192 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492546127] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492546179] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (15 min average)
[1492546302] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5156 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492546360] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492546727] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9346 found since last run. oneLevelSearchOps OK. searchOps is too high! 9501 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492546779] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1492546787] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492546839] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1492546899] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1492547387] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5153 found since last run. oneLevelSearchOps OK. searchOps is too high! 5315 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492547447] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492547686] SERVICE ALERT: strongpk;Device Status;OK;HARD;3;Ready
[1492547833] SERVICE FLAPPING ALERT: alexanderpk;Paper Tray Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1492548099] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1492548160] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6483 found since last run. oneLevelSearchOps OK. searchOps is too high! 6755 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492548220] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492548699] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1492548820] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6267 found since last run. oneLevelSearchOps OK. searchOps is too high! 6479 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492548880] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492549247] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7922 found since last run. oneLevelSearchOps OK. searchOps is too high! 8037 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492549307] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492549480] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7327 found since last run. oneLevelSearchOps OK. searchOps is too high! 7649 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492549540] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492549899] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1492549907] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8061 found since last run. oneLevelSearchOps OK. searchOps is too high! 8155 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492549967] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492550140] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5739 found since last run. oneLevelSearchOps OK. searchOps is too high! 6020 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492550200] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492550499] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 78% (15 min average)
[1492550800] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9329 found since last run. oneLevelSearchOps OK. searchOps is too high! 9565 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492550860] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492551099] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1492551159] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1492551219] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1492551460] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8714 found since last run. oneLevelSearchOps OK. searchOps is too high! 8909 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492551520] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492552120] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5257 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492552180] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492552780] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6127 found since last run. oneLevelSearchOps OK. searchOps is too high! 6315 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492552840] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492553073] SERVICE ALERT: cava2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1492553133] SERVICE ALERT: cava2;CPU Load;OK;SOFT;2;CPU Load 76% (15 min average)
[1492553440] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6878 found since last run. oneLevelSearchOps OK. searchOps is too high! 7112 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492553500] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492553567] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6443 found since last run. oneLevelSearchOps OK. searchOps is too high! 6635 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492553619] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 71% (15 min average)
[1492553627] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492553711] SERVICE ALERT: owen101lab1;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1492553733] SERVICE ALERT: cava2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1492553771] SERVICE ALERT: owen101lab1;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1492553793] SERVICE ALERT: cava2;CPU Load;OK;SOFT;2;CPU Load 5% (15 min average)
[1492553831] SERVICE ALERT: owen101lab1;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1492553837] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1492553897] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1492553957] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1492553966] SERVICE ALERT: hallesecondsouth;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1492553977] SERVICE ALERT: owen101lab2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1492554025] SERVICE ALERT: hallesecondsouth;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1492554037] SERVICE ALERT: owen101lab2;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1492554085] SERVICE ALERT: hallesecondsouth;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1492554093] SERVICE ALERT: owen101lab2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1492554097] SERVICE ALERT: owen101lab2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1492554153] SERVICE ALERT: owen101lab2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1492554169] SERVICE ALERT: owen101lab1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1492554202] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1492554213] SERVICE ALERT: owen101lab2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1492554219] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1492554226] SERVICE FLAPPING ALERT: halleground2;Device Status;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1492554227] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6486 found since last run. oneLevelSearchOps OK. searchOps is too high! 6637 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492554229] SERVICE ALERT: owen101lab1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1492554262] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1492554279] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 90% (15 min average)
[1492554287] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492554289] SERVICE ALERT: owen101lab1;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1492554322] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1492554685] SERVICE ALERT: hallesecondsouth;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1492554701] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5430 found since last run. oneLevelSearchOps OK. searchOps is too high! 5663 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492554760] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492554887] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5949 found since last run. oneLevelSearchOps OK. searchOps is too high! 6109 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492554947] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492555359] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5854 found since last run. oneLevelSearchOps OK. searchOps is too high! 6091 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492555419] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492555479] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1492555539] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1492555547] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5106 found since last run. oneLevelSearchOps OK. searchOps is too high! 5242 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492555599] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1492555607] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492556019] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6057 found since last run. oneLevelSearchOps OK. searchOps is too high! 6150 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492556079] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492556199] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1492556262] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1492556321] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1492556381] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1492556485] SERVICE ALERT: hallesecondsouth;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1492556545] SERVICE ALERT: hallesecondsouth;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1492556570] SERVICE ALERT: halleground1;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1492556605] SERVICE ALERT: hallesecondsouth;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1492556631] SERVICE ALERT: halleground1;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1492556679] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7105 found since last run. oneLevelSearchOps OK. searchOps is too high! 7213 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492556691] SERVICE ALERT: halleground1;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1492556739] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492556799] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1492556807] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5050 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492556867] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492557399] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 63% (15 min average)
[1492557467] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5415 found since last run. oneLevelSearchOps OK. searchOps is too high! 5566 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492557527] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492557581] SERVICE ALERT: hallesecondsouth;Device Status;OK;HARD;3;Lower green tab. Close door. 54 6F 20 65 78 69 74 20 70 72 65 73 73 20 1F 2E
[1492557999] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1492558059] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1492558119] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1492558539] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6424 found since last run. oneLevelSearchOps OK. searchOps is too high! 6557 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492558599] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492558719] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1492559005] SERVICE ALERT: hallesecondsouth;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1492559005] SERVICE FLAPPING ALERT: hallesecondsouth;Paper Tray Status;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1492559170] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 1 Trays Empty
[1492559199] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5082 found since last run. oneLevelSearchOps OK. searchOps is too high! 5173 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492559230] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 1 Trays Empty
[1492559259] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492559291] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 1 Trays Empty
[1492559319] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1492559379] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1492559439] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1492560039] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1492560639] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1492561239] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1492562025] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1492562085] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1492562146] SERVICE ALERT: halleground2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1492562327] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5941 found since last run. oneLevelSearchOps OK. searchOps is too high! 6056 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492562387] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492562988] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5153 found since last run. oneLevelSearchOps OK. searchOps is too high! 5236 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492563039] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1492563047] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492563345] SERVICE ALERT: halleground2;Device Status;OK;HARD;3;Ready 164.76.106.52
[1492563639] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1492563646] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5034 found since last run. oneLevelSearchOps OK. searchOps is too high! 5125 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492563699] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1492563705] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492563759] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 78% (15 min average)
[1492564305] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5192 found since last run. oneLevelSearchOps OK. searchOps is too high! 5291 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492564365] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492565400] HOST DOWNTIME ALERT: IMAGES;STARTED; Host has entered a period of scheduled downtime
[1492565559] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1492565619] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1492565679] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1492566164] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5619 found since last run. oneLevelSearchOps OK. searchOps is too high! 5685 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492566224] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492566279] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 52% (15 min average)
[1492566369] SERVICE ALERT: IMAGES;MDT Monitor;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1492566399] HOST ALERT: IMAGES;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1492566488] HOST ALERT: IMAGES;DOWN;HARD;2;PING CRITICAL - Packet loss = 100%
[1492566508] HOST ALERT: IMAGES;UP;HARD;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1492566824] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5035 found since last run. oneLevelSearchOps OK. searchOps is too high! 5112 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492566884] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492566948] SERVICE ALERT: IMAGES;MDT Monitor;UNKNOWN;HARD;3;Failed to open service MDT_Monitor: 1060: The specified service does not exist as an installed service.
[1492567479] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1492567539] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 80% (15 min average)
[1492568090] SERVICE ALERT: halleground1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1492568139] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1492568199] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1492568259] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1492568259] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6154 found since last run. oneLevelSearchOps OK. searchOps is too high! 6369 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492568319] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492568605] SERVICE FLAPPING ALERT: hallesecondsouth;Paper Tray Status;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1492568859] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 63% (15 min average)
[1492569944] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1492570004] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1492570065] SERVICE ALERT: halleground2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1492571859] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1492571919] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 59% (15 min average)
[1492572519] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1492572579] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1492572600] HOST DOWNTIME ALERT: IMAGES;STOPPED; Host has exited from a period of scheduled downtime
[1492572639] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1492573839] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)