[1483765545] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1483765665] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1483765785] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1483766985] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 46% (15 min average)
[1483767585] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1483767705] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1483767825] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 92% (15 min average)
[1483767852] SERVICE FLAPPING ALERT: ADWSUSCLIENTS;E:\ Drive Space;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1483769025] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1483769145] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1483769265] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1483769536] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1483769654] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1483769774] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1483770465] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1483770465] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1483770600] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1483770974] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1483771574] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1483771665] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1483771785] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1483771905] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1483773974] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1483773974] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1483775174] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1483776123] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10022;CRITICAL;SOFT;1;Windows Server Update Services, eventlog: 1 > critical
[1483776243] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10022;CRITICAL;SOFT;2;Windows Server Update Services, eventlog: 1 > critical
[1483776364] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10022;CRITICAL;HARD;3;Windows Server Update Services(error, 10022, success)[The last catalog synchronization attempt was unsuccessful., ], eventlog: 1 > critical
[1483776705] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 83% (15 min average)
[1483776973] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1483777305] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1483777425] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1483777545] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1483777573] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1483779373] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1483779973] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1483780093] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1483780213] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1483780813] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1483781402] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1483781402] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (51.6% change >= 20.0% threshold)
[1483781402] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1483781402] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1483781402] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1483781402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1483781402] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1483781402] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1483781402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1483781402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1483781402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1483781402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1483781402] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1483781402] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1483781402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1483781402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1483781402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1483781402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1483781402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1483781413] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1483782345] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 70% (15 min average)
[1483783203] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1483783203] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (50.1% change >= 20.0% threshold)
[1483783203] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1483783203] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1483783203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1483783203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1483783203] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1483783203] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1483783203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1483783203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1483783203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1483783203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1483783203] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1483783203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1483783203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1483783203] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1483783203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1483783203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1483783203] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1483783563] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10022;OK;HARD;3;Eventlog check ok
[1483783722] SERVICE ALERT: EPPROGSQL2;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 249.66 Gb - used: 225.20 Gb (90%) - free 24.46 Gb (10%)
[1483783842] SERVICE ALERT: EPPROGSQL2;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 249.66 Gb - used: 225.20 Gb (90%) - free 24.46 Gb (10%)
[1483783962] SERVICE ALERT: EPPROGSQL2;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 249.66 Gb - used: 225.20 Gb (90%) - free 24.46 Gb (10%)
[1483784413] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1483784533] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1483786974] SERVICE ALERT: cava2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1483787094] SERVICE ALERT: cava2;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1483787145] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1483787214] SERVICE ALERT: cava2;CPU Load;OK;SOFT;3;CPU Load 5% (15 min average)
[1483787265] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1483787385] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1483787985] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1483788600] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service has exited from a period of scheduled downtime
[1483791478] SERVICE ALERT: hallegroundcolor;Yellow Cartridge HP CE262A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1483791532] SERVICE ALERT: hallegroundcolor;Yellow Cartridge HP CE262A;OK;SOFT;2;Yellow Cartridge HP CE262A is at 58% - OK!
[1483792933] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1483794585] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1483797963] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10022;CRITICAL;SOFT;1;Windows Server Update Services(error, 10022, success)[The last catalog synchronization attempt was unsuccessful., ], eventlog: 1 > critical
[1483798083] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10022;CRITICAL;SOFT;2;Windows Server Update Services(error, 10022, success)[The last catalog synchronization attempt was unsuccessful., ], eventlog: 1 > critical
[1483798204] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10022;CRITICAL;HARD;3;Windows Server Update Services(error, 10022, success)[The last catalog synchronization attempt was unsuccessful., ], eventlog: 1 > critical
[1483799403] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1483799403] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1483799403] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1483799403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1483799403] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1483799403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1483799403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1483799403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1483799403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1483799403] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1483799403] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1483799403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1483799403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1483799403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1483799403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1483799403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1483805148] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10022;OK;HARD;3;Eventlog check ok
[1483805530] SERVICE ALERT: ADRODC-LIV;Service: AD Web Service;CRITICAL;SOFT;1;connect to address 207.74.163.140 and port 12489: Connection refused
[1483805532] SERVICE ALERT: ADRODC-LIV;Service: Distributed File System Replication;CRITICAL;SOFT;1;connect to address 207.74.163.140 and port 12489: Connection refused
[1483805532] SERVICE ALERT: ADRODC-LIV;Service: DNS Client;CRITICAL;SOFT;1;connect to address 207.74.163.140 and port 12489: Connection refused
[1483805533] SERVICE ALERT: ADRODC-LIV;NSClient++ Version;CRITICAL;SOFT;1;connect to address 207.74.163.140 and port 12489: Connection refused
[1483805650] SERVICE ALERT: ADRODC-LIV;Service: AD Web Service;OK;SOFT;2;ADWS: Started
[1483805650] SERVICE ALERT: ADRODC-LIV;Service: Distributed File System Replication;OK;SOFT;2;DFSR: Started
[1483805650] SERVICE ALERT: ADRODC-LIV;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1483805653] SERVICE ALERT: ADRODC-LIV;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1483810203] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1483810203] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1483810203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1483810203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1483810203] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1483810203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1483810203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1483810203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1483810203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1483810203] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1483810203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1483810203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1483810203] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1483810203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1483810203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1483810203] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1483810295] SERVICE ALERT: epprogsql2;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 249.66 Gb - used: 225.23 Gb (90%) - free 24.43 Gb (10%)
[1483810412] SERVICE ALERT: epprogsql2;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 249.66 Gb - used: 225.23 Gb (90%) - free 24.43 Gb (10%)
[1483810532] SERVICE ALERT: epprogsql2;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 249.66 Gb - used: 225.23 Gb (90%) - free 24.43 Gb (10%)
[1483811063] SERVICE ALERT: cava2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1483811183] SERVICE ALERT: cava2;CPU Load;OK;SOFT;2;CPU Load 2% (15 min average)
[1483819547] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10022;CRITICAL;SOFT;1;Windows Server Update Services, eventlog: 1 > critical
[1483819667] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10022;CRITICAL;SOFT;2;Windows Server Update Services, eventlog: 1 > critical
[1483819788] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10022;CRITICAL;HARD;3;Windows Server Update Services, eventlog: 1 > critical
[1483820413] SERVICE ALERT: hallesouthcommons;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1483821583] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1483821703] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 33% (15 min average)
[1483822279] HOST ALERT: crossroadspk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1483822369] HOST ALERT: crossroadspk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1483822447] SERVICE ALERT: crossroadspk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1483822459] HOST ALERT: crossroadspk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1483822507] SERVICE ALERT: crossroadspk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1483822533] SERVICE ALERT: halleground3;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1483822537] HOST ALERT: crossroadspk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1483822567] SERVICE ALERT: crossroadspk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1483822593] SERVICE ALERT: halleground3;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1483822597] HOST ALERT: crossroadspk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1483822653] SERVICE ALERT: halleground3;Device Status;OK;SOFT;3;Move left accessory, if attached.
[1483822663] SERVICE ALERT: crossroadspk;Black Cartridge HP CC364X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1483822687] HOST ALERT: crossroadspk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1483822723] SERVICE ALERT: crossroadspk;Black Cartridge HP CC364X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1483822753] HOST ALERT: crossroadspk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1483822780] SERVICE ALERT: crossroadspk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1483822783] SERVICE ALERT: crossroadspk;Black Cartridge HP CC364X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1483822810] HOST ALERT: crossroadspk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1483822840] SERVICE ALERT: crossroadspk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1483822870] HOST ALERT: crossroadspk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1483822900] SERVICE ALERT: crossroadspk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1483822930] HOST ALERT: crossroadspk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1483824094] SERVICE ALERT: crossroadspk;Device Status;OK;HARD;3;Ready
[1483824098] HOST ALERT: crossroadspk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1483824362] SERVICE ALERT: crossroadspk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1483824577] SERVICE ALERT: crossroadspk;Black Cartridge HP CC364X;OK;HARD;3;Black Cartridge HP CC364X is at 95% - OK!
[1483826986] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10022;OK;HARD;3;Eventlog check ok
[1483832076] SERVICE ALERT: eateriespk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1483834303] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1483834423] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1483834543] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 85% (15 min average)
[1483836343] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1483836463] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 49% (15 min average)
[1483837663] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1483837783] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 23% (15 min average)
[1483841386] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10022;CRITICAL;SOFT;1;Windows Server Update Services(error, 10022, success)[The last catalog synchronization attempt was unsuccessful., ], eventlog: 1 > critical
[1483841506] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10022;CRITICAL;SOFT;2;Windows Server Update Services(error, 10022, success)[The last catalog synchronization attempt was unsuccessful., ], eventlog: 1 > critical
[1483841627] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10022;CRITICAL;HARD;3;Windows Server Update Services(error, 10022, success)[The last catalog synchronization attempt was unsuccessful., ], eventlog: 1 > critical
[1483841648] SERVICE ALERT: ADWSUSCLIENTS;E:\ Drive Space;CRITICAL;HARD;3;e:\ - total: 399.87 Gb - used: 361.11 Gb (90%) - free 38.77 Gb (10%)
[1483841945] SERVICE ALERT: TARGETX-2012;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1483842065] SERVICE ALERT: TARGETX-2012;CPU Load;OK;SOFT;2;CPU Load 86% (15 min average)
[1483844493] SERVICE ALERT: PROCARE;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1483844613] SERVICE ALERT: PROCARE;CPU Load;OK;SOFT;2;CPU Load 72% (15 min average)
[1483844983] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1483845103] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1483845187] SERVICE ALERT: ADDC3;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1483845223] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1483845287] SERVICE ALERT: ADDC3;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1483845823] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (15 min average)
[1483847648] SERVICE ALERT: ADWSUSCLIENTS;E:\ Drive Space;WARNING;HARD;3;e:\ - total: 399.87 Gb - used: 356.95 Gb (89%) - free 42.92 Gb (11%)
[1483848226] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10022;OK;HARD;3;Eventlog check ok
[1483850023] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1483850143] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1483850263] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 86% (15 min average)
[1483850863] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1483850983] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1483851103] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 72% (15 min average)