[1457241463] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 57% (5 min average)
[1457242063] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457242183] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1457242303] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1457253703] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1457254303] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 79% (5 min average)
[1457254903] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457255023] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1457255143] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1457256343] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 71% (5 min average)
[1457257543] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1457257663] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1457257783] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 80% (5 min average)
[1457258983] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457259103] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1457259223] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 96% (5 min average)
[1457259223] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1457259823] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 94% (5 min average)
[1457260423] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1457261023] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 66% (5 min average)
[1457262223] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457262343] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1457262463] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1457264263] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (5 min average)
[1457264863] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457264983] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1457265103] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1457265653] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;1;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1457265703] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 56% (5 min average)
[1457265773] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;2;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1457265893] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;HARD;3;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1457266303] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457266423] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1457266543] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1457266667] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;SOFT;1;WARNING: Replication with ADDC5 failed!
[1457266787] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;SOFT;2;WARNING: Replication with ADDC5 failed!
[1457266907] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;HARD;3;WARNING: Replication with ADDC5 failed!
[1457267142] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 86% (5 min average)
[1457267742] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457267861] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1457267980] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 96% (5 min average)
[1457270377] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 24% (5 min average)
[1457272493] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;OK;HARD;3;OK: No entries found
[1457273977] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457274097] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1457274217] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1457274817] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 68% (5 min average)
[1457275417] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1457275537] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 70% (5 min average)
[1457276137] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457276257] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1457276377] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 44% (5 min average)
[1457277577] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1457277697] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 78% (5 min average)
[1457277711] SERVICE ALERT: ADRODC-LIV;AD Replication;OK;HARD;3;OK: replication is up and running
[1457278297] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457278417] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1457278537] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1457279137] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 83% (5 min average)
[1457279737] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457279857] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1457279977] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 74% (5 min average)
[1457280577] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457280697] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 97% (5 min average)
[1457280817] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 97% (5 min average)
[1457282017] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 58% (5 min average)
[1457283217] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1457283337] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1457283457] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1457285257] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 62% (5 min average)
[1457287657] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457287777] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 96% (5 min average)
[1457287897] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1457293297] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 92% (5 min average)
[1457293897] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1457295097] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 76% (5 min average)
[1457295697] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1457295817] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1457295937] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1457296537] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 60% (5 min average)
[1457297137] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457297257] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1457297377] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 89% (5 min average)
[1457297977] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1457298097] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1457298217] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1457298817] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 97% (5 min average)
[1457299417] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (5 min average)
[1457300017] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1457300137] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1457300257] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 96% (5 min average)
[1457302057] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 94% (5 min average)
[1457302657] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1457304457] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 78% (5 min average)
[1457305057] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457305177] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1457305297] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1457306098] SERVICE ALERT: hallegroundcolor;Cyan Cartridge HP CE261A;CRITICAL;SOFT;1;Cyan Cartridge HP CE261A is at 1% - CRITICAL!
[1457306157] SERVICE ALERT: hallegroundcolor;Cyan Cartridge HP CE261A;CRITICAL;SOFT;2;Cyan Cartridge HP CE261A is at 1% - CRITICAL!
[1457306217] SERVICE ALERT: hallegroundcolor;Cyan Cartridge HP CE261A;CRITICAL;HARD;3;Cyan Cartridge HP CE261A is at 1% - CRITICAL!
[1457310107] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;SOFT;1;WARNING: Replication with ADDC5 failed!
[1457310227] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;SOFT;2;WARNING: Replication with ADDC5 failed!
[1457310347] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;HARD;3;WARNING: Replication with ADDC5 failed!
[1457314296] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 90% (5 min average)
[1457314896] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 97% (5 min average)
[1457320552] SERVICE ALERT: ADRODC-LIV;AD Replication;OK;HARD;3;OK: replication is up and running
[1457320896] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 82% (5 min average)
[1457322096] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457322216] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 95% (5 min average)
[1457322336] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)