[1476763686] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1476763806] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1476763926] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1476765003] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 110 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476765123] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476765126] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1476766326] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1476766446] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 88% (15 min average)
[1476767646] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1476767766] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1476767886] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1476768158] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1476768278] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1476768399] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1476768486] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1476768600] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1476769086] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1476769597] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1476769686] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1476769806] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 91% (15 min average)
[1476769923] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 102 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476770043] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476771006] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1476771126] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1476771246] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1476771997] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1476774530] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1476774650] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1476774770] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1476775446] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 91% (15 min average)
[1476776046] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1476776166] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1476776286] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 79% (15 min average)
[1476776643] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 106 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476776763] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476776797] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1476776797] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1476776886] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1476777006] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1476777126] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1476777170] SERVICE ALERT: epprogsql2;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1476777770] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1476778326] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 44% (15 min average)
[1476778970] SERVICE ALERT: epprogsql2;CPU Load;OK;HARD;3;CPU Load 33% (15 min average)
[1476778970] SERVICE FLAPPING ALERT: epprogsql2;CPU Load;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1476779526] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1476779646] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 65% (15 min average)
[1476780363] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 108 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476780483] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476783997] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1476785646] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1476785766] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 85% (15 min average)
[1476786366] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1476786486] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1476786600] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service has exited from a period of scheduled downtime
[1476786606] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1476788404] SERVICE ALERT: ADCA2;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1476788524] SERVICE ALERT: ADCA2;CPU Load;OK;SOFT;2;CPU Load 0% (15 min average)
[1476788883] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1247 found since last run.
[1476789003] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476789005] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1476789123] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1565 found since last run.
[1476789243] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476789603] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1552 found since last run.
[1476789605] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1476789723] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476789725] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1476789769] SERVICE FLAPPING ALERT: epprogsql2;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1476789843] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1569 found since last run.
[1476789845] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1476789963] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476790323] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1545 found since last run.
[1476790443] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476790563] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1545 found since last run.
[1476790683] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476790845] SERVICE ALERT: marshallpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1476791043] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1580 found since last run.
[1476791163] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476791283] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1578 found since last run.
[1476791403] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476791763] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1757 found since last run.
[1476791883] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476792003] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1835 found since last run.
[1476792123] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476792244] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 91% (15 min average)
[1476792844] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1476792964] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1476793084] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 94% (15 min average)
[1476793303] SERVICE ALERT: hallefirstcolor;Device Status;OK;HARD;3;Processing job from tray 1
[1476793307] HOST ALERT: hallefirstcolor;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.41 ms
[1476793408] SERVICE ALERT: hallefirstcolor;Image Fuser Kit HP 110V-CE246A;OK;HARD;3;Image Fuser Kit HP 110V-CE246A, 220V-CE247A is at 100% - OK!
[1476793473] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;OK;HARD;3;Magenta Cartridge HP CE263A is at 75% - OK!
[1476793647] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;OK;HARD;3;Image Transfer Kit HP CE249A is at 76% - OK!
[1476793658] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1476793663] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;OK;HARD;3;Black Cartridge HP CE260X is at 43% - OK!
[1476793684] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1476793713] SERVICE ALERT: hallefirstcolor;Toner Collection Unit HP CE265A;OK;HARD;3;
[1476793743] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;OK;HARD;3;Yellow Cartridge HP CE262A is at 93% - OK!
[1476793777] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;OK;HARD;3;Cyan Cartridge HP CE261A is at 76% - OK!
[1476793804] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 54% (15 min average)
[1476794134] SERVICE ALERT: hallesecondsouth;Device Status;OK;HARD;3;Genuine HP cartridge installed
[1476794138] HOST ALERT: hallesecondsouth;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.49 ms
[1476794507] SERVICE ALERT: hallesecondsouth;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1476794523] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5011 found since last run. oneLevelSearchOps OK. searchOps is too high! 5252 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476794643] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476794663] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 9% - OK!
[1476795604] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1476795724] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1476795844] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 84% (15 min average)
[1476796083] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6638 found since last run. oneLevelSearchOps OK. searchOps is too high! 6897 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476796203] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476796443] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5783 found since last run. oneLevelSearchOps OK. searchOps is too high! 6128 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476796563] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476797044] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1476797163] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 12205 found since last run. oneLevelSearchOps OK. searchOps is too high! 12685 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476797164] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1476797283] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476797284] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1476797764] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1476797823] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1476797883] SERVICE ALERT: crossroadspk;Device Status;OK;SOFT;3;Open top cover.
[1476797883] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10693 found since last run. oneLevelSearchOps OK. searchOps is too high! 11175 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476797884] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (15 min average)
[1476798003] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7669 found since last run. oneLevelSearchOps OK. searchOps is too high! 7747 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476798003] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476798041] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1476798049] SERVICE ALERT: owen101labcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1476798059] SERVICE ALERT: owen101labcolor;Yellow Cartridge HP CE262A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1476798062] HOST ALERT: owen101labcolor;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1476798069] SERVICE ALERT: owen101labcolor;Black Cartridge HP CE260X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1476798071] HOST ALERT: owen101lab3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1476798089] SERVICE ALERT: owen101labcolor;Image Transfer Kit HP CE249A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1476798099] HOST ALERT: owen101labcolor;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1476798101] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1476798109] SERVICE ALERT: owen101labcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1476798119] SERVICE ALERT: owen101labcolor;Yellow Cartridge HP CE262A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1476798123] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476798129] SERVICE ALERT: owen101labcolor;Black Cartridge HP CE260X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1476798129] SERVICE ALERT: owen101labcolor;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1476798131] HOST ALERT: owen101lab3;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1476798139] HOST ALERT: owen101labcolor;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1476798149] SERVICE ALERT: owen101labcolor;Image Transfer Kit HP CE249A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1476798157] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1476798161] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1476798169] SERVICE ALERT: owen101labcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1476798169] SERVICE ALERT: owen101labcolor;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1476798178] SERVICE ALERT: owen101lab1;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1476798179] HOST ALERT: owen101labcolor;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1476798179] HOST ALERT: owen101lab2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1476798179] SERVICE ALERT: owen101labcolor;Yellow Cartridge HP CE262A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1476798184] HOST ALERT: owen101lab1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1476798187] HOST ALERT: owen101lab3;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1476798189] SERVICE ALERT: owen101labcolor;Black Cartridge HP CE260X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1476798189] SERVICE ALERT: owen101labcolor;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1476798209] SERVICE ALERT: owen101labcolor;Image Transfer Kit HP CE249A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1476798209] HOST ALERT: owen101labcolor;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1476798217] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1476798219] SERVICE ALERT: owen101lab1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1476798220] SERVICE ALERT: owen101labcolor;Toner Collection Unit HP CE265A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1476798229] SERVICE ALERT: owen101labcolor;Magenta Cartridge HP CE263A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1476798229] SERVICE ALERT: owen101labcolor;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1476798238] SERVICE ALERT: owen101lab1;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1476798247] HOST ALERT: owen101lab3;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1476798249] HOST ALERT: owen101lab1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1476798249] SERVICE ALERT: owen101labcolor;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1476798250] HOST ALERT: owen101labcolor;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1476798269] HOST ALERT: owen101lab2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1476798277] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1476798279] SERVICE ALERT: owen101lab1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1476798280] SERVICE ALERT: owen101labcolor;Toner Collection Unit HP CE265A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1476798289] SERVICE ALERT: owen101labcolor;Magenta Cartridge HP CE263A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1476798289] SERVICE ALERT: owen101labcolor;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1476798298] SERVICE ALERT: owen101lab1;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1476798299] SERVICE ALERT: owen101lab3;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1476798307] HOST ALERT: owen101lab3;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1476798309] SERVICE ALERT: owen101lab2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1476798309] HOST ALERT: owen101lab1;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1476798310] HOST ALERT: owen101labcolor;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1476798339] HOST ALERT: owen101lab2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1476798339] SERVICE ALERT: owen101lab1;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1476798340] SERVICE ALERT: owen101labcolor;Toner Collection Unit HP CE265A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1476798349] SERVICE ALERT: owen101labcolor;Magenta Cartridge HP CE263A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1476798355] SERVICE ALERT: owen101labcolor;Cyan Cartridge HP CE261A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1476798359] SERVICE ALERT: owen101lab3;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1476798369] SERVICE ALERT: owen101lab2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1476798369] HOST ALERT: owen101lab1;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1476798370] HOST ALERT: owen101labcolor;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1476798389] HOST ALERT: owen101lab3;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1476798399] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1476798399] HOST ALERT: owen101lab2;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1476798399] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1476798415] SERVICE ALERT: owen101labcolor;Cyan Cartridge HP CE261A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1476798419] SERVICE ALERT: owen101lab3;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1476798429] HOST ALERT: owen101lab1;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1476798429] SERVICE ALERT: owen101lab2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1476798429] HOST ALERT: owen101lab2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1476798439] SERVICE ALERT: owen101lab2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1476798445] HOST ALERT: owen101labcolor;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1476798449] HOST ALERT: owen101lab3;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1476798459] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1476798459] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1476798469] HOST ALERT: owen101lab2;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1476798475] SERVICE ALERT: owen101labcolor;Cyan Cartridge HP CE261A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1476798489] HOST ALERT: owen101lab1;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1476798499] SERVICE ALERT: owen101lab2;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1476798505] HOST ALERT: owen101labcolor;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1476798519] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1476798519] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1476798529] HOST ALERT: owen101lab2;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1476798539] HOST ALERT: owen101lab3;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1476798549] HOST ALERT: owen101lab1;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1476798559] SERVICE ALERT: owen101lab2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1476798589] HOST ALERT: owen101lab2;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1476798603] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5222 found since last run. oneLevelSearchOps OK. searchOps is too high! 5710 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476798629] HOST ALERT: owen101lab3;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1476798639] HOST ALERT: owen101lab1;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1476798679] HOST ALERT: owen101lab2;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1476798719] HOST ALERT: owen101lab3;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1476798723] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476798729] HOST ALERT: owen101lab1;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1476798769] HOST ALERT: owen101lab2;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1476798819] HOST ALERT: owen101lab1;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1476799083] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1476799143] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1476799193] SERVICE ALERT: eateriespk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1476799203] SERVICE ALERT: crossroadspk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1476799253] SERVICE ALERT: eateriespk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1476799313] SERVICE ALERT: eateriespk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1476799803] SERVICE ALERT: crossroadspk;Device Status;OK;HARD;3;Press OK to continue.
[1476800284] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1476800404] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1476800524] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1476801123] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6784 found since last run. oneLevelSearchOps OK. searchOps is too high! 7069 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476801124] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 83% (15 min average)
[1476801243] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476801723] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5770 found since last run. oneLevelSearchOps OK. searchOps is too high! 5883 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476801843] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476801843] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8613 found since last run. oneLevelSearchOps OK. searchOps is too high! 9046 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476801939] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1476801963] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476801999] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1476802059] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1476802443] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 13527 found since last run. oneLevelSearchOps OK. searchOps is too high! 13591 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476802563] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476802563] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10269 found since last run. oneLevelSearchOps OK. searchOps is too high! 10741 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476802683] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476802924] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1476803044] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1476803163] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7816 found since last run. oneLevelSearchOps OK. searchOps is too high! 7926 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476803164] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 92% (15 min average)
[1476803283] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476803283] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6976 found since last run. oneLevelSearchOps OK. searchOps is too high! 7400 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476803404] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476803734] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1476803764] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1476803795] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1476803854] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1476803883] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5236 found since last run. oneLevelSearchOps OK. searchOps is too high! 5365 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476803884] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1476804003] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9251 found since last run. oneLevelSearchOps OK. searchOps is too high! 9651 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476804003] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476804004] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1476804123] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476804603] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5728 found since last run. oneLevelSearchOps OK. searchOps is too high! 5852 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476804723] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476804723] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5789 found since last run. oneLevelSearchOps OK. searchOps is too high! 6304 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476804843] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476805054] SERVICE ALERT: hallesecondsouth;Device Status;OK;HARD;3;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1476805071] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1476805112] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1476805131] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1476805171] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1476805191] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1476805217] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1476805231] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1476805277] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1476805329] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1476805337] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1476805389] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1476805426] HOST ALERT: owen101labcolor;UP;HARD;10;PING OK - Packet loss = 75%, RTA = 0.39 ms
[1476805443] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6036 found since last run. oneLevelSearchOps OK. searchOps is too high! 6419 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476805444] SERVICE ALERT: owen101labcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1476805449] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1476805449] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1476805453] HOST ALERT: owen101lab1;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.58 ms
[1476805459] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1476805471] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 7% - OK!
[1476805475] HOST ALERT: owen101lab3;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.47 ms
[1476805483] SERVICE ALERT: owen101labcolor;Device Status;OK;HARD;3;Ready
[1476805487] SERVICE ALERT: owen101lab1;Device Status;OK;HARD;3;Ready
[1476805493] HOST ALERT: owen101lab2;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1476805509] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1476805519] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1476805533] SERVICE ALERT: owen101lab1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1476805535] SERVICE ALERT: owen101labcolor;Toner Collection Unit HP CE265A;OK;HARD;3;
[1476805543] SERVICE ALERT: owen101labcolor;Magenta Cartridge HP CE263A;OK;HARD;3;Magenta Cartridge HP CE263A is at 32% - OK!
[1476805563] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476805569] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1476805579] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1476805613] SERVICE ALERT: owen101lab3;Device Status;OK;HARD;3;ORDER CARTRIDGE LESS THAN 2000 PAGES 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1476805624] SERVICE ALERT: owen101lab2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1476805668] SERVICE ALERT: owen101labcolor;Cyan Cartridge HP CE261A;OK;HARD;3;Cyan Cartridge HP CE261A is at 31% - OK!
[1476805713] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 64% - OK!
[1476805714] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 32% - OK!
[1476805753] SERVICE ALERT: owen101lab2;Device Status;OK;HARD;3;Ready
[1476805803] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1476805863] SERVICE ALERT: crossroadspk;Device Status;OK;SOFT;2;Remove cartridge. Remove any paper.
[1476805923] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6030 found since last run. oneLevelSearchOps OK. searchOps is too high! 6074 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476805955] SERVICE ALERT: owen101lab3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1476805964] SERVICE ALERT: owen101labcolor;Image Fuser Kit HP 110V-CE246A;OK;HARD;3;Image Fuser Kit HP 110V-CE246A, 220V-CE247A is at 34% - OK!
[1476805973] SERVICE ALERT: owen101labcolor;Yellow Cartridge HP CE262A;OK;HARD;3;Yellow Cartridge HP CE262A is at 33% - OK!
[1476805984] SERVICE ALERT: owen101labcolor;Black Cartridge HP CE260X;OK;HARD;3;Black Cartridge HP CE260X is at 17% - OK!
[1476806001] SERVICE ALERT: owen101labcolor;Image Transfer Kit HP CE249A;OK;HARD;3;Image Transfer Kit HP CE249A is at 35% - OK!
[1476806043] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476806163] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7444 found since last run. oneLevelSearchOps OK. searchOps is too high! 7865 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476806260] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 44% - OK!
[1476806283] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476806403] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (25.6% change >= 20.0% threshold)
[1476806403] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1476806403] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1476806403] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1476806403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1476806403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1476806403] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1476806403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1476806403] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1476806403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1476806403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1476806403] HOST DOWNTIME ALERT: VUM;STARTED; Host has entered a period of scheduled downtime
[1476806403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1476806403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1476806403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1476806403] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1476806403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1476806403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1476806439] HOST ALERT: hallefirst1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1476806503] HOST ALERT: hallefirst1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.52 ms
[1476806503] HOST ALERT: hallefirst2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1476806531] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 9% - OK!
[1476806535] HOST ALERT: hallefirst2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1476806643] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 229 found since last run. wholeSubtreeSearchOps is too high! 9627 found since last run. oneLevelSearchOps OK. searchOps is too high! 9727 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476806763] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476806883] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5585 found since last run. oneLevelSearchOps OK. searchOps is too high! 5870 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476807003] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476807113] SERVICE ALERT: eateriespk;Device Status;OK;HARD;3;CLOSE RIGHT DOOR  46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1476807137] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1476807197] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1476807257] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1476807315] SERVICE ALERT: eateriespk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1476807363] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 201 found since last run. wholeSubtreeSearchOps is too high! 8532 found since last run. oneLevelSearchOps OK. searchOps is too high! 8645 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476807375] SERVICE ALERT: eateriespk;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 75% - OK!
[1476807399] HOST ALERT: hallefirst2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1476807409] HOST ALERT: hallefirst1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1476807463] HOST ALERT: hallefirst2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.55 ms
[1476807463] HOST FLAPPING ALERT: hallefirst2;STARTED; Host appears to have started flapping (20.7% change > 20.0% threshold)
[1476807473] HOST ALERT: hallefirst1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1476807473] HOST FLAPPING ALERT: hallefirst1;STARTED; Host appears to have started flapping (22.1% change > 20.0% threshold)
[1476807483] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476807603] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 13746 found since last run. oneLevelSearchOps OK. searchOps is too high! 14009 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476807713] SERVICE ALERT: eateriespk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1476807723] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476807773] SERVICE ALERT: eateriespk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1476807834] SERVICE ALERT: eateriespk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1476808083] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 211 found since last run. wholeSubtreeSearchOps is too high! 11124 found since last run. oneLevelSearchOps OK. searchOps is too high! 11293 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476808203] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476808228] SERVICE ALERT: hallefirst1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1476808323] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 11085 found since last run. oneLevelSearchOps OK. searchOps is too high! 11484 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476808433] SERVICE ALERT: eateriespk;Device Status;OK;HARD;3;RETRIEVE JOB  NO STORED JOBS
[1476808433] SERVICE FLAPPING ALERT: eateriespk;Device Status;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1476808443] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476808581] SERVICE ALERT: eateriespk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.59.141. Is the printer off?  Is the network down?
[1476808587] SERVICE ALERT: eateriespk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.59.141. Is the printer off?  Is the network down?
[1476808635] SERVICE ALERT: eateriespk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1476808642] SERVICE ALERT: eateriespk;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! 2 of 3 Trays Still Full
[1476808695] SERVICE ALERT: eateriespk;Black Print Cartridge HP C8543X;OK;SOFT;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1476808803] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6431 found since last run. oneLevelSearchOps OK. searchOps is too high! 6584 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476808822] SERVICE ALERT: porterpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1476808830] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1476808881] SERVICE ALERT: porterpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1476808890] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1476808894] HOST FLAPPING ALERT: hallefirst1;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1476808909] SERVICE ALERT: hallefirstcolor;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1476808923] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476808939] HOST ALERT: hallefirstcolor;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1476808941] SERVICE ALERT: porterpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1476808950] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1476808969] SERVICE ALERT: hallefirstcolor;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1476808976] HOST ALERT: hallefirstcolor;UP;SOFT;2;PING OK - Packet loss = 37%, RTA = 0.48 ms
[1476809024] SERVICE ALERT: hallefirstcolor;Device Status;OK;SOFT;3;Ready
[1476809043] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5210 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476809163] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476809393] SERVICE ALERT: studentcenterpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1476809452] SERVICE ALERT: studentcenterpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1476809477] HOST FLAPPING ALERT: hallefirst2;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1476809512] SERVICE ALERT: studentcenterpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1476809523] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7167 found since last run. oneLevelSearchOps OK. searchOps is too high! 7389 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476809643] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476809837] SERVICE ALERT: porterpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1476809841] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1476809897] SERVICE ALERT: porterpk;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 51% - OK!
[1476809901] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1476809961] SERVICE ALERT: eateriespk;Paper Tray Status;OK;SOFT;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1476810112] SERVICE ALERT: studentcenterpk;Device Status;OK;HARD;3;Open right door.
[1476810141] SERVICE ALERT: porterpk;Device Status;OK;HARD;3;Ready
[1476810166] HOST ALERT: halleground1;UP;HARD;10;PING WARNING - Packet loss = 86%, RTA = 0.50 ms
[1476810168] SERVICE ALERT: halleground1;Device Status;OK;HARD;3;ORDER CARTRIDGE LESS THAN 1500 PAGES 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Processing job
[1476810193] SERVICE ALERT: halleground2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1476810210] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1476810219] HOST ALERT: halleground2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1476810222] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1476810243] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6202 found since last run. oneLevelSearchOps OK. searchOps is too high! 6314 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476810252] HOST ALERT: halleground2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1476810253] SERVICE ALERT: halleground2;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1476810270] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1476810282] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1476810283] HOST ALERT: halleground2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1476810313] SERVICE ALERT: halleground2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1476810330] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1476810342] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1476810343] HOST ALERT: halleground2;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1476810363] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476810424] SERVICE ALERT: halleground1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1476810430] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 6% - OK!
[1476810433] HOST ALERT: halleground2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1476810523] HOST ALERT: halleground2;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1476810604] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 91% (15 min average)
[1476810613] HOST ALERT: halleground2;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1476810703] HOST ALERT: halleground2;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1476810793] HOST ALERT: halleground2;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1476810883] HOST ALERT: halleground2;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1476810963] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5059 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476811083] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476811204] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1476811324] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1476811444] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1476811683] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8009 found since last run. oneLevelSearchOps OK. searchOps is too high! 8157 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476811721] SERVICE ALERT: marshallpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.137.49. Is the printer off?  Is the network down?
[1476811780] SERVICE ALERT: marshallpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.137.49. Is the printer off?  Is the network down?
[1476811803] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476811833] SERVICE ALERT: marshallpk;Black Print Cartridge HP C8543X;OK;SOFT;3;Black Print Cartridge HP C8543X is at 8% - OK!
[1476811845] HOST ALERT: halleground2;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.53 ms
[1476812044] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1476812058] HOST ALERT: mckennypk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1476812079] SERVICE ALERT: mckennypk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1476812102] HOST ALERT: mckennypk;UP;SOFT;2;PING OK - Packet loss = 79%, RTA = 0.79 ms
[1476812107] SERVICE ALERT: halleground2;Device Status;OK;HARD;3;Ready 164.76.106.52
[1476812125] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 97% - OK!
[1476812133] SERVICE ALERT: mckennypk;Device Status;OK;SOFT;2;Ready
[1476812135] SERVICE ALERT: halleground2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1476812163] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6781 found since last run. oneLevelSearchOps OK. searchOps is too high! 7066 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476812283] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476812403] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7513 found since last run. oneLevelSearchOps OK. searchOps is too high! 7606 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476812503] HOST ALERT: hallefirst2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1476812523] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476812593] HOST ALERT: hallefirst2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1476812644] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 74% (15 min average)
[1476812679] HOST ALERT: hallefirst2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1476812769] HOST ALERT: hallefirst2;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1476812859] HOST ALERT: hallefirst2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1476812883] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 11650 found since last run. oneLevelSearchOps OK. searchOps is too high! 12088 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476812949] HOST ALERT: hallefirst2;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1476812961] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1476813003] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476813021] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1476813039] HOST ALERT: hallefirst2;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1476813081] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1476813123] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 14941 found since last run. oneLevelSearchOps OK. searchOps is too high! 14971 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476813129] HOST ALERT: hallefirst2;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1476813219] HOST ALERT: hallefirst2;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1476813243] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476813279] HOST ALERT: hallefirst2;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1476813603] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 11742 found since last run. oneLevelSearchOps OK. searchOps is too high! 12372 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476813723] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476813843] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9969 found since last run. oneLevelSearchOps OK. searchOps is too high! 10038 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476813844] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1476813963] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476813964] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1476814084] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1476814323] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5289 found since last run. oneLevelSearchOps OK. searchOps is too high! 5708 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476814443] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476814821] HOST ALERT: hallefirst1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1476814902] HOST ALERT: hallefirst1;UP;SOFT;2;PING OK - Packet loss = 77%, RTA = 0.52 ms
[1476814903] HOST ALERT: hallefirst2;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.44 ms
[1476814945] SERVICE ALERT: hallefirst1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1476814945] SERVICE FLAPPING ALERT: hallefirst1;Paper Tray Status;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1476815043] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8852 found since last run. oneLevelSearchOps OK. searchOps is too high! 9170 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476815044] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1476815051] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 9% - OK!
[1476815163] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Ready
[1476815163] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476815173] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 6% - OK!
[1476815385] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;ORDER CARTRIDGE LESS THAN 1300 PAGES 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1476815763] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5661 found since last run. oneLevelSearchOps OK. searchOps is too high! 5949 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476815883] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476816363] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6134 found since last run. oneLevelSearchOps OK. searchOps is too high! 6190 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476816483] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476817083] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5190 found since last run. oneLevelSearchOps OK. searchOps is too high! 5726 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476817083] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1476817203] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476817257] SERVICE ALERT: SCCM-PRIM01;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1476817357] SERVICE ALERT: SCCM-PRIM01;Service: RPC;OK;SOFT;2;RpcSs: Started
[1476817683] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6605 found since last run. oneLevelSearchOps OK. searchOps is too high! 6637 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476817683] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1476817803] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476817803] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6291 found since last run. oneLevelSearchOps OK. searchOps is too high! 6667 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476817923] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476818083] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1476818143] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1476818203] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1476818403] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9066 found since last run. oneLevelSearchOps OK. searchOps is too high! 9100 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476818523] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476818523] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9633 found since last run. oneLevelSearchOps OK. searchOps is too high! 10083 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476818524] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1476818583] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1476818643] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476818644] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1476818706] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1476818767] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1476818826] SERVICE ALERT: halleground2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1476818883] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1476819123] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10735 found since last run. oneLevelSearchOps OK. searchOps is too high! 10781 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476819243] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476819243] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 100 found since last run. wholeSubtreeSearchOps is too high! 6582 found since last run. oneLevelSearchOps OK. searchOps is too high! 7082 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476819363] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476819426] SERVICE ALERT: halleground2;Device Status;OK;HARD;3;TRAY 4 EMPTY PLAIN LETTER   Processing job from tray 3
[1476819426] SERVICE FLAPPING ALERT: halleground2;Device Status;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1476819483] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 88% (15 min average)
[1476819833] SERVICE FLAPPING ALERT: eateriespk;Device Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1476819843] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6911 found since last run. oneLevelSearchOps OK. searchOps is too high! 7080 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476819963] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476819963] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5608 found since last run. oneLevelSearchOps OK. searchOps is too high! 5864 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476820083] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1476820083] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476820203] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1476820323] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1476820563] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7565 found since last run. oneLevelSearchOps OK. searchOps is too high! 7790 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476820596] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1476820656] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1476820683] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476820716] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1476820924] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1476820944] SERVICE FLAPPING ALERT: hallefirst1;Paper Tray Status;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1476821283] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8016 found since last run. oneLevelSearchOps OK. searchOps is too high! 8200 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476821403] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476821552] HOST ALERT: mckennypk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1476821642] HOST ALERT: mckennypk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1476821709] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1476821731] HOST ALERT: mckennypk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1476821739] SERVICE ALERT: mckennypk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1476821769] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1476821769] HOST ALERT: mckennypk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1476821799] SERVICE ALERT: mckennypk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1476821809] SERVICE ALERT: mckennypk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1476821829] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1476821829] HOST ALERT: mckennypk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1476821859] SERVICE ALERT: mckennypk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1476821869] SERVICE ALERT: mckennypk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1476821889] HOST ALERT: mckennypk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1476821929] SERVICE ALERT: mckennypk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1476821959] HOST ALERT: mckennypk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1476822003] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6220 found since last run. oneLevelSearchOps OK. searchOps is too high! 6349 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476822049] HOST ALERT: mckennypk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1476822080] SERVICE ALERT: eateriespk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1476822123] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1476822123] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476822139] HOST ALERT: mckennypk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1476822229] HOST ALERT: mckennypk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1476822723] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1476822723] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5880 found since last run. oneLevelSearchOps OK. searchOps is too high! 6038 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476822843] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476823323] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1476823444] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7809 found since last run. oneLevelSearchOps OK. searchOps is too high! 7916 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476823563] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476823923] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1476824043] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1476824163] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1476824163] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7827 found since last run. oneLevelSearchOps OK. searchOps is too high! 7944 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476824283] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476824283] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5040 found since last run. oneLevelSearchOps OK. searchOps is too high! 5220 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476824403] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476824644] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1476824883] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5014 found since last run. oneLevelSearchOps OK. searchOps is too high! 5150 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476825003] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476825003] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6502 found since last run. oneLevelSearchOps OK. searchOps is too high! 6720 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476825123] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476825603] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5630 found since last run. oneLevelSearchOps OK. searchOps is too high! 5754 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476825723] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476825723] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8215 found since last run. oneLevelSearchOps OK. searchOps is too high! 8479 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476825843] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476826116] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1476826323] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5930 found since last run. oneLevelSearchOps OK. searchOps is too high! 5966 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476826443] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476826443] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9885 found since last run. oneLevelSearchOps OK. searchOps is too high! 10362 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476826563] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476826603] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;CRITICAL;HARD;3;Black Print Cartridge HP C8543X is at 1% - CRITICAL!
[1476826716] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1476826775] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1476826836] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1476827043] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5371 found since last run. oneLevelSearchOps OK. searchOps is too high! 5398 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476827163] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476827163] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5737 found since last run. oneLevelSearchOps OK. searchOps is too high! 5993 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476827283] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476827883] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6164 found since last run. oneLevelSearchOps OK. searchOps is too high! 6353 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476828003] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476828963] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6363 found since last run. oneLevelSearchOps OK. searchOps is too high! 6453 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476829083] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476829563] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1476829573] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1476829633] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1476829683] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8861 found since last run. oneLevelSearchOps OK. searchOps is too high! 8911 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476829694] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1476829803] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476830163] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1476830403] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5868 found since last run. oneLevelSearchOps OK. searchOps is too high! 5891 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476830523] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476830827] SERVICE FLAPPING ALERT: halleground2;Device Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1476831123] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7338 found since last run. oneLevelSearchOps OK. searchOps is too high! 7396 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476831243] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476831843] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5744 found since last run. oneLevelSearchOps OK. searchOps is too high! 5839 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476831844] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1476831904] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1476831963] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476831964] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1476833404] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6264 found since last run. oneLevelSearchOps OK. searchOps is too high! 6690 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476833523] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476834493] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;CRITICAL;HARD;3;Black Print Cartridge HP C8543X is at 1% - CRITICAL!
[1476834681] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1476834741] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1476834801] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1476835094] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at WARNING level!
[1476835323] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5067 found since last run. oneLevelSearchOps OK. searchOps is too high! 5306 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476835443] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476836163] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1476836163] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5595 found since last run. oneLevelSearchOps OK. searchOps is too high! 5679 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476836264] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1476836283] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476836323] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1476836383] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1476836763] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1476836883] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1476837003] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1476837583] SERVICE ALERT: markjeffersonpk;Device Status;OK;HARD;3;Lower green tab. Close door. 54 6F 20 65 78 69 74 20 70 72 65 73 73 20 1F 2E
[1476839293] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;CRITICAL;HARD;3;Black Print Cartridge HP C8543X is at 1% - CRITICAL!
[1476839293] SERVICE FLAPPING ALERT: hallefirst1;Black Print Cartridge HP C8543X;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1476840109] SERVICE ALERT: studentcenterpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1476840169] SERVICE ALERT: studentcenterpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1476840229] SERVICE ALERT: studentcenterpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1476840483] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5337 found since last run. oneLevelSearchOps OK. searchOps is too high! 5510 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476840603] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1476840603] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476840653] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1476840713] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1476840773] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;OK;SOFT;3;Black Print Cartridge HP C8543X is at 68% - OK!
[1476840829] SERVICE ALERT: studentcenterpk;Device Status;OK;HARD;3;Ready 164.76.150.252
[1476840964] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1476841203] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1476842043] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5202 found since last run. oneLevelSearchOps OK. searchOps is too high! 5489 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1476842163] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1476843460] SERVICE ALERT: hallesecondsouth;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1476843490] HOST ALERT: hallesecondsouth;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1476843520] SERVICE ALERT: hallesecondsouth;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1476843550] HOST ALERT: hallesecondsouth;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1476843580] SERVICE ALERT: hallesecondsouth;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1476843610] HOST ALERT: hallesecondsouth;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1476843700] HOST ALERT: hallesecondsouth;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1476843709] SERVICE ALERT: hallesecondsouth;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1476843739] HOST ALERT: hallesecondsouth;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1476843769] SERVICE ALERT: hallesecondsouth;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1476843799] HOST ALERT: hallesecondsouth;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1476843829] SERVICE ALERT: hallesecondsouth;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1476843859] HOST ALERT: hallesecondsouth;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1476843869] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1476843899] HOST ALERT: hallesecondsouth;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1476843929] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1476843959] HOST ALERT: hallesecondsouth;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1476843989] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1476844019] HOST ALERT: hallesecondsouth;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1476846801] SERVICE ALERT: eateriespk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1476847093] SERVICE FLAPPING ALERT: hallefirst1;Black Print Cartridge HP C8543X;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)