[1479532272] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1479532392] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1479532512] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 91% (15 min average)
[1479533112] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1479533232] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 23% (15 min average)
[1479534432] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1479534552] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1479534672] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 83% (15 min average)
[1479535272] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1479535392] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 80% (15 min average)
[1479535958] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1479535992] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1479536080] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1479536112] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1479536232] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1479536679] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1479536798] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;2;Warning:  idm1  are more than 120 seconds behind!
[1479536919] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1479537002] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1479537518] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1479538032] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 78% (15 min average)
[1479538118] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1479538632] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1479538718] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1479538718] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1479538752] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 81% (15 min average)
[1479539318] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1479539438] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1479540038] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1479540158] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1479540278] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1479540402] SERVICE ALERT: EVENTSTEST-2012;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1479540522] SERVICE ALERT: EVENTSTEST-2012;CPU Load;OK;SOFT;2;CPU Load 86% (15 min average)
[1479540552] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1479540672] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1479540792] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1479540792] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1479541478] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1479541992] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (15 min average)
[1479542239] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 111 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479542359] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479542592] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1479542678] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1479542712] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 79% (15 min average)
[1479542902] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 106 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479543022] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479543877] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1479543912] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1479543997] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;2;Warning:  idm1  are more than 120 seconds behind!
[1479544032] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 84% (15 min average)
[1479544117] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;3;OK. All replicas within 120 seconds.
[1479544632] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1479544717] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1479544752] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1479544837] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1479544872] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1479544889] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1479544957] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1479545009] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1479545129] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1479545472] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (15 min average)
[1479545729] SERVICE ALERT: epprogsql2;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1479546072] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1479546157] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1479546192] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1479546312] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 89% (15 min average)
[1479546329] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1479547158] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 102 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479547278] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479548129] SERVICE ALERT: epprogsql2;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1479548129] SERVICE FLAPPING ALERT: epprogsql2;CPU Load;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1479548729] SERVICE ALERT: epprogsql2;CPU Load;OK;HARD;3;CPU Load 84% (15 min average)
[1479549182] SERVICE ALERT: INTLAPP-ISSMAP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1479549302] SERVICE ALERT: INTLAPP-ISSMAP;CPU Load;OK;SOFT;2;CPU Load 7% (15 min average)
[1479550878] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 111 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479550998] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479553398] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 195 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479553518] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479553822] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 391 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479553942] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479554118] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 221 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479554238] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479555002] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service has exited from a period of scheduled downtime
[1479555051] SERVICE ALERT: INTLAPP;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1479555171] SERVICE ALERT: INTLAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1479555291] SERVICE ALERT: INTLAPP;CPU Load;OK;SOFT;3;CPU Load 10% (15 min average)
[1479555312] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1479555432] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1479555552] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1479556954] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1479557352] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 79% (15 min average)
[1479560129] SERVICE FLAPPING ALERT: epprogsql2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1479560952] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1479561072] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 62% (15 min average)
[1479564197] HOST ALERT: hallefirst2;UP;HARD;10;PING WARNING - Packet loss = 83%, RTA = 0.50 ms
[1479564217] HOST ALERT: hallefirst1;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.59 ms
[1479564245] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1479564310] SERVICE ALERT: hallefirstcolor;Device Status;OK;HARD;3;RETRIEVE JOB  NO STORED JOBS
[1479564314] HOST ALERT: hallefirstcolor;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1479564333] SERVICE ALERT: hallefirstcolor;Image Fuser Kit HP 110V-CE246A;OK;HARD;3;Image Fuser Kit HP 110V-CE246A, 220V-CE247A is at 100% - OK!
[1479564352] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;OK;HARD;3;Cyan Cartridge HP CE261A is at 35% - OK!
[1479564352] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;OK;HARD;3;Yellow Cartridge HP CE262A is at 31% - OK!
[1479564361] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;ORDER CARTRIDGE LESS THAN 1800 PAGES 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Processing job from tray 3
[1479564361] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;OK;HARD;3;Magenta Cartridge HP CE263A is at 47% - OK!
[1479564472] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;OK;HARD;3;Black Cartridge HP CE260X is at 51% - OK!
[1479564506] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 8% - OK!
[1479564507] SERVICE ALERT: hallefirst1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1479564532] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1479564539] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Ready
[1479564626] SERVICE ALERT: hallefirstcolor;Toner Collection Unit HP CE265A;OK;HARD;3;
[1479564646] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;OK;HARD;3;Image Transfer Kit HP CE249A is at 63% - OK!
[1479564762] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 45% - OK!
[1479565452] SERVICE ALERT: marshallpk;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1479565513] SERVICE ALERT: marshallpk;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1479565572] SERVICE ALERT: marshallpk;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1479567672] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1479574272] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1479574392] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 89% (15 min average)
[1479574992] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1479575112] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 83% (15 min average)
[1479578112] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1479578232] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 90% (15 min average)
[1479578238] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 156 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479578358] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479580032] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1479580152] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1479584943] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1479584952] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1479584973] HOST ALERT: hallefirst2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1479585003] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1479585011] HOST ALERT: hallefirstcolor;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1479585031] SERVICE ALERT: hallefirstcolor;Toner Collection Unit HP CE265A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1479585033] HOST ALERT: hallefirst2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1479585052] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1479585061] HOST ALERT: hallefirstcolor;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1479585063] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1479585072] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 71% (15 min average)
[1479585091] SERVICE ALERT: hallefirstcolor;Toner Collection Unit HP CE265A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1479585093] HOST ALERT: hallefirst2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1479585112] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1479585121] HOST ALERT: hallefirstcolor;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1479585151] SERVICE ALERT: hallefirstcolor;Toner Collection Unit HP CE265A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1479585168] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1479585172] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1479585181] HOST ALERT: hallefirstcolor;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1479585183] HOST ALERT: hallefirst2;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1479585217] HOST ALERT: hallefirst1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1479585228] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1479585250] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1479585258] HOST ALERT: hallefirst2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1479585271] HOST ALERT: hallefirstcolor;DOWN;SOFT;5;(Host check timed out after 30.00 seconds)
[1479585288] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1479585307] HOST ALERT: hallefirst1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1479585310] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1479585314] SERVICE ALERT: hallefirstcolor;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1479585318] HOST ALERT: hallefirst2;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1479585338] SERVICE ALERT: hallefirstcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1479585344] HOST ALERT: hallefirstcolor;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1479585358] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1479585358] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1479585366] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1479585366] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1479585370] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1479585374] SERVICE ALERT: hallefirstcolor;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1479585388] HOST ALERT: hallefirstcolor;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1479585396] HOST ALERT: hallefirst1;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1479585398] SERVICE ALERT: hallefirstcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1479585400] HOST ALERT: hallefirst2;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1479585418] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1479585418] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1479585426] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1479585426] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1479585428] HOST ALERT: hallefirstcolor;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1479585434] SERVICE ALERT: hallefirstcolor;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1479585456] HOST ALERT: hallefirst1;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1479585458] SERVICE ALERT: hallefirstcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1479585464] HOST ALERT: hallefirstcolor;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1479585476] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1479585478] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1479585478] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1479585486] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1479585486] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1479585490] HOST ALERT: hallefirst2;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1479585506] HOST ALERT: hallefirstcolor;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1479585510] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1479585513] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1479585516] HOST ALERT: hallefirst1;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1479585535] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1479585538] SERVICE ALERT: hallefirstcolor;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1479585570] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1479585573] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1479585580] HOST ALERT: hallefirst2;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1479585595] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1479585598] SERVICE ALERT: hallefirstcolor;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1479585600] HOST ALERT: hallefirst1;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1479585630] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1479585633] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1479585658] SERVICE ALERT: hallefirstcolor;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1479585660] HOST ALERT: hallefirst1;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1479585670] HOST ALERT: hallefirst2;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1479585742] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 130 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479585750] HOST ALERT: hallefirst1;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1479585840] HOST ALERT: hallefirst1;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1479585862] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479585930] HOST ALERT: hallefirst1;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1479586157] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 112 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479586277] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479586482] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 45% - OK!
[1479586486] HOST ALERT: hallefirst2;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1479586537] HOST ALERT: hallefirst1;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.56 ms
[1479586565] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1479586680] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;ORDER CARTRIDGE LESS THAN 1800 PAGES 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1479586824] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 8% - OK!
[1479586828] SERVICE ALERT: hallefirst1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1479586857] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Ready
[1479588077] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 118 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479588197] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479588262] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 109 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479588382] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479588797] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 142 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479588917] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479589272] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1479589392] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 67% (15 min average)
[1479589517] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 161 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479589582] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 179 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479589637] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479589702] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479589992] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1479590112] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 87% (15 min average)
[1479590712] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1479590832] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1479590952] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 87% (15 min average)
[1479592152] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1479592252] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1479592256] HOST ALERT: hallefirstcolor;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1479592272] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 63% (15 min average)
[1479592343] SERVICE ALERT: hallefirstcolor;Toner Collection Unit HP CE265A;OK;HARD;3;
[1479592366] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;OK;HARD;3;Image Transfer Kit HP CE249A is at 63% - OK!
[1479592626] SERVICE ALERT: hallefirstcolor;Device Status;OK;HARD;3;Ready
[1479592652] SERVICE ALERT: hallefirstcolor;Image Fuser Kit HP 110V-CE246A;OK;HARD;3;Image Fuser Kit HP 110V-CE246A, 220V-CE247A is at 100% - OK!
[1479592672] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;OK;HARD;3;Cyan Cartridge HP CE261A is at 35% - OK!
[1479592672] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;OK;HARD;3;Yellow Cartridge HP CE262A is at 31% - OK!
[1479592680] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;OK;HARD;3;Magenta Cartridge HP CE263A is at 46% - OK!
[1479592788] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;OK;HARD;3;Black Cartridge HP CE260X is at 51% - OK!
[1479592872] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1479592992] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1479593112] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1479593712] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 45% (15 min average)
[1479594652] SERVICE ALERT: eateriespk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1479594712] SERVICE ALERT: eateriespk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1479594772] SERVICE ALERT: eateriespk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1479594912] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1479595032] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1479595101] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 108 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479595152] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1479595221] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479595752] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 68% (15 min average)
[1479595752] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1479595972] SERVICE ALERT: eateriespk;Device Status;OK;HARD;3;TRAY 4 EMPTY PLAIN LETTER   Ready
[1479596352] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1479596472] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1479596592] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1479598992] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 63% (15 min average)
[1479600191] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1479600310] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 90% (15 min average)
[1479600437] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 129 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479600557] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479600621] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 111 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479600741] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479600910] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1479601029] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1479601149] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1479602349] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1479602541] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 106 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479602661] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479606547] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1479606667] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1479607267] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1479607387] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1479607507] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 93% (15 min average)
[1479608707] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1479608827] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1479608947] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1479609547] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 75% (15 min average)
[1479610147] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1479610267] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1479610387] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1479610987] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1479611587] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1479611707] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1479611827] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1479612049] SERVICE ALERT: ADDC2;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1479612049] SERVICE ALERT: ADDC2;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1479612052] SERVICE ALERT: ADDC2;Service: DNS Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1479612052] SERVICE ALERT: ADDC2;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1479612149] SERVICE ALERT: ADDC2;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1479612149] SERVICE ALERT: ADDC2;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1479612152] SERVICE ALERT: ADDC2;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1479612152] SERVICE ALERT: ADDC2;Service: DNS Service;OK;SOFT;2;DNS: Started
[1479613027] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 88% (15 min average)
[1479613627] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1479613747] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1479613757] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 103 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479613867] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1479613877] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479616866] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 90% (15 min average)
[1479617466] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1479617586] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)