[1457586707] SERVICE ALERT: halleground1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1457586767] SERVICE ALERT: halleground1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1457586827] SERVICE ALERT: halleground1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1457587800] HOST DOWNTIME ALERT: UHSPHARM-2012;STOPPED; Host has exited from a period of scheduled downtime
[1457587800] HOST DOWNTIME ALERT: TI22-PSYCLINIC;STOPPED; Host has exited from a period of scheduled downtime
[1457587800] HOST DOWNTIME ALERT: MBAM;STOPPED; Host has exited from a period of scheduled downtime
[1457587800] HOST DOWNTIME ALERT: PNCDB-2012;STOPPED; Host has exited from a period of scheduled downtime
[1457587800] HOST DOWNTIME ALERT: REFWIN;STOPPED; Host has exited from a period of scheduled downtime
[1457587800] HOST DOWNTIME ALERT: LANSWEEPER;STOPPED; Host has exited from a period of scheduled downtime
[1457587800] HOST DOWNTIME ALERT: IMAGES;STOPPED; Host has exited from a period of scheduled downtime
[1457587800] HOST DOWNTIME ALERT: WINKMS;STOPPED; Host has exited from a period of scheduled downtime
[1457587800] HOST DOWNTIME ALERT: OFFICEKMS;STOPPED; Host has exited from a period of scheduled downtime
[1457587800] HOST DOWNTIME ALERT: PROCARE;STOPPED; Host has exited from a period of scheduled downtime
[1457587800] HOST DOWNTIME ALERT: FTK;STOPPED; Host has exited from a period of scheduled downtime
[1457587800] HOST DOWNTIME ALERT: HDBACKUPS;STOPPED; Host has exited from a period of scheduled downtime
[1457587800] HOST DOWNTIME ALERT: XDM;STOPPED; Host has exited from a period of scheduled downtime
[1457587800] HOST DOWNTIME ALERT: TI22;STOPPED; Host has exited from a period of scheduled downtime
[1457587800] HOST DOWNTIME ALERT: Echo-01;STOPPED; Host has exited from a period of scheduled downtime
[1457587800] HOST DOWNTIME ALERT: EMUSQLEXPRESS;STOPPED; Host has exited from a period of scheduled downtime
[1457587800] HOST DOWNTIME ALERT: PNCWEB-2012;STOPPED; Host has exited from a period of scheduled downtime
[1457587800] HOST DOWNTIME ALERT: EMUSCRIPT;STOPPED; Host has exited from a period of scheduled downtime
[1457587800] HOST DOWNTIME ALERT: TABLEAU;STOPPED; Host has exited from a period of scheduled downtime
[1457587800] HOST DOWNTIME ALERT: CEDARON;STOPPED; Host has exited from a period of scheduled downtime
[1457587800] HOST DOWNTIME ALERT: MBAMDB;STOPPED; Host has exited from a period of scheduled downtime
[1457587800] HOST DOWNTIME ALERT: OPIE;STOPPED; Host has exited from a period of scheduled downtime
[1457587800] HOST DOWNTIME ALERT: NESSUS;STOPPED; Host has exited from a period of scheduled downtime
[1457587800] HOST DOWNTIME ALERT: DPS-SECURE02;STOPPED; Host has exited from a period of scheduled downtime
[1457587800] HOST DOWNTIME ALERT: AD-RMAD;STOPPED; Host has exited from a period of scheduled downtime
[1457587800] HOST DOWNTIME ALERT: EMU-VALUE01;STOPPED; Host has exited from a period of scheduled downtime
[1457587800] HOST DOWNTIME ALERT: SANDBOX-DB;STOPPED; Host has exited from a period of scheduled downtime
[1457588027] SERVICE ALERT: halleground1;Device Status;OK;HARD;3;Processing job from tray 4
[1457589756] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 78% (5 min average)
[1457590310] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 84.90 Gb - used: 78.73 Gb (93%) - free 6.17 Gb (7%)
[1457590356] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1457590476] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1457590596] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1457590783] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;SOFT;1;WARNING: Replication with ADDC5 failed!
[1457590902] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;SOFT;2;WARNING: Replication with ADDC5 failed!
[1457591023] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;HARD;3;WARNING: Replication with ADDC5 failed!
[1457591796] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 94% (5 min average)
[1457592396] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1457592396] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1457592442] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1457592561] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 68% (5 min average)
[1457593273] SERVICE ALERT: OFFICEKMS;CPU Load;WARNING;SOFT;1;CPU Load 94% (5 min average)
[1457593393] SERVICE ALERT: OFFICEKMS;CPU Load;OK;SOFT;2;CPU Load 85% (5 min average)
[1457594241] SERVICE ALERT: XDM;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457594361] SERVICE ALERT: XDM;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1457594961] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1457595081] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 75% (5 min average)
[1457596618] SERVICE ALERT: MBAM;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457596737] SERVICE ALERT: MBAM;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1457597274] SERVICE ALERT: TARGETX-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457597394] SERVICE ALERT: TARGETX-2012;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1457597481] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1457597514] SERVICE ALERT: TARGETX-2012;CPU Load;OK;SOFT;3;CPU Load 36% (5 min average)
[1457597601] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 81% (5 min average)
[1457600001] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1457600121] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 77% (5 min average)
[1457601227] SERVICE ALERT: ADRODC-LIV;AD Replication;OK;HARD;3;OK: replication is up and running
[1457601996] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 75% (5 min average)
[1457602011] SERVICE ALERT: SANDBOX-DB;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457602131] SERVICE ALERT: SANDBOX-DB;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1457602251] SERVICE ALERT: SANDBOX-DB;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1457602851] SERVICE ALERT: SANDBOX-DB;CPU Load;OK;HARD;3;CPU Load 0% (5 min average)
[1457603196] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1457603316] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 46% (5 min average)
[1457605116] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1457605236] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1457605356] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 87% (5 min average)
[1457605956] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457606076] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1457606196] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1457607396] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 53% (5 min average)
[1457607996] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457608116] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1457608236] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1457610636] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 72% (5 min average)
[1457611250] SERVICE ALERT: SANDBOX-DB;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1457611370] SERVICE ALERT: SANDBOX-DB;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1457611836] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457611956] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1457612076] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1457612676] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 61% (5 min average)
[1457613276] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457613396] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1457613516] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1457615377] HOST ALERT: owen101lab3;UP;HARD;10;PING WARNING - Packet loss = 86%, RTA = 0.66 ms
[1457615541] SERVICE ALERT: owen101lab3;Device Status;OK;HARD;3;Processing job from tray 4
[1457615761] SERVICE ALERT: owen101lab3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1457615941] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 54% - OK!
[1457616516] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 54% (5 min average)
[1457618916] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1457619036] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1457619156] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1457619709] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 84.90 Gb - used: 75.57 Gb (89%) - free 9.33 Gb (11%)
[1457620203] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (38.2% change >= 20.0% threshold)
[1457620956] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 86% (5 min average)
[1457622156] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457622276] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1457622396] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1457623803] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (33.9% change >= 20.0% threshold)
[1457625906] SERVICE ALERT: ADDC3;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457626006] SERVICE ALERT: ADDC3;Service: Windows Time;CRITICAL;SOFT;2;W32Time: Stopped
[1457626126] SERVICE ALERT: ADDC3;Service: Windows Time;OK;SOFT;3;W32Time: Started
[1457631141] SERVICE ALERT: owen101lab3;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1457631201] SERVICE ALERT: owen101lab3;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1457631261] SERVICE ALERT: owen101lab3;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1457631336] SERVICE ALERT: ADDC5;GPO Replication;UNKNOWN;SOFT;1;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1457631356] SERVICE ALERT: ADDC5;GPO Replication;OK;SOFT;2;Export-CSV : The process cannot access the file 'C:\Program
[1457631541] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1457631601] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 52% - OK!
[1457631861] SERVICE ALERT: owen101lab3;Device Status;OK;HARD;3;SIZE MISMATCH IN TRAY 4 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1457640994] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 87% (5 min average)
[1457641594] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457641714] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 88% (5 min average)
[1457641979] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1457642039] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1457642099] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1457642914] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457643034] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1457643150] SERVICE ALERT: VCENTERSSO;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457643154] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1457643164] SERVICE ALERT: mckennypk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1457643223] SERVICE ALERT: mckennypk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1457643270] SERVICE ALERT: VCENTERSSO;CPU Load;OK;SOFT;2;CPU Load 7% (5 min average)
[1457643284] SERVICE ALERT: mckennypk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1457643633] HOST ALERT: ADAUTH-RMAD;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1457643649] SERVICE ALERT: ADAUTH-RMAD;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457643649] SERVICE ALERT: ADAUTH-RMAD;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457643649] SERVICE ALERT: ADAUTH-RMAD;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457643649] SERVICE ALERT: ADAUTH-RMAD;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457643649] SERVICE ALERT: ADAUTH-RMAD;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457643649] SERVICE ALERT: ADAUTH-RMAD;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457643649] SERVICE ALERT: ADAUTH-RMAD;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457643649] SERVICE ALERT: ADAUTH-RMAD;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457643649] SERVICE ALERT: ADAUTH-RMAD;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457643649] SERVICE ALERT: ADAUTH-RMAD;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457643649] SERVICE ALERT: ADAUTH-RMAD;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457643649] SERVICE ALERT: ADAUTH-RMAD;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457643649] SERVICE ALERT: ADAUTH-RMAD;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457643649] SERVICE ALERT: ADAUTH-RMAD;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457643679] HOST ALERT: ADAUTH-RMAD;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1457643769] SERVICE ALERT: ADAUTH-RMAD;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457643769] SERVICE ALERT: ADAUTH-RMAD;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457643769] SERVICE ALERT: ADAUTH-RMAD;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457643769] SERVICE ALERT: ADAUTH-RMAD;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457643769] SERVICE ALERT: ADAUTH-RMAD;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457643769] SERVICE ALERT: ADAUTH-RMAD;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457643769] SERVICE ALERT: ADAUTH-RMAD;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457643769] SERVICE ALERT: ADAUTH-RMAD;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457643769] SERVICE ALERT: ADAUTH-RMAD;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457643769] SERVICE ALERT: ADAUTH-RMAD;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457643769] SERVICE ALERT: ADAUTH-RMAD;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457643769] SERVICE ALERT: ADAUTH-RMAD;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457643769] SERVICE ALERT: ADAUTH-RMAD;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457643769] HOST ALERT: ADAUTH-RMAD;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1457643769] SERVICE ALERT: ADAUTH-RMAD;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457643799] HOST ALERT: ADAUTH-RMAD;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1457643889] SERVICE ALERT: ADAUTH-RMAD;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457643889] SERVICE ALERT: ADAUTH-RMAD;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457643889] SERVICE ALERT: ADAUTH-RMAD;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457643889] SERVICE ALERT: ADAUTH-RMAD;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457643889] SERVICE ALERT: ADAUTH-RMAD;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457643889] SERVICE ALERT: ADAUTH-RMAD;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457643889] SERVICE ALERT: ADAUTH-RMAD;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457643889] SERVICE ALERT: ADAUTH-RMAD;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457643889] SERVICE ALERT: ADAUTH-RMAD;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457643889] SERVICE ALERT: ADAUTH-RMAD;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457643889] SERVICE ALERT: ADAUTH-RMAD;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457643889] SERVICE ALERT: ADAUTH-RMAD;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457643889] SERVICE ALERT: ADAUTH-RMAD;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457643889] SERVICE ALERT: ADAUTH-RMAD;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457643889] HOST ALERT: ADAUTH-RMAD;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1457643979] HOST ALERT: ADAUTH-RMAD;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1457644069] HOST ALERT: ADAUTH-RMAD;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1457644159] HOST ALERT: ADAUTH-RMAD;DOWN;SOFT;8;(Host check timed out after 30.00 seconds)
[1457644249] HOST ALERT: ADAUTH-RMAD;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1457644339] HOST ALERT: ADAUTH-RMAD;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1457644394] HOST ALERT: ADAUTH-RMAD;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.41 ms
[1457644491] HOST ALERT: ADAUTH-RMAD;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1457644581] HOST ALERT: ADAUTH-RMAD;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1457644671] HOST ALERT: ADAUTH-RMAD;DOWN;SOFT;3;(Host check timed out after 30.00 seconds)
[1457644761] HOST ALERT: ADAUTH-RMAD;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1457644851] HOST ALERT: ADAUTH-RMAD;DOWN;SOFT;5;(Host check timed out after 30.00 seconds)
[1457644941] HOST ALERT: ADAUTH-RMAD;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1457645031] HOST ALERT: ADAUTH-RMAD;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1457645091] HOST ALERT: ADAUTH-RMAD;DOWN;SOFT;8;(Host check timed out after 30.01 seconds)
[1457645181] HOST ALERT: ADAUTH-RMAD;DOWN;SOFT;9;(Host check timed out after 30.00 seconds)
[1457645271] HOST ALERT: ADAUTH-RMAD;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1457645554] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1457646154] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1457646154] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1457646825] SERVICE ALERT: halleground1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1457646885] SERVICE ALERT: halleground1;Device Status;OK;SOFT;2;MENUS  RETRIEVE JOB  INFORMATION  PAPER HANDLING
[1457647485] SERVICE ALERT: halleground1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1457647545] SERVICE ALERT: halleground1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1457647605] SERVICE ALERT: halleground1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1457649405] SERVICE ALERT: halleground1;Device Status;OK;HARD;3;Ready
[1457650019] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1457650079] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1457650139] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1457651938] SERVICE ALERT: hallesecondsouth;Device Status;OK;HARD;3;Ready
[1457657098] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1457657554] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1457658754] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 75% (5 min average)
[1457659354] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457659474] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 92% (5 min average)
[1457659594] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1457661711] SERVICE ALERT: TARGETX-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457661831] SERVICE ALERT: TARGETX-2012;CPU Load;WARNING;SOFT;2;CPU Load 93% (5 min average)
[1457661951] SERVICE ALERT: TARGETX-2012;CPU Load;CRITICAL;HARD;3;CPU Load 95% (5 min average)
[1457661976] SERVICE ALERT: SP2010DB;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 69.90 Gb - used: 55.99 Gb (80%) - free 13.91 Gb (20%)
[1457662096] SERVICE ALERT: SP2010DB;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 69.90 Gb - used: 55.99 Gb (80%) - free 13.91 Gb (20%)
[1457662216] SERVICE ALERT: SP2010DB;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 69.90 Gb - used: 55.99 Gb (80%) - free 13.91 Gb (20%)
[1457662551] SERVICE ALERT: TARGETX-2012;CPU Load;OK;HARD;3;CPU Load 70% (5 min average)
[1457662816] SERVICE ALERT: SP2010DB;C:\ Drive Space;OK;HARD;3;c:\ - total: 69.90 Gb - used: 55.28 Gb (79%) - free 14.62 Gb (21%)
[1457662931] SERVICE ALERT: EVENTS-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1457663051] SERVICE ALERT: EVENTS-2012;CPU Load;OK;SOFT;2;CPU Load 76% (5 min average)
[1457666743] SERVICE ALERT: CBORD;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457666794] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 68% (5 min average)
[1457666843] SERVICE ALERT: CBORD;C:\ Drive Space;OK;SOFT;2;c:\ - total: 59.66 Gb - used: 36.11 Gb (61%) - free 23.55 Gb (39%)
[1457667255] SERVICE ALERT: ADWSUSSERVERS;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.102.133 and port 12489: No route to host
[1457667255] SERVICE ALERT: ADWSUSSERVERS;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.102.133 and port 12489: No route to host
[1457667258] HOST ALERT: ADWSUSSERVERS;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.133)
[1457667277] HOST ALERT: ADWSUSSERVERS;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.36 ms
[1457667363] SERVICE ALERT: ADWSUSSERVERS;Memory Usage;OK;SOFT;2;Memory usage: total:32767.77 MB - used: 7605.01 MB (23%) - free: 25162.76 MB (77%)
[1457667369] SERVICE ALERT: ADWSUSSERVERS;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1457667394] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457667514] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1457667634] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1457667634] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1457667824] SERVICE ALERT: VC;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457667924] SERVICE ALERT: VC;Memory Usage;OK;SOFT;2;Memory usage: total:28669.44 MB - used: 6809.47 MB (24%) - free: 21859.96 MB (76%)
[1457670633] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 92% (5 min average)
[1457671233] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1457671833] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 85% (5 min average)