[1479618186] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1479618306] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1479618426] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 61% (15 min average)
[1479620826] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1479620945] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1479621066] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1479622352] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1479622473] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;2;Warning:  idm1  are more than 120 seconds behind!
[1479622593] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;3;OK. All replicas within 120 seconds.
[1479623195] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1479623312] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1479623400] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1479623433] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;3;OK. All replicas within 120 seconds.
[1479624033] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1479624062] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 87% (15 min average)
[1479624152] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;2;Warning:  idm1  are more than 120 seconds behind!
[1479624272] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1479624872] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1479625262] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1479625382] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1479625502] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 69% (15 min average)
[1479626102] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1479626222] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1479626342] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1479627272] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1479627542] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 82% (15 min average)
[1479627872] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1479627872] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1479628142] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1479628262] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1479628382] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1479629072] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1479629672] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1479631472] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1479631526] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1479631646] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1479631766] SERVICE ALERT: epprogsql2;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1479631982] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1479632366] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1479632582] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1479632672] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1479632702] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1479632792] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1479632822] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1479633392] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1479633422] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 46% (15 min average)
[1479633512] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;2;Warning:  idm1  are more than 120 seconds behind!
[1479633632] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;3;OK. All replicas within 120 seconds.
[1479635366] SERVICE ALERT: epprogsql2;CPU Load;OK;HARD;3;CPU Load 74% (15 min average)
[1479640022] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1479640142] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1479640262] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1479640862] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1479641400] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service has exited from a period of scheduled downtime
[1479642032] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1479642062] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 72% (15 min average)
[1479642662] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1479642782] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 79% (15 min average)
[1479645182] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1479645302] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1479645421] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 65% (15 min average)
[1479646019] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1479646139] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 91% (15 min average)
[1479653339] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1479661739] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1479661859] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 90% (15 min average)
[1479664259] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1479664379] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1479664499] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1479665099] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 65% (15 min average)
[1479665225] SERVICE ALERT: hallefirstcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1479665285] SERVICE ALERT: hallefirstcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1479665345] SERVICE ALERT: hallefirstcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1479665945] SERVICE ALERT: hallefirstcolor;Device Status;OK;HARD;3;Ready
[1479667499] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1479667619] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 85% (15 min average)
[1479669654] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1479669714] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1479669774] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1479670375] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;INPUT PAPER PATH OPEN 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1479672002] SERVICE ALERT: BBTS_DB;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.227.4 and port 12489: Connection refused
[1479672122] SERVICE ALERT: BBTS_DB;Service: RPC;OK;SOFT;2;RpcSs: Started
[1479673403] SERVICE ALERT: eateriespk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1479677219] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1479677339] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1479678182] SERVICE ALERT: studentcenterpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1479678208] SERVICE ALERT: hallegroundcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1479678228] SERVICE ALERT: hallegroundcolor;Cyan Cartridge HP CE261A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1479678238] SERVICE ALERT: hallegroundcolor;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1479678242] SERVICE ALERT: studentcenterpk;Device Status;OK;SOFT;2;Genuine HP cartridge installed
[1479678268] SERVICE ALERT: hallegroundcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1479678288] SERVICE ALERT: hallegroundcolor;Cyan Cartridge HP CE261A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1479678292] SERVICE ALERT: hallegroundcolor;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! All Trays Still Full
[1479678323] SERVICE ALERT: hallegroundcolor;Image Fuser Kit HP 110V-CE246A;OK;SOFT;3;Image Fuser Kit HP 110V-CE246A, 220V-CE247A is at 53% - OK!
[1479678342] SERVICE ALERT: hallegroundcolor;Cyan Cartridge HP CE261A;OK;SOFT;3;Cyan Cartridge HP CE261A is at 47% - OK!
[1479679738] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1479679858] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 28% (15 min average)
[1479680458] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1479680578] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 51% (15 min average)
[1479682352] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1479682412] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1479682472] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1479682778] SERVICE ALERT: hallegroundcolor;Magenta Cartridge HP CE263A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1479682833] SERVICE ALERT: hallegroundcolor;Magenta Cartridge HP CE263A;OK;SOFT;2;Magenta Cartridge HP CE263A is at 53% - OK!
[1479684178] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1479684298] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 84% (15 min average)
[1479687073] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5024 found since last run. oneLevelSearchOps OK. searchOps is too high! 5155 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1479687193] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479688993] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5898 found since last run. oneLevelSearchOps OK. searchOps is too high! 6017 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1479689113] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479690898] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1479691018] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 84% (15 min average)
[1479692218] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1479692338] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 86% (15 min average)
[1479694138] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1479694258] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1479694378] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1479695113] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5005 found since last run. oneLevelSearchOps OK. searchOps is too high! 5113 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1479695233] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479695656] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 121 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479695776] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479695833] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 118 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479695953] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479696178] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1479696778] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1479696898] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 54% (15 min average)
[1479697498] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1479697557] SERVICE ALERT: SANDBOX-APP;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1479697576] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 102 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479697618] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1479697657] SERVICE ALERT: SANDBOX-APP;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1479697696] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479697738] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1479697753] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 103 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479697873] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1479698938] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1479698938] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1479699538] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1479699658] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1479699778] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 90% (15 min average)
[1479700378] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1479700498] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1479701098] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1479701192] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1479701218] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 91% (15 min average)
[1479701252] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1479701312] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1479701409] SERVICE ALERT: hallesecondsouth;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1479701469] SERVICE ALERT: hallesecondsouth;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1479701529] SERVICE ALERT: hallesecondsouth;Paper Tray Status;OK;SOFT;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1479701912] SERVICE ALERT: hallesecondsouth;Device Status;OK;HARD;3;Ready
[1479703018] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1479703138] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 86% (15 min average)
[1479704338] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)