[1493006414] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1493006474] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1493006534] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 71% (15 min average)
[1493006778] SERVICE ALERT: TABLEAU;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1493006838] SERVICE ALERT: TABLEAU;CPU Load;OK;SOFT;2;CPU Load 0% (15 min average)
[1493007734] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1493007794] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 92% (15 min average)
[1493008201] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1493008394] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1493008454] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1493009054] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1493009114] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1493009174] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 94% (15 min average)
[1493010974] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1493011034] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1493011094] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1493011222] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1493011282] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1493011342] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1493011941] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1493012540] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1493012894] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 75% (15 min average)
[1493014094] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1493014154] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1493014214] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1493014814] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1493015414] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1493015540] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1493015540] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1493016014] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 79% (15 min average)
[1493016614] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1493016674] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1493016734] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1493016741] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1493017334] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1493017934] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1493020934] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 78% (15 min average)
[1493020938] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1493021538] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1493022600] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service has exited from a period of scheduled downtime
[1493022602] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (48.7% change >= 20.0% threshold)
[1493022602] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (34.0% change >= 20.0% threshold)
[1493022602] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1493022602] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1493022602] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1493022602] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1493022602] HOST DOWNTIME ALERT: INTLAPP;STARTED; Host has entered a period of scheduled downtime
[1493022602] HOST DOWNTIME ALERT: INTLDB;STARTED; Host has entered a period of scheduled downtime
[1493022602] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1493022602] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1493022602] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1493022602] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1493022602] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1493022602] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1493022602] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1493022602] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1493022602] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1493022602] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1493022602] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1493024402] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (41.2% change >= 20.0% threshold)
[1493024402] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1493024402] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1493024402] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1493024402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1493024402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1493024402] HOST DOWNTIME ALERT: INTLAPP;STARTED; Host has entered a period of scheduled downtime
[1493024402] HOST DOWNTIME ALERT: INTLDB;STARTED; Host has entered a period of scheduled downtime
[1493024402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1493024402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1493024402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1493024402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1493024402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1493024402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1493024402] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1493024402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1493024402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1493024402] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1493024402] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1493025138] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1493025198] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;2;Warning:  idm1  are more than 120 seconds behind!
[1493025258] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1493025858] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1493026998] SERVICE ALERT: ADWU1;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1493027058] SERVICE ALERT: ADWU1;CPU Load;OK;SOFT;2;CPU Load 16% (15 min average)
[1493028734] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1493028794] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1493028854] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1493029454] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1493030054] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1493030114] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1493030174] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1493034974] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 61% (15 min average)
[1493036362] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5579 found since last run. oneLevelSearchOps OK. searchOps is too high! 5627 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493036422] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493036774] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1493036834] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1493036894] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 92% (15 min average)
[1493037022] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5031 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493037082] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493037258] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1493037494] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1493037554] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1493037614] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1493037682] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5072 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493037742] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493037824] SERVICE ALERT: ADARC;Service: WMI;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1493037884] SERVICE ALERT: ADARC;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1493038214] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 75% (15 min average)
[1493038342] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7138 found since last run. oneLevelSearchOps OK. searchOps is too high! 7287 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493038402] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493039002] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5808 found since last run. oneLevelSearchOps OK. searchOps is too high! 5977 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493039062] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493039297] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5554 found since last run. oneLevelSearchOps OK. searchOps is too high! 5733 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493039357] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493040324] SERVICE ALERT: halleground2;Device Status;OK;HARD;3;Ready 164.76.106.52
[1493040603] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (28.8% change >= 20.0% threshold)
[1493040603] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1493040603] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1493040603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1493040603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1493040603] HOST DOWNTIME ALERT: INTLAPP;STARTED; Host has entered a period of scheduled downtime
[1493040603] HOST DOWNTIME ALERT: INTLDB;STARTED; Host has entered a period of scheduled downtime
[1493040603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1493040603] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1493040603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1493040603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1493040603] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1493040603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1493040603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1493040603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1493040603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1493040603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1493040603] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1493040862] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5043 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493040922] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493041157] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5548 found since last run. oneLevelSearchOps OK. searchOps is too high! 5723 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493041217] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493041817] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5525 found since last run. oneLevelSearchOps OK. searchOps is too high! 5757 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493041877] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493042324] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1493042338] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1493042384] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1493042398] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1493042444] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1493042458] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1493042501] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1493042501] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1493042504] SERVICE ALERT: CAVA2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1493042561] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1493042561] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1493042564] SERVICE ALERT: CAVA2;CPU Load;OK;SOFT;2;CPU Load 0% (15 min average)
[1493042621] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1493042622] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1493042742] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1493042758] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1493042777] SERVICE ALERT: hallefirst3;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1493042802] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1493042818] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1493042832] HOST ALERT: hallefirst1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1493042838] SERVICE ALERT: hallefirst3;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1493042842] SERVICE ALERT: hallefirst3;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1493042862] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1493042866] HOST ALERT: hallefirst1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.55 ms
[1493042878] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1493042897] SERVICE ALERT: hallefirst3;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1493042902] SERVICE ALERT: hallefirst3;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1493042962] SERVICE ALERT: hallefirst3;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1493043038] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 40% - OK!
[1493043053] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1493043077] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6127 found since last run. oneLevelSearchOps OK. searchOps is too high! 6278 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493043137] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493043215] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 32% - OK!
[1493043215] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Ready
[1493043456] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;Ready
[1493043737] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6557 found since last run. oneLevelSearchOps OK. searchOps is too high! 6711 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493043797] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493044397] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6190 found since last run. oneLevelSearchOps OK. searchOps is too high! 6598 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493044457] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493045057] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7248 found since last run. oneLevelSearchOps OK. searchOps is too high! 7549 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493045117] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493045122] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5476 found since last run. oneLevelSearchOps OK. searchOps is too high! 5761 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493045182] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493045782] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 11299 found since last run. oneLevelSearchOps OK. searchOps is too high! 11610 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493045842] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493046071] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1493046131] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1493046174] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 99.90 Gb - used: 90.23 Gb (90%) - free 9.67 Gb (10%)
[1493046191] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1493046214] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;WARNING;SOFT;1;Image Transfer Kit HP CE249A is at 2% - WARNING!
[1493046234] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 99.90 Gb - used: 90.23 Gb (90%) - free 9.67 Gb (10%)
[1493046273] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;WARNING;SOFT;2;Image Transfer Kit HP CE249A is at 2% - WARNING!
[1493046294] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 99.90 Gb - used: 90.23 Gb (90%) - free 9.67 Gb (10%)
[1493046324] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1493046333] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;WARNING;HARD;3;Image Transfer Kit HP CE249A is at 2% - WARNING!
[1493046384] SERVICE ALERT: halleground2;Device Status;OK;SOFT;2;Close door.
[1493047042] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7571 found since last run. oneLevelSearchOps OK. searchOps is too high! 7860 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493047102] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493047702] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5873 found since last run. oneLevelSearchOps OK. searchOps is too high! 6168 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493047762] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493048117] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5530 found since last run. oneLevelSearchOps OK. searchOps is too high! 5958 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493048177] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493048362] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7156 found since last run. oneLevelSearchOps OK. searchOps is too high! 7388 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493048422] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493049022] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5139 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493049082] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493049377] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9600 found since last run. oneLevelSearchOps OK. searchOps is too high! 9678 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493049437] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493049613] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1493049682] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7592 found since last run. oneLevelSearchOps OK. searchOps is too high! 7782 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493049742] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493049763] SERVICE ALERT: CAVA2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1493049823] SERVICE ALERT: CAVA2;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1493049883] SERVICE ALERT: CAVA2;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1493050213] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1493050273] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 89% (15 min average)
[1493050483] SERVICE ALERT: CAVA2;CPU Load;OK;HARD;3;CPU Load 5% (15 min average)
[1493050591] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;CRITICAL;SOFT;1;Black Cartridge HP CE260X is at CRITICAL level!
[1493050637] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6106 found since last run. oneLevelSearchOps OK. searchOps is too high! 6418 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493050652] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;CRITICAL;SOFT;2;Black Cartridge HP CE260X is at CRITICAL level!
[1493050697] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493050711] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;CRITICAL;HARD;3;Black Cartridge HP CE260X is at CRITICAL level!
[1493050790] HOST ALERT: hallefirst3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1493050873] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1493050880] HOST ALERT: hallefirst3;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1493050933] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1493050942] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5844 found since last run. oneLevelSearchOps OK. searchOps is too high! 6014 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493050970] HOST ALERT: hallefirst3;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1493050993] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1493051002] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493051060] HOST ALERT: hallefirst3;UP;SOFT;4;PING WARNING - Packet loss = 96%, RTA = 5.97 ms
[1493051184] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1493051245] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1493051271] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 60% - OK!
[1493051290] SERVICE ALERT: hallefirst3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1493051297] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8603 found since last run. oneLevelSearchOps OK. searchOps is too high! 8963 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493051310] SERVICE ALERT: halleground2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1493051340] HOST ALERT: halleground2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1493051354] SERVICE ALERT: hallefirst3;Device Status;OK;HARD;3;Ready
[1493051357] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493051387] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1493051403] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1493051403] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1493051403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1493051403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1493051403] HOST DOWNTIME ALERT: INTLAPP;STARTED; Host has entered a period of scheduled downtime
[1493051403] HOST DOWNTIME ALERT: INTLDB;STARTED; Host has entered a period of scheduled downtime
[1493051403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1493051403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1493051403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1493051403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1493051403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1493051403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1493051403] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1493051403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1493051403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1493051403] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1493051403] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1493051443] HOST ALERT: halleground2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1493051447] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1493051477] HOST ALERT: halleground2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1493051507] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1493051527] HOST ALERT: halleground2;UP;SOFT;4;PING OK - Packet loss = 76%, RTA = 7.10 ms
[1493051590] SERVICE ALERT: hallesecondsouth;Device Status;OK;HARD;3;TRAY 3 EMPTY ANY TYPE LETTER   Ready
[1493051602] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7338 found since last run. oneLevelSearchOps OK. searchOps is too high! 7501 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493051662] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493051903] SERVICE ALERT: halleground2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1493051910] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;OK;HARD;3;Black Cartridge HP CE260X is at 100% - OK!
[1493051957] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6618 found since last run. oneLevelSearchOps OK. searchOps is too high! 6992 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493052017] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493052102] SERVICE ALERT: halleground2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1493052617] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5518 found since last run. oneLevelSearchOps OK. searchOps is too high! 5930 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493052677] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493052793] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1493052894] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;OK;HARD;3;c:\ - total: 99.90 Gb - used: 89.87 Gb (90%) - free 10.03 Gb (10%)
[1493053393] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1493053901] SERVICE ALERT: hallegroundcolor;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 1 Trays Empty
[1493053960] SERVICE ALERT: hallegroundcolor;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 1 Trays Empty
[1493053993] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 71% (15 min average)
[1493053993] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1493054020] SERVICE ALERT: hallegroundcolor;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 1 Trays Empty
[1493054062] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7552 found since last run. oneLevelSearchOps OK. searchOps is too high! 7596 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493054122] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493054650] HOST ALERT: hallesecondsouth;DOWN;SOFT;1;PING CRITICAL - Packet loss = 64%, RTA = 10879.78 ms
[1493054714] HOST ALERT: hallesecondsouth;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.72 ms
[1493054955] SERVICE ALERT: hallefirst3;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1493055014] SERVICE ALERT: hallefirst3;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1493055074] SERVICE ALERT: hallefirst3;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1493055220] SERVICE ALERT: hallegroundcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1493055322] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7993 found since last run. oneLevelSearchOps OK. searchOps is too high! 8042 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493055382] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493055793] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1493055853] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1493055913] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1493055982] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6548 found since last run. oneLevelSearchOps OK. searchOps is too high! 6576 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493056042] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493056513] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1493056877] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5757 found since last run. oneLevelSearchOps OK. searchOps is too high! 6411 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493056937] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493057537] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5654 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493057597] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493058197] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5121 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493058257] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493058442] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6686 found since last run. oneLevelSearchOps OK. searchOps is too high! 6742 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493058502] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493058800] SERVICE ALERT: owen101lab3;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1493058857] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5124 found since last run. oneLevelSearchOps OK. searchOps is too high! 5545 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493058859] SERVICE ALERT: owen101lab3;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1493058917] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493058919] SERVICE ALERT: owen101lab3;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1493059517] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6068 found since last run. oneLevelSearchOps OK. searchOps is too high! 6521 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493059577] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493059702] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5971 found since last run. oneLevelSearchOps OK. searchOps is too high! 5992 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493059762] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493060177] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6357 found since last run. oneLevelSearchOps OK. searchOps is too high! 6599 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493060237] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493060362] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5086 found since last run. oneLevelSearchOps OK. searchOps is too high! 5098 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493060422] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493061074] SERVICE ALERT: hallefirst3;Device Status;OK;HARD;3;Clearing paper path
[1493061622] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5171 found since last run. oneLevelSearchOps OK. searchOps is too high! 5181 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493061682] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493061913] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1493061973] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1493062033] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1493062381] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1493062440] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1493062501] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1493062633] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1493063833] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1493063837] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6651 found since last run. oneLevelSearchOps OK. searchOps is too high! 6880 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493063897] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493064318] SERVICE ALERT: owen101lab3;Device Status;OK;HARD;3;Clearing paper path
[1493064433] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1493064497] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5157 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493064557] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493065033] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 72% (15 min average)
[1493065157] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5274 found since last run. oneLevelSearchOps OK. searchOps is too high! 5644 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493065217] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493065817] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6621 found since last run. oneLevelSearchOps OK. searchOps is too high! 6932 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493065877] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493066117] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;1;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1493066237] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;2;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1493066357] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;HARD;3;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1493066482] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5420 found since last run. oneLevelSearchOps OK. searchOps is too high! 5442 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493066542] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493067142] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5922 found since last run. oneLevelSearchOps OK. searchOps is too high! 5957 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493067202] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493067677] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7013 found since last run. oneLevelSearchOps OK. searchOps is too high! 7281 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493067737] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493068033] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1493068093] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1493068153] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1493068337] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5819 found since last run. oneLevelSearchOps OK. searchOps is too high! 6124 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493068397] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493068583] SERVICE ALERT: sillpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1493068642] SERVICE ALERT: sillpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1493068703] SERVICE ALERT: sillpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1493068753] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1493068783] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1493068843] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1493068903] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1493068997] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6893 found since last run. oneLevelSearchOps OK. searchOps is too high! 7127 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493069057] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493069657] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5258 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493069717] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493069953] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1493070553] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1493070613] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1493070673] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1493070703] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;BAD DUPLEXER CONNECTION 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1493071103] SERVICE ALERT: halleground2;Device Status;OK;HARD;3;59.40 ERROR  To continue turn off then on 59.40 ERROR  46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1493071306] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1493071336] HOST ALERT: halleground2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1493071349] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1493071367] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1493071379] HOST ALERT: halleground2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1493071409] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1493071426] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1493071439] HOST ALERT: halleground2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1493071469] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1493071499] HOST ALERT: halleground2;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1493071589] HOST ALERT: halleground2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1493071679] HOST ALERT: halleground2;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1493071709] SERVICE ALERT: halleground2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1493071739] HOST ALERT: halleground2;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1493071769] SERVICE ALERT: halleground2;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1493071799] HOST ALERT: halleground2;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1493071829] SERVICE ALERT: halleground2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1493071859] HOST ALERT: halleground2;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1493071949] HOST ALERT: halleground2;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1493072602] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5309 found since last run. oneLevelSearchOps OK. searchOps is too high! 5395 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493072662] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493073073] SERVICE ALERT: hallefirst3;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1493073073] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (15 min average)
[1493073133] SERVICE ALERT: hallefirst3;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1493073193] SERVICE ALERT: hallefirst3;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1493073262] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5040 found since last run. oneLevelSearchOps OK. searchOps is too high! 5087 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493073322] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493073557] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;OK;HARD;3;OK: No entries found
[1493073673] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1493073733] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1493073793] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 92% (15 min average)
[1493074701] SERVICE ALERT: sillpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1493074903] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1493074964] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1493075023] SERVICE ALERT: sillpk;Device Status;OK;SOFT;3;RETRIEVE JOB  NO STORED JOBS
[1493075118] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5100 found since last run. oneLevelSearchOps OK. searchOps is too high! 5229 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1493075177] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493075456] SERVICE ALERT: sillpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1493075516] SERVICE ALERT: sillpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1493075576] SERVICE ALERT: sillpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1493077393] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1493077453] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1493077513] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1493078113] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 52% (15 min average)
[1493078713] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1493078773] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1493078833] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1493080394] SERVICE ALERT: hallefirst3;Device Status;OK;HARD;3;Ready
[1493080633] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1493081833] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1493081893] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 65% (15 min average)
[1493082310] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;SOFT;1;e:\ - total: 949.87 Gb - used: 764.03 Gb (80%) - free 185.84 Gb (20%)
[1493082370] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;SOFT;2;e:\ - total: 949.87 Gb - used: 764.03 Gb (80%) - free 185.84 Gb (20%)
[1493082430] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;HARD;3;e:\ - total: 949.87 Gb - used: 764.03 Gb (80%) - free 185.84 Gb (20%)
[1493083093] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1493083153] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 75% (15 min average)
[1493083630] SERVICE ALERT: AD-RMAD;E:\ Drive Space;OK;HARD;3;e:\ - total: 949.87 Gb - used: 733.15 Gb (77%) - free 216.72 Gb (23%)
[1493083800] HOST DOWNTIME ALERT: IMAGES;STARTED; Host has entered a period of scheduled downtime
[1493084744] SERVICE ALERT: IMAGES;MDT Monitor;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1493084774] HOST ALERT: IMAGES;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1493084830] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;SOFT;1;e:\ - total: 949.87 Gb - used: 764.03 Gb (80%) - free 185.84 Gb (20%)
[1493084864] HOST ALERT: IMAGES;DOWN;HARD;2;PING CRITICAL - Packet loss = 100%
[1493084890] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;SOFT;2;e:\ - total: 949.87 Gb - used: 764.03 Gb (80%) - free 185.84 Gb (20%)
[1493084901] HOST ALERT: IMAGES;UP;HARD;2;PING OK - Packet loss = 0%, RTA = 0.20 ms
[1493084950] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;HARD;3;e:\ - total: 949.87 Gb - used: 764.03 Gb (80%) - free 185.84 Gb (20%)
[1493085193] SERVICE ALERT: hallefirst3;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1493085253] SERVICE ALERT: hallefirst3;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1493085313] SERVICE ALERT: hallefirst3;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1493085324] SERVICE ALERT: IMAGES;MDT Monitor;UNKNOWN;HARD;3;Failed to open service MDT_Monitor: 1060: The specified service does not exist as an installed service.
[1493087350] SERVICE ALERT: AD-RMAD;E:\ Drive Space;OK;HARD;3;e:\ - total: 949.87 Gb - used: 733.15 Gb (77%) - free 216.72 Gb (23%)
[1493087350] SERVICE FLAPPING ALERT: AD-RMAD;E:\ Drive Space;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1493089153] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1493089213] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1493089273] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1493091000] HOST DOWNTIME ALERT: IMAGES;STOPPED; Host has exited from a period of scheduled downtime
[1493091673] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 61% (15 min average)
[1493092091] SERVICE ALERT: hallefirst3;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1493092151] SERVICE ALERT: hallefirst3;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1493092211] SERVICE ALERT: hallefirst3;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1493092273] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1493092333] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 90% (15 min average)