[1462510040] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1462510160] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 82% (15 min average)
[1462512560] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1462512680] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 88% (15 min average)
[1462513280] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462513400] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 81% (15 min average)
[1462514000] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462514120] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1462514240] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 84% (15 min average)
[1462514840] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462514960] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462515080] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1462516280] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1462517368] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1462517481] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462517600] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462517720] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1462518568] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462518688] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462518808] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 42% (15 min average)
[1462518920] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1462520008] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462520120] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1462520128] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 83% (15 min average)
[1462520240] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462520360] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 91% (15 min average)
[1462520728] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462520848] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462520960] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462520968] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 60% (15 min average)
[1462521081] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1462521200] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 93% (15 min average)
[1462521800] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462521920] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1462522040] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 87% (15 min average)
[1462522640] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1462522761] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1462522880] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1462523480] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1462523968] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1462524080] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462524199] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462524318] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1462525768] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1462525888] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 45% (15 min average)
[1462526118] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 90% (15 min average)
[1462526718] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1462526838] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1462526958] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 88% (15 min average)
[1462527558] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1462527678] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462527688] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462527798] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1462527808] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462527928] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1462528998] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462530197] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1462530328] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 66% (15 min average)
[1462530797] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462530917] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 71% (15 min average)
[1462530928] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462531048] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1462531168] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462531768] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1462531768] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1462534168] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462534288] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 77% (15 min average)
[1462537075] SERVICE ALERT: NETIDADPROXY;Service: Windows Time;OK;HARD;3;W32Time: Started
[1462537888] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1462538008] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 31% (15 min average)
[1462538972] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1462539033] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1462539093] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1462539208] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462539328] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 71% (15 min average)
[1462539928] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462540048] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1462541117] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1462541248] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462541368] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462541488] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462541717] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1462541837] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 75% (15 min average)
[1462542088] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 56% (15 min average)
[1462543287] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462543407] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 92% (15 min average)
[1462544007] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462544127] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462544234] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1462544247] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1462544354] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462544474] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1462546047] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 33% (15 min average)
[1462546274] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 84% (15 min average)
[1462546874] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462546994] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1462548194] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462548314] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462548434] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1462549034] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 82% (15 min average)
[1462549034] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1462550696] HOST ALERT: crossroadspk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462550786] HOST ALERT: crossroadspk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1462550813] SERVICE ALERT: crossroadspk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1462550843] HOST ALERT: crossroadspk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1462550873] SERVICE ALERT: crossroadspk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1462550903] HOST ALERT: crossroadspk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1462550911] SERVICE ALERT: crossroadspk;Black Cartridge HP CC364X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1462550933] SERVICE ALERT: crossroadspk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1462550941] HOST ALERT: crossroadspk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1462550969] SERVICE ALERT: crossroadspk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1462550971] SERVICE ALERT: crossroadspk;Black Cartridge HP CC364X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1462550999] HOST ALERT: crossroadspk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1462551029] SERVICE ALERT: crossroadspk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1462551031] SERVICE ALERT: crossroadspk;Black Cartridge HP CC364X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1462551059] HOST ALERT: crossroadspk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1462551089] SERVICE ALERT: crossroadspk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1462551092] SERVICE ALERT: marshallpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1462551119] HOST ALERT: crossroadspk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1462551209] HOST ALERT: crossroadspk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1462551299] HOST ALERT: crossroadspk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1462551447] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462551567] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 90% (15 min average)
[1462552167] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462552215] SERVICE ALERT: SNAP;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1462552215] SERVICE ALERT: SNAP;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1462552218] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.40)
[1462552282] HOST ALERT: SNAP;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1462552287] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462552323] SERVICE ALERT: SNAP;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1462552325] SERVICE ALERT: SNAP;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1462552407] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462552634] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1462552754] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1462552874] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 93% (15 min average)
[1462553007] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 55% (15 min average)
[1462553425] SERVICE ALERT: crossroadspk;Black Cartridge HP CC364X;OK;HARD;3;Black Cartridge HP CC364X is at 86% - OK!
[1462553429] HOST ALERT: crossroadspk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.36 ms
[1462553483] SERVICE ALERT: crossroadspk;Device Status;OK;HARD;3;Sleep mode on
[1462553759] HOST ALERT: crossroadspk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462553849] HOST ALERT: crossroadspk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1462553939] HOST ALERT: crossroadspk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1462554029] HOST ALERT: crossroadspk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1462554031] SERVICE ALERT: crossroadspk;Black Cartridge HP CC364X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1462554061] HOST ALERT: crossroadspk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1462554089] SERVICE ALERT: crossroadspk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1462554091] SERVICE ALERT: crossroadspk;Black Cartridge HP CC364X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1462554119] HOST ALERT: crossroadspk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1462554149] SERVICE ALERT: crossroadspk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1462554151] SERVICE ALERT: crossroadspk;Black Cartridge HP CC364X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1462554179] HOST ALERT: crossroadspk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1462554209] SERVICE ALERT: crossroadspk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1462554239] HOST ALERT: crossroadspk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1462554329] HOST ALERT: crossroadspk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1462554419] HOST ALERT: crossroadspk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1462555274] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1462555345] SERVICE ALERT: crossroadspk;Black Cartridge HP CC364X;OK;HARD;3;Black Cartridge HP CC364X is at 86% - OK!
[1462555345] SERVICE FLAPPING ALERT: crossroadspk;Black Cartridge HP CC364X;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1462555349] HOST ALERT: crossroadspk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1462555394] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1462555403] SERVICE ALERT: crossroadspk;Device Status;OK;HARD;3;Sleep mode on
[1462555403] SERVICE FLAPPING ALERT: crossroadspk;Device Status;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1462555407] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1462555527] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1462555647] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 90% (15 min average)
[1462555727] SERVICE ALERT: crossroadspk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1462557447] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1462557567] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 18% (15 min average)
[1462557794] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462557914] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1462558034] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1462558535] SERVICE ALERT: SNAP;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1462558538] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.40)
[1462558602] HOST ALERT: SNAP;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.44 ms
[1462558634] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 84% (15 min average)
[1462558655] SERVICE ALERT: SNAP;Service: Firewall;CRITICAL;SOFT;2;connect to address 164.76.102.40 and port 12489: No route to host
[1462558658] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.40)
[1462558658] SERVICE ALERT: SNAP;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1462558763] SERVICE ALERT: SNAP;Service: Firewall;OK;SOFT;3;MpsSvc: Started
[1462558767] SERVICE ALERT: SNAP;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1462558767] HOST ALERT: SNAP;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.44 ms
[1462558767] HOST FLAPPING ALERT: SNAP;STARTED; Host appears to have started flapping (22.5% change > 20.0% threshold)
[1462559497] HOST FLAPPING ALERT: SNAP;STOPPED; Host appears to have stopped flapping (4.1% change < 5.0% threshold)
[1462559685] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.40)
[1462559749] HOST ALERT: SNAP;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.45 ms
[1462561489] SERVICE ALERT: ADWSUSSERVERS;Service: Group Policy;CRITICAL;SOFT;1;gpsvc: Stopped
[1462561538] SERVICE ALERT: ADWSUSSERVERS;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462561609] SERVICE ALERT: ADWSUSSERVERS;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1462561638] SERVICE ALERT: ADWSUSSERVERS;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 2 minute(s)
[1462564634] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462564754] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 82% (15 min average)
[1462564767] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1462565954] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1462565967] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462566074] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 84% (15 min average)
[1462566087] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1462566144] SERVICE FLAPPING ALERT: crossroadspk;Black Cartridge HP CC364X;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1462566203] SERVICE FLAPPING ALERT: crossroadspk;Device Status;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1462566207] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 77% (15 min average)
[1462568007] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462568127] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462568247] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462568474] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462568594] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 66% (15 min average)
[1462569194] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1462569314] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 86% (15 min average)
[1462569914] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1462570034] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1462570154] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 84% (15 min average)
[1462570754] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462570874] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1462570994] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1462571247] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 78% (15 min average)
[1462571594] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 74% (15 min average)
[1462571847] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462571967] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462572087] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1462573287] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1462573287] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1462573887] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462574007] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1462574127] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462575927] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1462576527] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462576647] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462576767] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462578793] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462578913] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 77% (15 min average)
[1462579167] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1462579767] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1462579887] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462580007] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462581807] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 74% (15 min average)
[1462583007] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462583113] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1462583127] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462583247] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1462587793] SERVICE ALERT: LANSWEEPER;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462587893] SERVICE ALERT: LANSWEEPER;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1462588047] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462588477] SERVICE ALERT: INTLTESTAPP;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462588577] SERVICE ALERT: INTLTESTAPP;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1462588647] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462589713] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1462589833] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1462589953] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462590553] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1462591047] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462591153] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1462591645] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1462591753] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462591873] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 91% (15 min average)
[1462592245] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462592365] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462592485] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462593085] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)