[1481259908] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;CRITICAL;HARD;3;Yellow Cartridge HP CE262A is at 1% - CRITICAL!
[1481259933] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1481260053] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1481260173] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 86% (15 min average)
[1481260773] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1481260893] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1481261013] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 82% (15 min average)
[1481263086] SERVICE ALERT: SP2010;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1481263206] SERVICE ALERT: SP2010;CPU Load;OK;SOFT;2;CPU Load 4% (15 min average)
[1481263413] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1481263533] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1481263653] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1481264153] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1481264274] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1481264873] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1481264993] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1481265000] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1481265453] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 85% (15 min average)
[1481265594] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1481265713] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;2;Warning:  idm1  are more than 120 seconds behind!
[1481265833] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;3;OK. All replicas within 120 seconds.
[1481266053] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1481266173] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1481266433] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1481266553] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1481266673] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1481267373] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1481267493] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1481267613] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1481267873] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1481267890] SERVICE ALERT: FITNESSGRAM;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1481268010] SERVICE ALERT: FITNESSGRAM;CPU Load;OK;SOFT;2;CPU Load 0% (15 min average)
[1481268473] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1481268813] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1481269674] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1481269674] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1481270273] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1481270873] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1481271213] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1481271473] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1481271593] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;2;Warning:  idm1  are more than 120 seconds behind!
[1481271713] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;3;OK. All replicas within 120 seconds.
[1481271813] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1481271933] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1481272053] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1481272289] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1481272313] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1481272409] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1481272433] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1481272529] SERVICE ALERT: epprogsql2;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1481272553] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1481273129] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1481273153] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1481273753] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1481274353] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1481274473] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1481275053] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 65% (15 min average)
[1481275491] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 125 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481275611] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481275653] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1481275773] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 42% (15 min average)
[1481275802] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (34.1% change >= 20.0% threshold)
[1481275802] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (46.6% change >= 20.0% threshold)
[1481275802] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1481275802] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1481275802] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1481275802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1481275802] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1481275802] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1481275802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1481275802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1481275802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1481275802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1481275802] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1481275802] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1481275802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1481275802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1481275802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1481275802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1481275802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1481276129] SERVICE ALERT: epprogsql2;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1481276729] SERVICE ALERT: epprogsql2;CPU Load;OK;HARD;3;CPU Load 23% (15 min average)
[1481276729] SERVICE FLAPPING ALERT: epprogsql2;CPU Load;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1481276973] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1481277093] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1481277213] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 71% (15 min average)
[1481277602] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (28.8% change >= 20.0% threshold)
[1481277602] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (43.1% change >= 20.0% threshold)
[1481277602] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1481277602] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1481277602] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1481277602] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1481277602] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1481277602] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1481277602] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1481277602] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1481277602] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1481277602] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1481277602] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1481277602] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1481277602] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1481277602] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1481277602] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1481277602] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1481277602] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1481277813] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1481277933] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1481278053] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 44% (15 min average)
[1481281053] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1481281173] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1481281293] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1481283000] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service has exited from a period of scheduled downtime
[1481284893] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1481285273] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1481285493] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1481288493] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1481288810] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 108 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481288930] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481289052] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1481289093] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1481289113] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 31% - OK!
[1481290093] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 124 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481290130] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 140 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481290213] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481290250] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481290850] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 141 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481290970] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481291493] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (15 min average)
[1481291587] HOST ALERT: santa;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1481291651] HOST ALERT: santa;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 7.09 ms
[1481292013] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5761 found since last run. oneLevelSearchOps OK. searchOps is too high! 5869 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481292093] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1481292133] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481292213] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1481292333] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1481292379] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481292439] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481292499] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1481293099] SERVICE ALERT: markjeffersonpk;Device Status;OK;HARD;3;Clearing paper path
[1481293533] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1481293802] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (37.6% change >= 20.0% threshold)
[1481293802] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1481293802] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1481293802] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1481293802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1481293802] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1481293802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1481293802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1481293802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1481293802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1481293802] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1481293802] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1481293802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1481293802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1481293802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1481293802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1481293802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1481293970] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5112 found since last run. oneLevelSearchOps OK. searchOps is too high! 5392 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481294090] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481294133] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1481294253] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 67% (15 min average)
[1481295453] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1481295573] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1481295693] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1481295890] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5344 found since last run. oneLevelSearchOps OK. searchOps is too high! 5521 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481296010] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481296168] SERVICE ALERT: marshallpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.137.49. Is the printer off?  Is the network down?
[1481296198] HOST ALERT: marshallpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1481296222] SERVICE ALERT: marshallpk;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 68% - OK!
[1481296226] HOST ALERT: marshallpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.82 ms
[1481296293] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1481296333] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5265 found since last run. oneLevelSearchOps OK. searchOps is too high! 5436 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481296453] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481296610] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5025 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481296730] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481296893] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1481297493] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1481297613] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 65% (15 min average)
[1481299130] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7814 found since last run. oneLevelSearchOps OK. searchOps is too high! 8096 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481299250] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481299345] SERVICE ALERT: rooseveltpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1481299413] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1481299533] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1481299653] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 94% (15 min average)
[1481300253] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1481300373] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 60% (15 min average)
[1481300973] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1481301093] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1481301213] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1481301813] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1481302249] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5137 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481302369] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481302413] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1481302533] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1481302653] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1481302969] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5974 found since last run. oneLevelSearchOps OK. searchOps is too high! 6470 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481302974] SERVICE ALERT: CAVA2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1481303089] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481303094] SERVICE ALERT: CAVA2;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1481303213] SERVICE ALERT: CAVA2;CPU Load;OK;SOFT;3;CPU Load 1% (15 min average)
[1481303689] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7583 found since last run. oneLevelSearchOps OK. searchOps is too high! 7789 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481303808] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481304253] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8324 found since last run. oneLevelSearchOps OK. searchOps is too high! 8356 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481304373] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481304513] SERVICE ALERT: SNAP;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1481304513] SERVICE ALERT: SNAP;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1481304515] SERVICE ALERT: SNAP;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1481304519] SERVICE ALERT: SNAP;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1481304603] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (43.0% change >= 20.0% threshold)
[1481304603] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1481304603] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1481304603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1481304603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1481304603] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1481304603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1481304603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1481304603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1481304603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1481304603] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1481304603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1481304603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1481304603] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1481304603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1481304603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1481304603] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1481304613] SERVICE ALERT: SNAP;Uptime;OK;SOFT;2;System Uptime - 6 day(s) 17 hour(s) 32 minute(s)
[1481304613] SERVICE ALERT: SNAP;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1481304615] SERVICE ALERT: SNAP;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1481304619] SERVICE ALERT: SNAP;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1481304973] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6436 found since last run. oneLevelSearchOps OK. searchOps is too high! 6474 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481305093] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481305099] SERVICE ALERT: SNAP;Service: User Profile;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1481305102] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.138)
[1481305137] HOST ALERT: SNAP;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.44 ms
[1481305207] SERVICE ALERT: SNAP;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1481305443] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.138)
[1481305451] SERVICE ALERT: SNAP;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1481305466] SERVICE ALERT: SNAP;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1481305469] HOST ALERT: SNAP;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.4.138)
[1481305533] HOST ALERT: SNAP;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.49 ms
[1481305533] HOST FLAPPING ALERT: SNAP;STARTED; Host appears to have started flapping (20.3% change > 20.0% threshold)
[1481305551] SERVICE ALERT: SNAP;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1481305583] SERVICE ALERT: SNAP;C:\ Drive Space;OK;SOFT;2;c:\ - total: 1116.65 Gb - used: 37.61 Gb (3%) - free 1079.04 Gb (97%)
[1481305693] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6438 found since last run. oneLevelSearchOps OK. searchOps is too high! 6481 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481305813] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481306150] HOST FLAPPING ALERT: SNAP;STOPPED; Host appears to have stopped flapping (4.1% change < 5.0% threshold)
[1481306413] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5790 found since last run. oneLevelSearchOps OK. searchOps is too high! 5834 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481306533] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481306756] SERVICE ALERT: SNAP;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1481306756] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.138)
[1481306787] HOST ALERT: SNAP;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1481306808] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 102 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481306853] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 86% (15 min average)
[1481306870] SERVICE ALERT: SNAP;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1481306928] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481306945] SERVICE ALERT: SNAP;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1481306948] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.138)
[1481306984] HOST ALERT: SNAP;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.40 ms
[1481306984] HOST FLAPPING ALERT: SNAP;STARTED; Host appears to have started flapping (22.8% change > 20.0% threshold)
[1481307053] SERVICE ALERT: SNAP;Service: RPC;OK;SOFT;2;RpcSs: Started
[1481307133] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5623 found since last run. oneLevelSearchOps OK. searchOps is too high! 5661 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481307253] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481307394] SERVICE ALERT: SNAP;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1481307397] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.138)
[1481307453] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1481307460] HOST ALERT: SNAP;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.49 ms
[1481307502] SERVICE ALERT: SNAP;C:\ Drive Space;OK;SOFT;2;c:\ - total: 1116.65 Gb - used: 37.61 Gb (3%) - free 1079.04 Gb (97%)
[1481307573] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1481307693] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 90% (15 min average)
[1481308181] SERVICE ALERT: SNAP;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1481308181] HOST FLAPPING ALERT: SNAP;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1481308184] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.138)
[1481308209] SERVICE ALERT: SNAP;Service: User Profile;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1481308212] HOST ALERT: SNAP;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.4.138)
[1481308216] SERVICE ALERT: SNAP;Uptime;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1481308216] SERVICE ALERT: SNAP;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1481308216] SERVICE ALERT: SNAP;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1481308221] SERVICE ALERT: SNAP;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1481308255] SERVICE ALERT: SNAP;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1481308258] HOST ALERT: SNAP;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.4.138)
[1481308266] SERVICE ALERT: SNAP;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1481308266] SERVICE ALERT: SNAP;CPU Load;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1481308293] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1481308298] SERVICE ALERT: SNAP;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1481308302] HOST ALERT: SNAP;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1481308346] SERVICE ALERT: SNAP;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1481308352] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.138)
[1481308353] SERVICE ALERT: SNAP;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1481308353] SERVICE ALERT: SNAP;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1481308353] SERVICE ALERT: SNAP;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1481308356] SERVICE ALERT: SNAP;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1481308358] SERVICE ALERT: SNAP;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1481308375] SERVICE ALERT: SNAP;Service: RPC;CRITICAL;SOFT;2;connect to address 164.76.4.138 and port 12489: No route to host
[1481308378] HOST ALERT: SNAP;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.4.138)
[1481308386] SERVICE ALERT: SNAP;Service: Group Policy;CRITICAL;SOFT;2;connect to address 164.76.4.138 and port 12489: No route to host
[1481308386] SERVICE ALERT: SNAP;CPU Load;CRITICAL;SOFT;2;connect to address 164.76.4.138 and port 12489: No route to host
[1481308413] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1481308446] SERVICE ALERT: SNAP;Service: User Profile;OK;SOFT;3;ProfSvc: Started
[1481308450] HOST ALERT: SNAP;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.39 ms
[1481308453] SERVICE ALERT: SNAP;Uptime;OK;SOFT;3;System Uptime - 6 day(s) 18 hour(s) 36 minute(s)
[1481308453] SERVICE ALERT: SNAP;Service: Windows Time;OK;SOFT;3;W32Time: Started
[1481308453] SERVICE ALERT: SNAP;Service: DNS Client;OK;SOFT;3;Dnscache: Started
[1481308458] SERVICE ALERT: SNAP;Service: WMI;OK;SOFT;3;Winmgmt: Started
[1481308464] SERVICE ALERT: SNAP;Memory Usage;OK;SOFT;2;Memory usage: total:16362.75 MB - used: 1863.25 MB (11%) - free: 14499.50 MB (89%)
[1481308492] SERVICE ALERT: SNAP;Service: RPC;OK;SOFT;3;RpcSs: Started
[1481308503] SERVICE ALERT: SNAP;Service: Group Policy;OK;SOFT;3;gpsvc: Started
[1481308503] SERVICE ALERT: SNAP;CPU Load;OK;SOFT;3;CPU Load 0% (15 min average)
[1481308533] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1481308728] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5019 found since last run. oneLevelSearchOps OK. searchOps is too high! 5297 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481308848] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481309448] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5802 found since last run. oneLevelSearchOps OK. searchOps is too high! 5994 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481309568] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481309733] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 91% (15 min average)
[1481310168] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5520 found since last run. oneLevelSearchOps OK. searchOps is too high! 5893 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481310288] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481310333] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1481310453] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1481310573] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1481310888] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5090 found since last run. oneLevelSearchOps OK. searchOps is too high! 5342 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481311008] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481311453] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6269 found since last run. oneLevelSearchOps OK. searchOps is too high! 6327 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481311573] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481311773] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 79% (15 min average)
[1481312208] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5679 found since last run. oneLevelSearchOps OK. searchOps is too high! 5953 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481312328] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481312973] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1481313093] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 92% (15 min average)
[1481313373] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5123 found since last run. oneLevelSearchOps OK. searchOps is too high! 5165 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481313493] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481313528] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5207 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481313648] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481313693] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1481313813] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1481313933] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1481314533] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1481315133] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1481315253] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1481315293] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5836 found since last run. oneLevelSearchOps OK. searchOps is too high! 5950 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481315413] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481315853] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1481315973] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1481316093] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1481317213] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5989 found since last run. oneLevelSearchOps OK. searchOps is too high! 6101 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481317333] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481317848] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5907 found since last run. oneLevelSearchOps OK. searchOps is too high! 6160 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481317933] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7862 found since last run. oneLevelSearchOps OK. searchOps is too high! 7964 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481317968] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481318053] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481318653] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7664 found since last run. oneLevelSearchOps OK. searchOps is too high! 7765 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481318773] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481319093] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 49% (15 min average)
[1481319373] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5516 found since last run. oneLevelSearchOps OK. searchOps is too high! 5658 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481319493] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481319693] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1481319813] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1481319933] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1481320533] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1481320968] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5158 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481321088] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481321133] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1481321689] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5250 found since last run. oneLevelSearchOps OK. searchOps is too high! 5435 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481321808] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481323533] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1481324733] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1481326693] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5227 found since last run. oneLevelSearchOps OK. searchOps is too high! 5318 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481326813] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481327733] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1481328013] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5059 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481328133] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481328333] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1481328733] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5056 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481328853] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481328933] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 74% (15 min average)
[1481329453] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5019 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481329573] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481330133] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1481330253] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 92% (15 min average)
[1481330853] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1481330973] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1481331093] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1481332343] SERVICE ALERT: rooseveltpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1481332404] SERVICE ALERT: rooseveltpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1481332464] SERVICE ALERT: rooseveltpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1481335673] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;SOFT;1;e:\ - total: 799.87 Gb - used: 645.77 Gb (81%) - free 154.10 Gb (19%)
[1481335793] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;SOFT;2;e:\ - total: 799.87 Gb - used: 645.77 Gb (81%) - free 154.10 Gb (19%)
[1481335815] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.138)
[1481335879] HOST ALERT: SNAP;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.49 ms
[1481335913] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;HARD;3;e:\ - total: 799.87 Gb - used: 645.77 Gb (81%) - free 154.10 Gb (19%)
[1481336057] SERVICE ALERT: SNAP;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1481336057] SERVICE ALERT: SNAP;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1481336057] SERVICE ALERT: SNAP;Service: User Profile;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1481336060] SERVICE ALERT: SNAP;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1481336060] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.138)
[1481336063] SERVICE ALERT: SNAP;Uptime;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1481336066] SERVICE ALERT: SNAP;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1481336095] HOST ALERT: SNAP;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1481336095] HOST FLAPPING ALERT: SNAP;STARTED; Host appears to have started flapping (21.7% change > 20.0% threshold)
[1481336165] SERVICE ALERT: SNAP;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1481336171] SERVICE ALERT: SNAP;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1481336171] SERVICE ALERT: SNAP;Uptime;OK;SOFT;2;System Uptime - 7 day(s) 2 hour(s) 18 minute(s)
[1481336173] SERVICE ALERT: SNAP;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1481336177] SERVICE ALERT: SNAP;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1481336183] SERVICE ALERT: SNAP;Memory Usage;OK;SOFT;2;Memory usage: total:16362.75 MB - used: 1848.41 MB (11%) - free: 14514.34 MB (89%)
[1481336703] HOST FLAPPING ALERT: SNAP;STOPPED; Host appears to have stopped flapping (4.1% change < 5.0% threshold)
[1481336880] SERVICE ALERT: SNAP;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1481336883] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.138)
[1481336904] SERVICE ALERT: SNAP;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1481336908] HOST ALERT: SNAP;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.79 ms
[1481336968] SERVICE ALERT: SNAP;Service: NetLogin;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1481336971] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.138)
[1481336988] SERVICE ALERT: SNAP;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1481336992] HOST ALERT: SNAP;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1481336992] HOST FLAPPING ALERT: SNAP;STARTED; Host appears to have started flapping (23.4% change > 20.0% threshold)
[1481337033] SERVICE ALERT: SNAP;C:\ Drive Space;CRITICAL;SOFT;2;connect to address 164.76.4.138 and port 12489: No route to host
[1481337088] SERVICE ALERT: SNAP;Service: NetLogin;CRITICAL;SOFT;2;connect to address 164.76.4.138 and port 12489: No route to host
[1481337091] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.138)
[1481337150] SERVICE ALERT: SNAP;C:\ Drive Space;OK;SOFT;3;c:\ - total: 1116.65 Gb - used: 37.61 Gb (3%) - free 1079.04 Gb (97%)
[1481337154] HOST ALERT: SNAP;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1481337196] SERVICE ALERT: SNAP;Service: NetLogin;OK;SOFT;3;Netlogon: Started
[1481337761] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.138)
[1481337761] SERVICE ALERT: SNAP;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1481337800] HOST ALERT: SNAP;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.39 ms
[1481337873] SERVICE ALERT: SNAP;C:\ Drive Space;CRITICAL;SOFT;2;connect to address 164.76.4.138 and port 12489: No route to host
[1481337876] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.138)
[1481337895] HOST ALERT: SNAP;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.40 ms
[1481337977] SERVICE ALERT: SNAP;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1481337977] SERVICE ALERT: SNAP;Service: User Profile;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1481337977] SERVICE ALERT: SNAP;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1481337980] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.138)
[1481337980] SERVICE ALERT: SNAP;Uptime;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1481337983] SERVICE ALERT: SNAP;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1481337983] SERVICE ALERT: SNAP;C:\ Drive Space;CRITICAL;HARD;3;connect to address 164.76.4.138 and port 12489: No route to host
[1481337991] SERVICE ALERT: SNAP;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1481338097] SERVICE ALERT: SNAP;Service: DNS Client;CRITICAL;SOFT;2;connect to address 164.76.4.138 and port 12489: No route to host
[1481338097] SERVICE ALERT: SNAP;Service: User Profile;CRITICAL;SOFT;2;connect to address 164.76.4.138 and port 12489: No route to host
[1481338097] SERVICE ALERT: SNAP;Uptime;CRITICAL;SOFT;2;connect to address 164.76.4.138 and port 12489: No route to host
[1481338100] HOST ALERT: SNAP;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.4.138)
[1481338103] SERVICE ALERT: SNAP;Service: WMI;CRITICAL;SOFT;2;connect to address 164.76.4.138 and port 12489: No route to host
[1481338103] SERVICE ALERT: SNAP;Service: Windows Time;CRITICAL;SOFT;2;connect to address 164.76.4.138 and port 12489: No route to host
[1481338105] SERVICE ALERT: SNAP;Memory Usage;OK;SOFT;2;Memory usage: total:16362.75 MB - used: 1892.51 MB (12%) - free: 14470.24 MB (88%)
[1481338200] SERVICE ALERT: SNAP;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1481338203] HOST ALERT: SNAP;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.4.138)
[1481338208] SERVICE ALERT: SNAP;Service: User Profile;CRITICAL;HARD;3;connect to address 164.76.4.138 and port 12489: No route to host
[1481338208] HOST FLAPPING ALERT: SNAP;STOPPED; Host appears to have stopped flapping (4.4% change < 5.0% threshold)
[1481338211] SERVICE ALERT: SNAP;Uptime;CRITICAL;HARD;3;connect to address 164.76.4.138 and port 12489: No route to host
[1481338211] SERVICE ALERT: SNAP;Service: Windows Time;CRITICAL;HARD;3;connect to address 164.76.4.138 and port 12489: No route to host
[1481338216] SERVICE ALERT: SNAP;Service: DNS Client;CRITICAL;HARD;3;connect to address 164.76.4.138 and port 12489: No route to host
[1481338221] SERVICE ALERT: SNAP;Service: WMI;OK;SOFT;3;Winmgmt: Started
[1481338222] HOST ALERT: SNAP;UP;SOFT;4;PING OK - Packet loss = 28%, RTA = 13.62 ms
[1481338308] SERVICE ALERT: SNAP;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1481338313] SERVICE ALERT: AD-RMAD;E:\ Drive Space;OK;HARD;3;e:\ - total: 799.87 Gb - used: 620.18 Gb (78%) - free 179.69 Gb (22%)
[1481338805] SERVICE ALERT: SNAP;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1481338811] SERVICE ALERT: SNAP;Uptime;OK;HARD;3;System Uptime - 7 day(s) 3 hour(s) 2 minute(s)
[1481338811] SERVICE ALERT: SNAP;Service: Windows Time;OK;HARD;3;W32Time: Started
[1481338813] SERVICE ALERT: SNAP;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1481339103] SERVICE ALERT: SNAP;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1481339103] SERVICE ALERT: SNAP;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1481339103] SERVICE ALERT: SNAP;CPU Load;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1481339106] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.138)
[1481339170] HOST ALERT: SNAP;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.47 ms
[1481339181] SERVICE ALERT: SNAP;C:\ Drive Space;OK;HARD;3;c:\ - total: 1116.65 Gb - used: 37.62 Gb (3%) - free 1079.04 Gb (97%)
[1481339211] SERVICE ALERT: SNAP;Service: RPC;OK;SOFT;2;RpcSs: Started
[1481339222] SERVICE ALERT: SNAP;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1481339222] SERVICE ALERT: SNAP;CPU Load;OK;SOFT;2;CPU Load 0% (15 min average)
[1481339417] SERVICE ALERT: SNAP;Service: User Profile;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1481339417] SERVICE ALERT: SNAP;Uptime;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1481339417] SERVICE ALERT: SNAP;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1481339420] SERVICE ALERT: SNAP;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1481339420] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.138)
[1481339423] SERVICE ALERT: SNAP;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1481339493] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 91% (15 min average)
[1481339509] SERVICE ALERT: SNAP;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1481339509] SERVICE ALERT: SNAP;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1481339512] HOST ALERT: SNAP;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.4.138)
[1481339528] SERVICE ALERT: SNAP;Service: User Profile;CRITICAL;SOFT;2;connect to address 164.76.4.138 and port 12489: No route to host
[1481339531] HOST ALERT: SNAP;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.4.138)
[1481339534] SERVICE ALERT: SNAP;Service: Windows Time;CRITICAL;SOFT;2;connect to address 164.76.4.138 and port 12489: No route to host
[1481339534] SERVICE ALERT: SNAP;Service: DNS Client;CRITICAL;SOFT;2;connect to address 164.76.4.138 and port 12489: No route to host
[1481339534] SERVICE ALERT: SNAP;Uptime;CRITICAL;SOFT;2;connect to address 164.76.4.138 and port 12489: No route to host
[1481339540] SERVICE ALERT: SNAP;Service: WMI;CRITICAL;SOFT;2;connect to address 164.76.4.138 and port 12489: No route to host
[1481339599] SERVICE ALERT: SNAP;Service: NetLogin;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1481339602] HOST ALERT: SNAP;DOWN;SOFT;4;CRITICAL - Host Unreachable (164.76.4.138)
[1481339620] SERVICE ALERT: SNAP;Service: Firewall;CRITICAL;SOFT;2;connect to address 164.76.4.138 and port 12489: No route to host
[1481339623] HOST ALERT: SNAP;DOWN;SOFT;5;CRITICAL - Host Unreachable (164.76.4.138)
[1481339629] SERVICE ALERT: SNAP;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1481339657] SERVICE ALERT: SNAP;Service: User Profile;CRITICAL;HARD;3;connect to address 164.76.4.138 and port 12489: No route to host
[1481339657] SERVICE ALERT: SNAP;Service: DNS Client;CRITICAL;HARD;3;connect to address 164.76.4.138 and port 12489: No route to host
[1481339657] SERVICE ALERT: SNAP;Uptime;CRITICAL;HARD;3;connect to address 164.76.4.138 and port 12489: No route to host
[1481339660] HOST ALERT: SNAP;DOWN;SOFT;6;CRITICAL - Host Unreachable (164.76.4.138)
[1481339663] SERVICE ALERT: SNAP;Service: Windows Time;CRITICAL;HARD;3;connect to address 164.76.4.138 and port 12489: No route to host
[1481339663] SERVICE ALERT: SNAP;Service: WMI;CRITICAL;HARD;3;connect to address 164.76.4.138 and port 12489: No route to host
[1481339716] SERVICE ALERT: SNAP;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1481339720] HOST ALERT: SNAP;UP;SOFT;7;PING OK - Packet loss = 0%, RTA = 0.40 ms
[1481339737] SERVICE ALERT: SNAP;Service: Firewall;OK;SOFT;3;MpsSvc: Started
[1481340093] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1481340213] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1481340245] SERVICE ALERT: SNAP;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1481340245] SERVICE FLAPPING ALERT: SNAP;Service: User Profile;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1481340251] SERVICE ALERT: SNAP;Uptime;OK;HARD;3;System Uptime - 7 day(s) 3 hour(s) 26 minute(s)
[1481340251] SERVICE FLAPPING ALERT: SNAP;Uptime;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1481340251] SERVICE ALERT: SNAP;Service: Windows Time;OK;HARD;3;W32Time: Started
[1481340251] SERVICE FLAPPING ALERT: SNAP;Service: Windows Time;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1481340253] SERVICE ALERT: SNAP;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1481340253] SERVICE FLAPPING ALERT: SNAP;Service: DNS Client;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1481340257] SERVICE ALERT: SNAP;Service: WMI;OK;HARD;3;Winmgmt: Started
[1481340333] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1481342040] SERVICE ALERT: SNAP;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1481342043] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.138)
[1481342046] SERVICE ALERT: SNAP;Service: User Profile;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1481342054] SERVICE ALERT: SNAP;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1481342054] SERVICE ALERT: SNAP;Uptime;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1481342054] SERVICE ALERT: SNAP;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1481342057] SERVICE ALERT: SNAP;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1481342061] HOST ALERT: SNAP;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 201.02 ms
[1481342133] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1481342148] SERVICE ALERT: SNAP;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1481342165] SERVICE ALERT: SNAP;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1481342171] SERVICE ALERT: SNAP;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1481342171] SERVICE ALERT: SNAP;Uptime;OK;SOFT;2;System Uptime - 7 day(s) 3 hour(s) 58 minute(s)
[1481342173] SERVICE ALERT: SNAP;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1481342177] SERVICE ALERT: SNAP;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1481343333] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 71% (15 min average)
[1481343933] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1481344053] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1481344173] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 94% (15 min average)
[1481345373] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1481345493] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1481345614] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)