[1458793800] HOST DOWNTIME ALERT: REFWIN;STOPPED; Host has exited from a period of scheduled downtime
[1458793800] HOST DOWNTIME ALERT: PROCARE;STOPPED; Host has exited from a period of scheduled downtime
[1458793800] HOST DOWNTIME ALERT: UHSPHARM-2012;STOPPED; Host has exited from a period of scheduled downtime
[1458793800] HOST DOWNTIME ALERT: ADWSUSCLIENTS;STOPPED; Host has exited from a period of scheduled downtime
[1458793800] HOST DOWNTIME ALERT: EMUSCRIPT;STOPPED; Host has exited from a period of scheduled downtime
[1458793800] HOST DOWNTIME ALERT: TABLEAU;STOPPED; Host has exited from a period of scheduled downtime
[1458793800] HOST DOWNTIME ALERT: SANDBOX-DB;STOPPED; Host has exited from a period of scheduled downtime
[1458793800] HOST DOWNTIME ALERT: EMUSQLEXPRESS;STOPPED; Host has exited from a period of scheduled downtime
[1458793800] HOST DOWNTIME ALERT: Echo-01;STOPPED; Host has exited from a period of scheduled downtime
[1458793800] HOST DOWNTIME ALERT: TI22-PSYCLINIC;STOPPED; Host has exited from a period of scheduled downtime
[1458793800] HOST DOWNTIME ALERT: TI22;STOPPED; Host has exited from a period of scheduled downtime
[1458793800] HOST DOWNTIME ALERT: ADWU1;STOPPED; Host has exited from a period of scheduled downtime
[1458793800] HOST DOWNTIME ALERT: PNCWEB-2012;STOPPED; Host has exited from a period of scheduled downtime
[1458793800] HOST DOWNTIME ALERT: CEDARON;STOPPED; Host has exited from a period of scheduled downtime
[1458793800] HOST DOWNTIME ALERT: CBORD;STOPPED; Host has exited from a period of scheduled downtime
[1458793800] HOST DOWNTIME ALERT: FTK;STOPPED; Host has exited from a period of scheduled downtime
[1458793800] HOST DOWNTIME ALERT: LANSWEEPER;STOPPED; Host has exited from a period of scheduled downtime
[1458793800] HOST DOWNTIME ALERT: WINKMS;STOPPED; Host has exited from a period of scheduled downtime
[1458793800] HOST DOWNTIME ALERT: XDM;STOPPED; Host has exited from a period of scheduled downtime
[1458793800] HOST DOWNTIME ALERT: HDBACKUPS;STOPPED; Host has exited from a period of scheduled downtime
[1458793800] HOST DOWNTIME ALERT: EMU-VALUE01;STOPPED; Host has exited from a period of scheduled downtime
[1458793800] HOST DOWNTIME ALERT: IMAGES;STOPPED; Host has exited from a period of scheduled downtime
[1458793800] HOST DOWNTIME ALERT: AD-RMAD;STOPPED; Host has exited from a period of scheduled downtime
[1458793800] HOST DOWNTIME ALERT: NETENG-PA;STOPPED; Host has exited from a period of scheduled downtime
[1458793800] HOST DOWNTIME ALERT: MBAMDB;STOPPED; Host has exited from a period of scheduled downtime
[1458793800] HOST DOWNTIME ALERT: MBAM;STOPPED; Host has exited from a period of scheduled downtime
[1458793800] HOST DOWNTIME ALERT: ADAUTH-RMAD;STOPPED; Host has exited from a period of scheduled downtime
[1458793800] HOST DOWNTIME ALERT: COE-WEBDAV;STOPPED; Host has exited from a period of scheduled downtime
[1458793800] HOST DOWNTIME ALERT: PNCDB-2012;STOPPED; Host has exited from a period of scheduled downtime
[1458793800] HOST DOWNTIME ALERT: DPS-SECURE02;STOPPED; Host has exited from a period of scheduled downtime
[1458793800] HOST DOWNTIME ALERT: NESSUS;STOPPED; Host has exited from a period of scheduled downtime
[1458793800] HOST DOWNTIME ALERT: PNCAPP-2012;STOPPED; Host has exited from a period of scheduled downtime
[1458793800] HOST DOWNTIME ALERT: DIRECTACCESS;STOPPED; Host has exited from a period of scheduled downtime
[1458793800] HOST DOWNTIME ALERT: OFFICEKMS;STOPPED; Host has exited from a period of scheduled downtime
[1458793800] HOST DOWNTIME ALERT: OPIE;STOPPED; Host has exited from a period of scheduled downtime
[1458799356] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 72% (5 min average)
[1458799548] SERVICE ALERT: NETENG-PA;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1458799668] SERVICE ALERT: NETENG-PA;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1458799676] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 84.90 Gb - used: 78.83 Gb (93%) - free 6.07 Gb (7%)
[1458799796] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 84.90 Gb - used: 78.83 Gb (93%) - free 6.07 Gb (7%)
[1458799916] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 84.90 Gb - used: 78.83 Gb (93%) - free 6.07 Gb (7%)
[1458799956] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1458800076] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1458800196] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1458800515] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1458800635] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 95% (5 min average)
[1458800755] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 95% (5 min average)
[1458801124] SERVICE ALERT: INTLAPP-ISSMAP;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1458801244] SERVICE ALERT: INTLAPP-ISSMAP;CPU Load;OK;SOFT;2;CPU Load 25% (5 min average)
[1458802250] SERVICE ALERT: INTLTESTAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1458802370] SERVICE ALERT: INTLTESTAPP;CPU Load;OK;SOFT;2;CPU Load 22% (5 min average)
[1458802555] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 90% (5 min average)
[1458803155] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 96% (5 min average)
[1458803755] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 80% (5 min average)
[1458806196] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 92% (5 min average)
[1458806796] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1458806796] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1458807996] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 77% (5 min average)
[1458808596] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1458808716] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 94% (5 min average)
[1458808836] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 95% (5 min average)
[1458810036] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 93% (5 min average)
[1458810636] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1458811236] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 58% (5 min average)
[1458813036] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458813156] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 85% (5 min average)
[1458814956] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458815076] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1458815130] SERVICE ALERT: INTLAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1458815155] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1458815196] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1458815250] SERVICE ALERT: INTLAPP;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1458815796] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 78% (5 min average)
[1458816396] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1458816516] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 84% (5 min average)
[1458817116] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1458817236] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1458817356] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1458817956] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 82% (5 min average)
[1458818556] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1458818676] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1458818796] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1458820605] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1458820725] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1458820846] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1458821445] SERVICE ALERT: BOEAPPTEST;CPU Load;OK;HARD;3;CPU Load 51% (5 min average)
[1458821525] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1458821585] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1458821645] SERVICE ALERT: halleground2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1458822027] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1458823370] SERVICE ALERT: INTLTESTAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1458823490] SERVICE ALERT: INTLTESTAPP;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1458824196] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 45% (5 min average)
[1458824796] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1458824915] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 77% (5 min average)
[1458825771] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1458825774] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10022;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1458825774] HOST ALERT: ADWSUSCLIENTS;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.134)
[1458825783] SERVICE ALERT: halleground1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1458825783] SERVICE ALERT: ADWSUSCLIENTS;E:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458825786] SERVICE ALERT: ADWSUSCLIENTS;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.102.134 and port 12489: No route to host
[1458825786] SERVICE ALERT: ADWSUSCLIENTS;Uptime;CRITICAL;SOFT;1;connect to address 164.76.102.134 and port 12489: No route to host
[1458825789] HOST ALERT: ADWSUSCLIENTS;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.134)
[1458825806] SERVICE ALERT: ADWSUSCLIENTS;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.102.134 and port 12489: No route to host
[1458825809] HOST ALERT: ADWSUSCLIENTS;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.102.134)
[1458825843] SERVICE ALERT: halleground1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1458825854] SERVICE ALERT: ADWSUSCLIENTS;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.102.134 and port 12489: No route to host
[1458825854] SERVICE ALERT: ADWSUSCLIENTS;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.102.134 and port 12489: No route to host
[1458825857] HOST ALERT: ADWSUSCLIENTS;DOWN;SOFT;4;CRITICAL - Host Unreachable (164.76.102.134)
[1458825882] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1458825885] SERVICE ALERT: ADWSUSCLIENTS;E:\ Drive Space;CRITICAL;SOFT;2;connect to address 164.76.102.134 and port 12489: No route to host
[1458825885] HOST ALERT: ADWSUSCLIENTS;DOWN;SOFT;5;CRITICAL - Host Unreachable (164.76.102.134)
[1458825896] SERVICE ALERT: ADWSUSCLIENTS;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.102.134 and port 12489: No route to host
[1458825896] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10022;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1458825903] SERVICE ALERT: halleground1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1458825906] SERVICE ALERT: ADWSUSCLIENTS;C:\ Drive Space;CRITICAL;SOFT;2;connect to address 164.76.102.134 and port 12489: No route to host
[1458825906] SERVICE ALERT: ADWSUSCLIENTS;Uptime;CRITICAL;SOFT;2;connect to address 164.76.102.134 and port 12489: No route to host
[1458825909] HOST ALERT: ADWSUSCLIENTS;DOWN;SOFT;6;CRITICAL - Host Unreachable (164.76.102.134)
[1458825926] SERVICE ALERT: ADWSUSCLIENTS;Memory Usage;CRITICAL;SOFT;2;connect to address 164.76.102.134 and port 12489: No route to host
[1458825929] HOST ALERT: ADWSUSCLIENTS;DOWN;SOFT;7;CRITICAL - Host Unreachable (164.76.102.134)
[1458825956] SERVICE ALERT: ADWSUSCLIENTS;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.102.134 and port 12489: No route to host
[1458825959] HOST ALERT: ADWSUSCLIENTS;DOWN;SOFT;8;CRITICAL - Host Unreachable (164.76.102.134)
[1458825974] SERVICE ALERT: ADWSUSCLIENTS;Service: Windows Time;CRITICAL;SOFT;2;connect to address 164.76.102.134 and port 12489: No route to host
[1458825974] SERVICE ALERT: ADWSUSCLIENTS;NSClient++ Version;CRITICAL;SOFT;2;connect to address 164.76.102.134 and port 12489: No route to host
[1458825977] HOST ALERT: ADWSUSCLIENTS;DOWN;SOFT;9;CRITICAL - Host Unreachable (164.76.102.134)
[1458826002] SERVICE ALERT: ADWSUSCLIENTS;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.102.134 and port 12489: No route to host
[1458826002] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1458826005] HOST ALERT: ADWSUSCLIENTS;DOWN;HARD;10;CRITICAL - Host Unreachable (164.76.102.134)
[1458826005] SERVICE ALERT: ADWSUSCLIENTS;E:\ Drive Space;CRITICAL;HARD;3;connect to address 164.76.102.134 and port 12489: No route to host
[1458826016] SERVICE ALERT: ADWSUSCLIENTS;Service: DNS Client;CRITICAL;SOFT;2;connect to address 164.76.102.134 and port 12489: No route to host
[1458826016] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10022;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1458826026] SERVICE ALERT: ADWSUSCLIENTS;C:\ Drive Space;CRITICAL;HARD;3;connect to address 164.76.102.134 and port 12489: No route to host
[1458826026] SERVICE ALERT: ADWSUSCLIENTS;Uptime;CRITICAL;HARD;3;connect to address 164.76.102.134 and port 12489: No route to host
[1458826026] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 364;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1458826041] SERVICE ALERT: ADWSUSCLIENTS;Service: NetLogin;CRITICAL;SOFT;1;connect to address 164.76.102.134 and port 12489: No route to host
[1458826044] SERVICE ALERT: ADWSUSCLIENTS;Memory Usage;CRITICAL;HARD;3;connect to address 164.76.102.134 and port 12489: No route to host
[1458826076] SERVICE ALERT: ADWSUSCLIENTS;Service: Firewall;CRITICAL;SOFT;2;connect to address 164.76.102.134 and port 12489: No route to host
[1458826086] SERVICE ALERT: ADWSUSCLIENTS;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.102.134 and port 12489: No route to host
[1458826094] SERVICE ALERT: ADWSUSCLIENTS;Service: Windows Time;CRITICAL;HARD;3;connect to address 164.76.102.134 and port 12489: No route to host
[1458826094] SERVICE ALERT: ADWSUSCLIENTS;NSClient++ Version;CRITICAL;HARD;3;connect to address 164.76.102.134 and port 12489: No route to host
[1458826100] HOST ALERT: HDBACKUPS;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.104)
[1458826108] SERVICE ALERT: HDBACKUPS;F:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.102.104 and port 12489: No route to host
[1458826111] HOST ALERT: HDBACKUPS;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.104)
[1458826122] SERVICE ALERT: ADWSUSCLIENTS;Service: Group Policy;CRITICAL;SOFT;2;connect to address 164.76.102.134 and port 12489: No route to host
[1458826126] SERVICE ALERT: strongpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1458826138] SERVICE ALERT: ADWSUSCLIENTS;Service: User Profile;CRITICAL;SOFT;1;connect to address 164.76.102.134 and port 12489: No route to host
[1458826138] SERVICE ALERT: ADWSUSCLIENTS;Service: DNS Client;CRITICAL;HARD;3;connect to address 164.76.102.134 and port 12489: No route to host
[1458826144] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 364;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1458826149] SERVICE ALERT: HDBACKUPS;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.102.104 and port 12489: No route to host
[1458826152] HOST ALERT: HDBACKUPS;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.102.104)
[1458826161] SERVICE ALERT: ADWSUSCLIENTS;Service: NetLogin;CRITICAL;SOFT;2;connect to address 164.76.102.134 and port 12489: No route to host
[1458826172] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10042;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1458826186] SERVICE ALERT: ADWSUSCLIENTS;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.102.134 and port 12489: No route to host
[1458826187] SERVICE ALERT: strongpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1458826196] SERVICE ALERT: HDBACKUPS;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.102.104 and port 12489: No route to host
[1458826196] SERVICE ALERT: ADWSUSCLIENTS;Service: Firewall;CRITICAL;HARD;3;connect to address 164.76.102.134 and port 12489: No route to host
[1458826199] HOST ALERT: HDBACKUPS;DOWN;SOFT;4;CRITICAL - Host Unreachable (164.76.102.104)
[1458826206] SERVICE ALERT: ADWSUSCLIENTS;Service: RPC;CRITICAL;SOFT;2;connect to address 164.76.102.134 and port 12489: No route to host
[1458826206] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 12072;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1458826220] SERVICE ALERT: HDBACKUPS;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.102.104 and port 12489: No route to host
[1458826223] HOST ALERT: HDBACKUPS;DOWN;SOFT;5;CRITICAL - Host Unreachable (164.76.102.104)
[1458826228] SERVICE ALERT: HDBACKUPS;F:\ Drive Space;CRITICAL;SOFT;2;connect to address 164.76.102.104 and port 12489: No route to host
[1458826242] SERVICE ALERT: ADWSUSCLIENTS;Service: Group Policy;CRITICAL;HARD;3;connect to address 164.76.102.134 and port 12489: No route to host
[1458826247] SERVICE ALERT: strongpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1458826258] SERVICE ALERT: ADWSUSCLIENTS;Service: User Profile;CRITICAL;SOFT;2;connect to address 164.76.102.134 and port 12489: No route to host
[1458826266] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 364;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1458826266] SERVICE ALERT: HDBACKUPS;Memory Usage;OK;SOFT;2;Memory usage: total:37532.83 MB - used: 1430.79 MB (4%) - free: 36102.04 MB (96%)
[1458826281] SERVICE ALERT: ADWSUSCLIENTS;Service: NetLogin;CRITICAL;HARD;3;connect to address 164.76.102.134 and port 12489: No route to host
[1458826292] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10042;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1458826296] HOST ALERT: HDBACKUPS;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1458826306] SERVICE ALERT: ADWSUSCLIENTS;Service: WMI;CRITICAL;SOFT;2;connect to address 164.76.102.134 and port 12489: No route to host
[1458826313] SERVICE ALERT: HDBACKUPS;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1458826318] SERVICE ALERT: ADWSUSCLIENTS;VMWare Tools;CRITICAL;SOFT;1;connect to address 164.76.102.134 and port 12489: No route to host
[1458826326] SERVICE ALERT: ADWSUSCLIENTS;CPU Load;CRITICAL;SOFT;1;connect to address 164.76.102.134 and port 12489: No route to host
[1458826326] SERVICE ALERT: ADWSUSCLIENTS;Service: RPC;CRITICAL;HARD;3;connect to address 164.76.102.134 and port 12489: No route to host
[1458826326] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 12072;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1458826337] SERVICE ALERT: HDBACKUPS;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1458826343] HOST ALERT: HDBACKUPS;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1458826345] SERVICE ALERT: HDBACKUPS;F:\ Drive Space;OK;SOFT;3;f:\ - total: 3315.22 Gb - used: 1614.71 Gb (49%) - free 1700.51 Gb (51%)
[1458826375] HOST ALERT: HDBACKUPS;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1458826378] SERVICE ALERT: ADWSUSCLIENTS;Service: User Profile;CRITICAL;HARD;3;connect to address 164.76.102.134 and port 12489: No route to host
[1458826411] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10042;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1458826426] SERVICE ALERT: ADWSUSCLIENTS;Service: WMI;CRITICAL;HARD;3;connect to address 164.76.102.134 and port 12489: No route to host
[1458826438] SERVICE ALERT: ADWSUSCLIENTS;VMWare Tools;CRITICAL;SOFT;2;connect to address 164.76.102.134 and port 12489: No route to host
[1458826438] SERVICE ALERT: ADWSUSCLIENTS;CPU Load;CRITICAL;SOFT;2;connect to address 164.76.102.134 and port 12489: No route to host
[1458826446] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 12072;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1458826446] HOST ALERT: HDBACKUPS;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1458826493] HOST ALERT: HDBACKUPS;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1458826558] SERVICE ALERT: ADWSUSCLIENTS;CPU Load;CRITICAL;HARD;3;connect to address 164.76.102.134 and port 12489: No route to host
[1458826558] SERVICE ALERT: ADWSUSCLIENTS;VMWare Tools;CRITICAL;HARD;3;connect to address 164.76.102.134 and port 12489: No route to host
[1458826933] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1458826986] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1458826993] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1458827045] SERVICE ALERT: halleground2;Device Status;OK;HARD;3;Clearing paper path
[1458827046] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1458827054] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;OK;SOFT;3;Black Print Cartridge HP C8543X is at 81% - OK!
[1458827057] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1458827103] SERVICE ALERT: halleground1;Device Status;OK;HARD;3;Clearing paper path
[1458827105] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1458827117] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1458827177] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1458827313] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1458827433] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1458827447] SERVICE ALERT: strongpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1458827553] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1458827640] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1458827645] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1458827670] HOST ALERT: hallefirst2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1458827700] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1458827706] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1458827730] HOST ALERT: hallefirst2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1458827760] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1458827765] SERVICE ALERT: halleground2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1458827783] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1458827790] HOST ALERT: hallefirst2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1458827880] HOST ALERT: hallefirst2;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1458827962] HOST ALERT: hallefirst2;UP;SOFT;5;PING OK - Packet loss = 78%, RTA = 0.52 ms
[1458827974] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1458828034] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 5% - OK!
[1458828354] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1458828377] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1458828634] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1458828694] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1458828754] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1458828827] HOST ALERT: HDBACKUPS;UP;HARD;10;PING OK - Packet loss = 54%, RTA = 0.32 ms
[1458829315] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;OK;HARD;3;c:\ - total: 84.90 Gb - used: 75.74 Gb (89%) - free 9.16 Gb (11%)
[1458829390] HOST ALERT: hallefirst2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1458829480] HOST ALERT: hallefirst2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1458829560] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1458829570] HOST ALERT: hallefirst2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1458829584] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1458829614] HOST ALERT: hallefirst2;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1458829620] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1458829643] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1458829650] HOST ALERT: hallefirst2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1458829680] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1458829703] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1458829703] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1458829710] HOST ALERT: hallefirst2;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1458829800] HOST ALERT: hallefirst2;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1458829890] HOST ALERT: hallefirst2;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1458829953] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 46% (5 min average)
[1458829980] HOST ALERT: hallefirst2;DOWN;SOFT;9;(Host check timed out after 30.00 seconds)
[1458830070] HOST ALERT: hallefirst2;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1458830553] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1458830570] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1458830629] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1458830673] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 94% (5 min average)
[1458830689] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1458830793] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1458832278] SERVICE ALERT: ADWSUSCLIENTS;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1458832323] SERVICE ALERT: ADWSUSCLIENTS;Service: RPC;OK;HARD;3;RpcSs: Started
[1458832373] SERVICE ALERT: ADWSUSCLIENTS;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1458832408] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10042;OK;HARD;3;Eventlog check ok
[1458832423] SERVICE ALERT: ADWSUSCLIENTS;Service: WMI;OK;HARD;3;Winmgmt: Started
[1458832443] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 12072;OK;HARD;3;Eventlog check ok
[1458832505] SERVICE ALERT: crossroadspk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1458832554] SERVICE ALERT: ADWSUSCLIENTS;CPU Load;OK;HARD;3;CPU Load 11% (5 min average)
[1458832555] SERVICE ALERT: ADWSUSCLIENTS;VMWare Tools;OK;HARD;3;VMTools: Started
[1458832599] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;OK;HARD;3;Eventlog check ok
[1458832603] SERVICE ALERT: ADWSUSCLIENTS;E:\ Drive Space;OK;HARD;3;e:\ - total: 299.87 Gb - used: 198.50 Gb (66%) - free 101.38 Gb (34%)
[1458832613] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10022;OK;HARD;3;Eventlog check ok
[1458832623] SERVICE ALERT: ADWSUSCLIENTS;C:\ Drive Space;OK;HARD;3;c:\ - total: 79.66 Gb - used: 38.60 Gb (48%) - free 41.06 Gb (52%)
[1458832623] SERVICE ALERT: ADWSUSCLIENTS;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 6 minute(s)
[1458832643] SERVICE ALERT: ADWSUSCLIENTS;Memory Usage;OK;HARD;3;Memory usage: total:14335.61 MB - used: 2027.71 MB (14%) - free: 12307.90 MB (86%)
[1458832691] SERVICE ALERT: ADWSUSCLIENTS;NSClient++ Version;OK;HARD;3;NSClient++ 0,4,1,73 2012-12-17
[1458832691] SERVICE ALERT: ADWSUSCLIENTS;Service: Windows Time;OK;HARD;3;W32Time: Started
[1458832734] SERVICE ALERT: ADWSUSCLIENTS;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1458832793] SERVICE ALERT: ADWSUSCLIENTS;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1458832839] SERVICE ALERT: ADWSUSCLIENTS;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1458832863] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 364;OK;HARD;3;Eventlog check ok
[1458833104] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1458833158] HOST ALERT: ADWSUSCLIENTS;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 1.89 ms
[1458833165] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1458833225] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1458833403] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1458833403] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1458833403] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1458833403] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1458833403] HOST DOWNTIME ALERT: MAGICDBPROD;STARTED; Host has entered a period of scheduled downtime
[1458833403] SERVICE DOWNTIME ALERT: BOEAPPTEST;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1458833403] HOST DOWNTIME ALERT: MAGICAPPTEST;STARTED; Host has entered a period of scheduled downtime
[1458833689] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;CRITICAL;HARD;3;Black Print Cartridge HP C8543X is at 1% - CRITICAL!
[1458833824] SERVICE ALERT: NETIDADPROXY;Service: NetLogin;CRITICAL;SOFT;1;Netlogon: Stopped
[1458833944] SERVICE ALERT: NETIDADPROXY;Service: NetLogin;CRITICAL;SOFT;2;Netlogon: Stopped
[1458834064] SERVICE ALERT: NETIDADPROXY;Service: NetLogin;CRITICAL;HARD;3;Netlogon: Stopped
[1458835245] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458835365] SERVICE ALERT: BOEAPPTEST;CPU Load;OK;SOFT;2;CPU Load 25% (5 min average)
[1458835443] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 12072;UNKNOWN;SOFT;1;Could not open the 'application' event log: 1717: The interface is unknown.
[1458835563] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 12072;OK;SOFT;2;Eventlog check ok
[1458836247] SERVICE ALERT: C4SAS-SE;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1458836324] SERVICE ALERT: cava2;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458836326] SERVICE ALERT: cava2;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458836326] SERVICE ALERT: cava2;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458836354] HOST ALERT: cava2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1458836363] SERVICE ALERT: cava2;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458836367] SERVICE ALERT: C4SAS-SE;CPU Load;OK;SOFT;2;CPU Load 8% (5 min average)
[1458836373] SERVICE ALERT: cava2;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458836373] SERVICE ALERT: cava2;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1458836382] SERVICE ALERT: cava2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1458836393] HOST ALERT: cava2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1458836404] HOST ALERT: cava2;UP;SOFT;3;PING OK - Packet loss = 58%, RTA = 0.37 ms
[1458836424] SERVICE ALERT: cava2;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 1 minute(s)
[1458836426] SERVICE ALERT: cava2;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1458836426] SERVICE ALERT: cava2;Memory Usage;OK;SOFT;2;Memory usage: total:43391.92 MB - used: 6262.90 MB (14%) - free: 37129.03 MB (86%)
[1458836436] SERVICE ALERT: VUM;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1458836463] SERVICE ALERT: cava2;Service: RPC;OK;SOFT;2;RpcSs: Started
[1458836473] SERVICE ALERT: cava2;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1458836493] SERVICE ALERT: cava2;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1458836502] SERVICE ALERT: cava2;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1458836556] SERVICE ALERT: VUM;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1458836622] SERVICE ALERT: cava2;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1458837260] HOST ALERT: cava2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1458837324] HOST ALERT: cava2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1458837822] SERVICE ALERT: cava2;CPU Load;OK;HARD;3;CPU Load 0% (5 min average)
[1458838076] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1458838076] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1458838076] HOST DOWNTIME ALERT: MAGICDBPROD;STARTED; Host has entered a period of scheduled downtime
[1458838076] SERVICE DOWNTIME ALERT: BOEAPPTEST;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1458838076] HOST DOWNTIME ALERT: MAGICAPPTEST;STARTED; Host has entered a period of scheduled downtime
[1458838489] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1458840068] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1458840068] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1458840068] HOST DOWNTIME ALERT: MAGICDBPROD;STARTED; Host has entered a period of scheduled downtime
[1458840068] SERVICE DOWNTIME ALERT: BOEAPPTEST;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1458840068] HOST DOWNTIME ALERT: MAGICAPPTEST;STARTED; Host has entered a period of scheduled downtime
[1458840316] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1458840316] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1458840316] HOST DOWNTIME ALERT: MAGICDBPROD;STARTED; Host has entered a period of scheduled downtime
[1458840316] SERVICE DOWNTIME ALERT: BOEAPPTEST;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1458840316] HOST DOWNTIME ALERT: MAGICAPPTEST;STARTED; Host has entered a period of scheduled downtime
[1458840422] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1458840431] HOST ALERT: epprogsql2;DOWN;SOFT;1;(Return code of 13 is out of bounds)
[1458840432] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1458840432] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1458840432] HOST DOWNTIME ALERT: MAGICDBPROD;STARTED; Host has entered a period of scheduled downtime
[1458840432] SERVICE DOWNTIME ALERT: BOEAPPTEST;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1458840432] HOST DOWNTIME ALERT: MAGICAPPTEST;STARTED; Host has entered a period of scheduled downtime
[1458840457] HOST ALERT: epprogsql2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1458840476] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1458840476] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1458840476] HOST DOWNTIME ALERT: MAGICDBPROD;STARTED; Host has entered a period of scheduled downtime
[1458840476] SERVICE DOWNTIME ALERT: BOEAPPTEST;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1458840476] HOST DOWNTIME ALERT: MAGICAPPTEST;STARTED; Host has entered a period of scheduled downtime
[1458840483] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1458840542] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1458842053] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1458842113] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 81% - OK!
[1458842165] SERVICE ALERT: halleground2;Device Status;OK;HARD;3;Warming up
[1458842343] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;CRITICAL;HARD;3;Black Print Cartridge HP C8543X is at 1% - CRITICAL!
[1458842713] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1458842773] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 100% - OK!
[1458843351] HOST ALERT: halleground2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1458843371] SERVICE ALERT: halleground2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1458843379] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1458843401] HOST ALERT: halleground2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1458843402] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1458843424] SERVICE ALERT: crossroadspk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1458843431] SERVICE ALERT: halleground2;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1458843432] HOST ALERT: halleground2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1458843440] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1458843462] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1458843470] HOST ALERT: halleground2;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1458843491] SERVICE ALERT: halleground2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1458843499] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1458843521] HOST ALERT: halleground2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1458843522] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1458843552] HOST ALERT: halleground2;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1458843642] HOST ALERT: halleground2;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1458843732] HOST ALERT: halleground2;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1458843822] HOST ALERT: halleground2;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1458843912] HOST ALERT: halleground2;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1458844685] SERVICE ALERT: halleground2;Device Status;OK;HARD;3;Push green lever.
[1458844689] HOST ALERT: halleground2;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.56 ms
[1458844717] SERVICE ALERT: halleground2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1458844742] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1458845293] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 81% - OK!
[1458846993] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 63% (5 min average)
[1458847593] SERVICE ALERT: prayharroldpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1458847593] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458847653] SERVICE ALERT: prayharroldpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1458847713] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 84% (5 min average)
[1458847714] SERVICE ALERT: prayharroldpk;Device Status;OK;SOFT;3;50 72 65 73 73 20 17 20 74 6F 20 73 65 65 20 73  74 65 70  by step information
[1458848313] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1458848433] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1458848553] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1458848784] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1458848838] SERVICE ALERT: prayharroldpk;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! 2 of 3 Trays Still Full
[1458848846] SERVICE ALERT: ADAUTH-RMAD;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.102.174 and port 12489: No route to host
[1458848849] HOST ALERT: ADAUTH-RMAD;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.174)
[1458848878] SERVICE ALERT: ADAUTH-RMAD;Service: NetLogin;CRITICAL;SOFT;1;connect to address 164.76.102.174 and port 12489: No route to host
[1458848881] HOST ALERT: ADAUTH-RMAD;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.174)
[1458848927] SERVICE ALERT: ADAUTH-RMAD;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.102.174 and port 12489: No route to host
[1458848930] HOST ALERT: ADAUTH-RMAD;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.102.174)
[1458848957] SERVICE ALERT: ADAUTH-RMAD;Service: Group Policy;CRITICAL;SOFT;2;connect to address 164.76.102.174 and port 12489: No route to host
[1458848960] HOST ALERT: ADAUTH-RMAD;DOWN;SOFT;4;CRITICAL - Host Unreachable (164.76.102.174)
[1458848974] SERVICE ALERT: ADAUTH-RMAD;Service: User Profile;CRITICAL;SOFT;1;connect to address 164.76.102.174 and port 12489: No route to host
[1458848977] HOST ALERT: ADAUTH-RMAD;DOWN;SOFT;5;CRITICAL - Host Unreachable (164.76.102.174)
[1458848998] SERVICE ALERT: ADAUTH-RMAD;Service: NetLogin;CRITICAL;SOFT;2;connect to address 164.76.102.174 and port 12489: No route to host
[1458849001] HOST ALERT: ADAUTH-RMAD;DOWN;SOFT;6;CRITICAL - Host Unreachable (164.76.102.174)
[1458849047] SERVICE ALERT: ADAUTH-RMAD;Service: RPC;CRITICAL;SOFT;2;connect to address 164.76.102.174 and port 12489: No route to host
[1458849050] HOST ALERT: ADAUTH-RMAD;DOWN;SOFT;7;CRITICAL - Host Unreachable (164.76.102.174)
[1458849063] SERVICE ALERT: AD-RMAD;CPU Load;CRITICAL;SOFT;1;connect to address 164.76.102.135 and port 12489: No route to host
[1458849063] SERVICE ALERT: AD-RMAD;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.102.135 and port 12489: No route to host
[1458849063] SERVICE ALERT: AD-RMAD;Service: NetLogin;CRITICAL;SOFT;1;connect to address 164.76.102.135 and port 12489: No route to host
[1458849066] SERVICE ALERT: AD-RMAD;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.102.135 and port 12489: No route to host
[1458849066] HOST ALERT: AD-RMAD;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.135)
[1458849071] SERVICE ALERT: AD-RMAD;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458849073] SERVICE ALERT: AD-RMAD;E:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458849073] SERVICE ALERT: AD-RMAD;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458849073] SERVICE ALERT: AD-RMAD;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458849073] SERVICE ALERT: AD-RMAD;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458849073] SERVICE ALERT: AD-RMAD;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458849074] SERVICE ALERT: AD-RMAD;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458849074] SERVICE ALERT: AD-RMAD;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458849074] SERVICE ALERT: AD-RMAD;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458849074] SERVICE ALERT: AD-RMAD;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458849074] SERVICE ALERT: AD-RMAD;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458849076] SERVICE ALERT: ADAUTH-RMAD;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.102.174 and port 12489: No route to host
[1458849076] SERVICE ALERT: ADAUTH-RMAD;Service: Group Policy;CRITICAL;HARD;3;connect to address 164.76.102.174 and port 12489: No route to host
[1458849079] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1458849079] HOST ALERT: ADAUTH-RMAD;DOWN;SOFT;8;CRITICAL - Host Unreachable (164.76.102.174)
[1458849094] SERVICE ALERT: ADAUTH-RMAD;Service: User Profile;CRITICAL;SOFT;2;connect to address 164.76.102.174 and port 12489: No route to host
[1458849097] HOST ALERT: ADAUTH-RMAD;DOWN;SOFT;9;CRITICAL - Host Unreachable (164.76.102.174)
[1458849108] SERVICE ALERT: ADAUTH-RMAD;Uptime;CRITICAL;SOFT;1;connect to address 164.76.102.174 and port 12489: No route to host
[1458849109] HOST ALERT: prayharroldpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1458849118] SERVICE ALERT: ADAUTH-RMAD;Service: NetLogin;CRITICAL;HARD;3;connect to address 164.76.102.174 and port 12489: No route to host
[1458849121] HOST ALERT: ADAUTH-RMAD;DOWN;HARD;10;CRITICAL - Host Unreachable (164.76.102.174)
[1458849139] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1458849152] SERVICE ALERT: ADAUTH-RMAD;VMWare Tools;CRITICAL;SOFT;1;connect to address 164.76.102.174 and port 12489: No route to host
[1458849155] SERVICE ALERT: ADAUTH-RMAD;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.102.174 and port 12489: No route to host
[1458849167] SERVICE ALERT: ADAUTH-RMAD;Service: RPC;CRITICAL;HARD;3;connect to address 164.76.102.174 and port 12489: No route to host
[1458849169] HOST ALERT: prayharroldpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1458849174] SERVICE ALERT: AD-RMAD;Service: DNS Client;CRITICAL;SOFT;2;connect to address 164.76.102.135 and port 12489: No route to host
[1458849174] SERVICE ALERT: AD-RMAD;Service: Windows Time;CRITICAL;SOFT;2;connect to address 164.76.102.135 and port 12489: No route to host
[1458849174] SERVICE ALERT: AD-RMAD;Service: Group Policy;CRITICAL;SOFT;2;connect to address 164.76.102.135 and port 12489: No route to host
[1458849177] SERVICE ALERT: AD-RMAD;Service: NetLogin;CRITICAL;SOFT;2;connect to address 164.76.102.135 and port 12489: No route to host
[1458849177] SERVICE ALERT: AD-RMAD;Memory Usage;CRITICAL;SOFT;2;connect to address 164.76.102.135 and port 12489: No route to host
[1458849177] SERVICE ALERT: AD-RMAD;E:\ Drive Space;CRITICAL;SOFT;2;connect to address 164.76.102.135 and port 12489: No route to host
[1458849178] SERVICE ALERT: ADWSUSCLIENTS;Service: User Profile;CRITICAL;SOFT;1;connect to address 164.76.102.134 and port 12489: No route to host
[1458849178] HOST ALERT: ADWSUSCLIENTS;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.134)
[1458849180] SERVICE ALERT: AD-RMAD;C:\ Drive Space;CRITICAL;SOFT;2;connect to address 164.76.102.135 and port 12489: No route to host
[1458849180] HOST ALERT: AD-RMAD;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.135)
[1458849183] SERVICE ALERT: AD-RMAD;CPU Load;CRITICAL;SOFT;2;connect to address 164.76.102.135 and port 12489: No route to host
[1458849183] SERVICE ALERT: AD-RMAD;NSClient++ Version;CRITICAL;SOFT;2;connect to address 164.76.102.135 and port 12489: No route to host
[1458849186] SERVICE ALERT: ADAUTH-RMAD;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.102.174 and port 12489: No route to host
[1458849186] SERVICE ALERT: AD-RMAD;Service: RPC;CRITICAL;SOFT;2;connect to address 164.76.102.135 and port 12489: No route to host
[1458849186] SERVICE ALERT: AD-RMAD;Service: WMI;CRITICAL;SOFT;2;connect to address 164.76.102.135 and port 12489: No route to host
[1458849186] SERVICE ALERT: AD-RMAD;Service: Firewall;CRITICAL;SOFT;2;connect to address 164.76.102.135 and port 12489: No route to host
[1458849191] SERVICE ALERT: AD-RMAD;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458849194] SERVICE ALERT: AD-RMAD;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458849194] HOST ALERT: AD-RMAD;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.102.135)
[1458849194] SERVICE ALERT: AD-RMAD;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458849196] SERVICE ALERT: ADAUTH-RMAD;CPU Load;CRITICAL;SOFT;1;connect to address 164.76.102.174 and port 12489: No route to host
[1458849196] SERVICE ALERT: ADAUTH-RMAD;Service: WMI;CRITICAL;SOFT;2;connect to address 164.76.102.174 and port 12489: No route to host
[1458849199] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1458849211] HOST ALERT: ADWSUSCLIENTS;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.134)
[1458849214] SERVICE ALERT: ADAUTH-RMAD;Service: User Profile;CRITICAL;HARD;3;connect to address 164.76.102.174 and port 12489: No route to host
[1458849226] SERVICE ALERT: ADWSUSCLIENTS;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.102.134 and port 12489: No route to host
[1458849228] SERVICE ALERT: ADAUTH-RMAD;Uptime;CRITICAL;SOFT;2;connect to address 164.76.102.174 and port 12489: No route to host
[1458849229] HOST ALERT: prayharroldpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1458849229] HOST ALERT: ADWSUSCLIENTS;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.102.134)
[1458849246] SERVICE ALERT: ADAUTH-RMAD;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.102.174 and port 12489: No route to host
[1458849272] SERVICE ALERT: ADAUTH-RMAD;VMWare Tools;CRITICAL;SOFT;2;connect to address 164.76.102.174 and port 12489: No route to host
[1458849275] SERVICE ALERT: ADAUTH-RMAD;C:\ Drive Space;CRITICAL;SOFT;2;connect to address 164.76.102.174 and port 12489: No route to host
[1458849292] HOST ALERT: ADWSUSCLIENTS;DOWN;SOFT;4;CRITICAL - Host Unreachable (164.76.102.134)
[1458849294] SERVICE ALERT: AD-RMAD;E:\ Drive Space;CRITICAL;HARD;3;connect to address 164.76.102.135 and port 12489: No route to host
[1458849294] SERVICE ALERT: AD-RMAD;CPU Load;CRITICAL;HARD;3;connect to address 164.76.102.135 and port 12489: No route to host
[1458849294] SERVICE ALERT: AD-RMAD;Uptime;CRITICAL;HARD;3;connect to address 164.76.102.135 and port 12489: No route to host
[1458849296] SERVICE ALERT: ADWSUSCLIENTS;Service: User Profile;CRITICAL;SOFT;2;connect to address 164.76.102.134 and port 12489: No route to host
[1458849297] SERVICE ALERT: AD-RMAD;Service: NetLogin;CRITICAL;HARD;3;connect to address 164.76.102.135 and port 12489: No route to host
[1458849297] HOST ALERT: AD-RMAD;DOWN;SOFT;4;CRITICAL - Host Unreachable (164.76.102.135)
[1458849297] SERVICE ALERT: AD-RMAD;Service: Windows Time;CRITICAL;HARD;3;connect to address 164.76.102.135 and port 12489: No route to host
[1458849297] SERVICE ALERT: ADAUTH-RMAD;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.102.174 and port 12489: No route to host
[1458849300] SERVICE ALERT: AD-RMAD;Service: Firewall;CRITICAL;HARD;3;connect to address 164.76.102.135 and port 12489: No route to host
[1458849300] SERVICE ALERT: AD-RMAD;C:\ Drive Space;CRITICAL;HARD;3;connect to address 164.76.102.135 and port 12489: No route to host
[1458849300] SERVICE ALERT: AD-RMAD;Service: User Profile;CRITICAL;HARD;3;connect to address 164.76.102.135 and port 12489: No route to host
[1458849303] SERVICE ALERT: AD-RMAD;Service: Group Policy;CRITICAL;HARD;3;connect to address 164.76.102.135 and port 12489: No route to host
[1458849303] SERVICE ALERT: AD-RMAD;Memory Usage;CRITICAL;HARD;3;connect to address 164.76.102.135 and port 12489: No route to host
[1458849303] SERVICE ALERT: AD-RMAD;NSClient++ Version;CRITICAL;HARD;3;connect to address 164.76.102.135 and port 12489: No route to host
[1458849306] SERVICE ALERT: ADAUTH-RMAD;Service: Windows Time;CRITICAL;SOFT;2;connect to address 164.76.102.174 and port 12489: No route to host
[1458849306] SERVICE ALERT: AD-RMAD;Service: RPC;CRITICAL;HARD;3;connect to address 164.76.102.135 and port 12489: No route to host
[1458849306] SERVICE ALERT: AD-RMAD;Service: WMI;CRITICAL;HARD;3;connect to address 164.76.102.135 and port 12489: No route to host
[1458849311] SERVICE ALERT: AD-RMAD;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458849313] SERVICE ALERT: AD-RMAD;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458849314] HOST ALERT: AD-RMAD;DOWN;SOFT;5;CRITICAL - Host Unreachable (164.76.102.135)
[1458849314] SERVICE ALERT: ADAUTH-RMAD;E:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.102.174 and port 12489: No route to host
[1458849314] SERVICE ALERT: ADAUTH-RMAD;Service: WMI;CRITICAL;HARD;3;connect to address 164.76.102.174 and port 12489: No route to host
[1458849317] SERVICE ALERT: ADAUTH-RMAD;CPU Load;CRITICAL;SOFT;2;connect to address 164.76.102.174 and port 12489: No route to host
[1458849319] HOST ALERT: prayharroldpk;DOWN;SOFT;4;(Host check timed out after 30.00 seconds)
[1458849343] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 12072;WARNING;SOFT;1;(No output on stdout) stderr: Could not resolve hostname ADWSUSSERVERS: Name or service not known
[1458849346] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10022;WARNING;SOFT;1;(No output on stdout) stderr: Could not resolve hostname ADWSUSSERVERS: Name or service not known
[1458849346] SERVICE ALERT: ADWSUSCLIENTS;Service: WMI;CRITICAL;SOFT;2;connect to address 164.76.102.134 and port 12489: No route to host
[1458849346] SERVICE ALERT: ADAUTH-RMAD;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.102.174 and port 12489: No route to host
[1458849346] SERVICE ALERT: ADAUTH-RMAD;Uptime;CRITICAL;HARD;3;connect to address 164.76.102.174 and port 12489: No route to host
[1458849349] HOST ALERT: ADWSUSCLIENTS;DOWN;SOFT;5;CRITICAL - Host Unreachable (164.76.102.134)
[1458849356] SERVICE ALERT: ADWSUSCLIENTS;CPU Load;CRITICAL;SOFT;1;connect to address 164.76.102.134 and port 12489: No route to host
[1458849356] SERVICE ALERT: ADWSUSCLIENTS;VMWare Tools;CRITICAL;SOFT;1;connect to address 164.76.102.134 and port 12489: No route to host
[1458849366] SERVICE ALERT: ADAUTH-RMAD;Memory Usage;CRITICAL;SOFT;2;connect to address 164.76.102.174 and port 12489: No route to host
[1458849367] HOST ALERT: ADWSUSCLIENTS;DOWN;SOFT;6;CRITICAL - Host Unreachable (164.76.102.134)
[1458849377] HOST ALERT: AD-RMAD;DOWN;SOFT;6;CRITICAL - Host Unreachable (164.76.102.135)
[1458849385] SERVICE ALERT: ADWSUSSERVERS;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.102.133 and port 12489: No route to host
[1458849386] SERVICE ALERT: ADAUTH-RMAD;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.102.174 and port 12489: No route to host
[1458849388] HOST ALERT: ADWSUSSERVERS;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.133)
[1458849389] SERVICE ALERT: ADAUTH-RMAD;VMWare Tools;CRITICAL;HARD;3;connect to address 164.76.102.174 and port 12489: No route to host
[1458849396] SERVICE ALERT: ADAUTH-RMAD;C:\ Drive Space;CRITICAL;HARD;3;connect to address 164.76.102.174 and port 12489: No route to host
[1458849402] HOST ALERT: ADWSUSCLIENTS;DOWN;SOFT;7;CRITICAL - Host Unreachable (164.76.102.134)
[1458849406] SERVICE ALERT: ADWSUSCLIENTS;E:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.102.134 and port 12489: No route to host
[1458849409] HOST ALERT: prayharroldpk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1458849416] SERVICE ALERT: ADWSUSCLIENTS;Service: User Profile;CRITICAL;HARD;3;connect to address 164.76.102.134 and port 12489: No route to host
[1458849417] SERVICE ALERT: ADAUTH-RMAD;NSClient++ Version;CRITICAL;SOFT;2;connect to address 164.76.102.174 and port 12489: No route to host
[1458849419] HOST ALERT: ADWSUSCLIENTS;DOWN;SOFT;8;CRITICAL - Host Unreachable (164.76.102.134)
[1458849419] SERVICE ALERT: ADWSUSSERVERS;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.102.133 and port 12489: No route to host
[1458849422] SERVICE ALERT: ADWSUSSERVERS;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.102.133 and port 12489: No route to host
[1458849422] HOST ALERT: ADWSUSSERVERS;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.133)
[1458849426] SERVICE ALERT: ADAUTH-RMAD;Service: Windows Time;CRITICAL;HARD;3;connect to address 164.76.102.174 and port 12489: No route to host
[1458849426] SERVICE ALERT: ADWSUSCLIENTS;Uptime;CRITICAL;SOFT;1;connect to address 164.76.102.134 and port 12489: No route to host
[1458849426] SERVICE ALERT: ADWSUSCLIENTS;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.102.134 and port 12489: No route to host
[1458849434] SERVICE ALERT: ADAUTH-RMAD;E:\ Drive Space;CRITICAL;SOFT;2;connect to address 164.76.102.174 and port 12489: No route to host
[1458849437] SERVICE ALERT: ADAUTH-RMAD;CPU Load;CRITICAL;HARD;3;connect to address 164.76.102.174 and port 12489: No route to host
[1458849440] HOST ALERT: AD-RMAD;DOWN;SOFT;7;CRITICAL - Host Unreachable (164.76.102.135)
[1458849444] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1458849446] SERVICE ALERT: ADWSUSCLIENTS;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.102.134 and port 12489: No route to host
[1458849449] HOST ALERT: ADWSUSCLIENTS;DOWN;SOFT;9;CRITICAL - Host Unreachable (164.76.102.134)
[1458849463] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 12072;OK;SOFT;2;Eventlog check ok
[1458849466] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10022;OK;SOFT;2;Eventlog check ok
[1458849466] SERVICE ALERT: ADWSUSCLIENTS;Service: WMI;CRITICAL;HARD;3;connect to address 164.76.102.134 and port 12489: No route to host
[1458849466] SERVICE ALERT: ADAUTH-RMAD;Service: DNS Client;CRITICAL;SOFT;2;connect to address 164.76.102.174 and port 12489: No route to host
[1458849466] HOST ALERT: ADWSUSSERVERS;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.102.133)
[1458849469] HOST ALERT: ADWSUSCLIENTS;DOWN;HARD;10;CRITICAL - Host Unreachable (164.76.102.134)
[1458849470] SERVICE ALERT: ADWSUSSERVERS;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.102.133 and port 12489: No route to host
[1458849470] SERVICE ALERT: ADWSUSSERVERS;Uptime;CRITICAL;SOFT;1;connect to address 164.76.102.133 and port 12489: No route to host
[1458849475] HOST ALERT: prayharroldpk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1458849476] SERVICE ALERT: ADWSUSCLIENTS;CPU Load;CRITICAL;SOFT;2;connect to address 164.76.102.134 and port 12489: No route to host
[1458849476] SERVICE ALERT: ADWSUSCLIENTS;VMWare Tools;CRITICAL;SOFT;2;connect to address 164.76.102.134 and port 12489: No route to host
[1458849486] SERVICE ALERT: ADAUTH-RMAD;Memory Usage;CRITICAL;HARD;3;connect to address 164.76.102.174 and port 12489: No route to host
[1458849487] SERVICE ALERT: ADWSUSSERVERS;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.102.133 and port 12489: No route to host
[1458849490] HOST ALERT: ADWSUSSERVERS;DOWN;SOFT;4;CRITICAL - Host Unreachable (164.76.102.133)
[1458849494] SERVICE ALERT: ADWSUSCLIENTS;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.102.134 and port 12489: No route to host
[1458849494] SERVICE ALERT: ADWSUSCLIENTS;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.102.134 and port 12489: No route to host
[1458849496] SERVICE ALERT: ADWSUSSERVERS;Service: WMI;CRITICAL;SOFT;2;connect to address 164.76.102.133 and port 12489: No route to host
[1458849503] HOST ALERT: AD-RMAD;DOWN;SOFT;8;CRITICAL - Host Unreachable (164.76.102.135)
[1458849504] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1458849506] SERVICE ALERT: ADAUTH-RMAD;Service: Firewall;CRITICAL;SOFT;2;connect to address 164.76.102.174 and port 12489: No route to host
[1458849518] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1458849526] SERVICE ALERT: ADWSUSCLIENTS;E:\ Drive Space;CRITICAL;SOFT;2;connect to address 164.76.102.134 and port 12489: No route to host
[1458849534] HOST ALERT: prayharroldpk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1458849536] SERVICE ALERT: ADWSUSCLIENTS;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.102.134 and port 12489: No route to host
[1458849537] SERVICE ALERT: ADAUTH-RMAD;NSClient++ Version;CRITICAL;HARD;3;connect to address 164.76.102.174 and port 12489: No route to host
[1458849539] SERVICE ALERT: ADWSUSSERVERS;Service: Windows Time;CRITICAL;SOFT;2;connect to address 164.76.102.133 and port 12489: No route to host
[1458849542] SERVICE ALERT: ADWSUSSERVERS;Service: Group Policy;CRITICAL;SOFT;2;connect to address 164.76.102.133 and port 12489: No route to host
[1458849542] HOST ALERT: ADWSUSSERVERS;DOWN;SOFT;5;CRITICAL - Host Unreachable (164.76.102.133)
[1458849546] SERVICE ALERT: ADWSUSCLIENTS;C:\ Drive Space;CRITICAL;SOFT;2;connect to address 164.76.102.134 and port 12489: No route to host
[1458849546] SERVICE ALERT: ADWSUSCLIENTS;Uptime;CRITICAL;SOFT;2;connect to address 164.76.102.134 and port 12489: No route to host
[1458849554] SERVICE ALERT: ADAUTH-RMAD;E:\ Drive Space;CRITICAL;HARD;3;connect to address 164.76.102.174 and port 12489: No route to host
[1458849563] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1458849566] SERVICE ALERT: C4SAS-SE;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1458849566] HOST ALERT: AD-RMAD;DOWN;SOFT;9;CRITICAL - Host Unreachable (164.76.102.135)
[1458849566] SERVICE ALERT: ADWSUSCLIENTS;Memory Usage;CRITICAL;SOFT;2;connect to address 164.76.102.134 and port 12489: No route to host
[1458849566] HOST ALERT: ADWSUSSERVERS;DOWN;SOFT;6;CRITICAL - Host Unreachable (164.76.102.133)
[1458849578] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1458849586] SERVICE ALERT: ADAUTH-RMAD;Service: DNS Client;CRITICAL;HARD;3;connect to address 164.76.102.174 and port 12489: No route to host
[1458849589] HOST ALERT: ADWSUSSERVERS;DOWN;SOFT;7;CRITICAL - Host Unreachable (164.76.102.133)
[1458849593] HOST ALERT: prayharroldpk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1458849593] SERVICE ALERT: ADWSUSSERVERS;Uptime;CRITICAL;SOFT;2;connect to address 164.76.102.133 and port 12489: No route to host
[1458849593] SERVICE ALERT: ADWSUSSERVERS;C:\ Drive Space;CRITICAL;SOFT;2;connect to address 164.76.102.133 and port 12489: No route to host
[1458849596] SERVICE ALERT: ADWSUSCLIENTS;VMWare Tools;CRITICAL;HARD;3;connect to address 164.76.102.134 and port 12489: No route to host
[1458849596] SERVICE ALERT: ADWSUSCLIENTS;CPU Load;CRITICAL;HARD;3;connect to address 164.76.102.134 and port 12489: No route to host
[1458849596] SERVICE ALERT: ADWSUSCLIENTS;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.102.134 and port 12489: No route to host
[1458849607] SERVICE ALERT: ADWSUSSERVERS;Memory Usage;CRITICAL;SOFT;2;connect to address 164.76.102.133 and port 12489: No route to host
[1458849610] HOST ALERT: ADWSUSSERVERS;DOWN;SOFT;8;CRITICAL - Host Unreachable (164.76.102.133)
[1458849614] SERVICE ALERT: ADWSUSCLIENTS;NSClient++ Version;CRITICAL;SOFT;2;connect to address 164.76.102.134 and port 12489: No route to host
[1458849616] SERVICE ALERT: ADWSUSSERVERS;Service: WMI;CRITICAL;HARD;3;connect to address 164.76.102.133 and port 12489: No route to host
[1458849623] SERVICE ALERT: ADWSUSCLIENTS;Service: Windows Time;CRITICAL;SOFT;2;connect to address 164.76.102.134 and port 12489: No route to host
[1458849626] SERVICE ALERT: ADAUTH-RMAD;Service: Firewall;CRITICAL;HARD;3;connect to address 164.76.102.174 and port 12489: No route to host
[1458849629] HOST ALERT: AD-RMAD;DOWN;HARD;10;CRITICAL - Host Unreachable (164.76.102.135)
[1458849632] SERVICE ALERT: ADWSUSSERVERS;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.102.133 and port 12489: No route to host
[1458849635] HOST ALERT: ADWSUSSERVERS;DOWN;SOFT;9;CRITICAL - Host Unreachable (164.76.102.133)
[1458849638] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1458849641] SERVICE ALERT: ADWSUSCLIENTS;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.102.134 and port 12489: No route to host
[1458849644] SERVICE ALERT: ADWSUSCLIENTS;E:\ Drive Space;CRITICAL;HARD;3;connect to address 164.76.102.134 and port 12489: No route to host
[1458849645] SERVICE ALERT: ADWSUSSERVERS;VMWare Tools;CRITICAL;SOFT;1;connect to address 164.76.102.133 and port 12489: No route to host
[1458849645] SERVICE ALERT: ADWSUSSERVERS;CPU Load;CRITICAL;SOFT;1;connect to address 164.76.102.133 and port 12489: No route to host
[1458849656] SERVICE ALERT: ADWSUSCLIENTS;Service: DNS Client;CRITICAL;SOFT;2;connect to address 164.76.102.134 and port 12489: No route to host
[1458849659] SERVICE ALERT: ADWSUSSERVERS;Service: Windows Time;CRITICAL;HARD;3;connect to address 164.76.102.133 and port 12489: No route to host
[1458849659] SERVICE ALERT: ADWSUSSERVERS;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.102.133 and port 12489: No route to host
[1458849662] SERVICE ALERT: ADWSUSSERVERS;Service: Group Policy;CRITICAL;HARD;3;connect to address 164.76.102.133 and port 12489: No route to host
[1458849662] HOST ALERT: ADWSUSSERVERS;DOWN;HARD;10;CRITICAL - Host Unreachable (164.76.102.133)
[1458849664] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1458849666] SERVICE ALERT: ADWSUSCLIENTS;Uptime;CRITICAL;HARD;3;connect to address 164.76.102.134 and port 12489: No route to host
[1458849666] SERVICE ALERT: ADWSUSCLIENTS;C:\ Drive Space;CRITICAL;HARD;3;connect to address 164.76.102.134 and port 12489: No route to host
[1458849668] HOST ALERT: prayharroldpk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1458849681] SERVICE ALERT: ADWSUSCLIENTS;Service: NetLogin;CRITICAL;SOFT;1;connect to address 164.76.102.134 and port 12489: No route to host
[1458849684] SERVICE ALERT: ADWSUSCLIENTS;Memory Usage;CRITICAL;HARD;3;connect to address 164.76.102.134 and port 12489: No route to host
[1458849684] SERVICE ALERT: C4SAS-SE;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1458849705] SERVICE ALERT: ADWSUSSERVERS;E:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.102.133 and port 12489: No route to host
[1458849708] SERVICE ALERT: ADWSUSSERVERS;Uptime;CRITICAL;HARD;3;connect to address 164.76.102.133 and port 12489: No route to host
[1458849708] SERVICE ALERT: ADWSUSSERVERS;C:\ Drive Space;CRITICAL;HARD;3;connect to address 164.76.102.133 and port 12489: No route to host
[1458849712] SERVICE ALERT: ADWSUSSERVERS;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.102.133 and port 12489: No route to host
[1458849715] SERVICE ALERT: ADWSUSSERVERS;Service: NetLogin;CRITICAL;SOFT;1;connect to address 164.76.102.133 and port 12489: No route to host
[1458849716] SERVICE ALERT: ADWSUSCLIENTS;Service: Firewall;CRITICAL;SOFT;2;connect to address 164.76.102.134 and port 12489: No route to host
[1458849724] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1458849726] SERVICE ALERT: ADWSUSCLIENTS;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.102.134 and port 12489: No route to host
[1458849727] SERVICE ALERT: ADWSUSSERVERS;Memory Usage;CRITICAL;HARD;3;connect to address 164.76.102.133 and port 12489: No route to host
[1458849732] HOST ALERT: prayharroldpk;UP;SOFT;10;PING OK - Packet loss = 0%, RTA = 0.73 ms
[1458849734] SERVICE ALERT: ADWSUSCLIENTS;Service: Windows Time;CRITICAL;HARD;3;connect to address 164.76.102.134 and port 12489: No route to host
[1458849734] SERVICE ALERT: ADWSUSCLIENTS;NSClient++ Version;CRITICAL;HARD;3;connect to address 164.76.102.134 and port 12489: No route to host
[1458849752] SERVICE ALERT: ADWSUSSERVERS;Service: Firewall;CRITICAL;SOFT;2;connect to address 164.76.102.133 and port 12489: No route to host
[1458849761] SERVICE ALERT: ADWSUSCLIENTS;Service: Group Policy;CRITICAL;SOFT;2;connect to address 164.76.102.134 and port 12489: No route to host
[1458849765] SERVICE ALERT: ADWSUSSERVERS;VMWare Tools;CRITICAL;SOFT;2;connect to address 164.76.102.133 and port 12489: No route to host
[1458849774] SERVICE ALERT: ADWSUSSERVERS;CPU Load;CRITICAL;SOFT;2;connect to address 164.76.102.133 and port 12489: No route to host
[1458849776] SERVICE ALERT: ADWSUSCLIENTS;Service: DNS Client;CRITICAL;HARD;3;connect to address 164.76.102.134 and port 12489: No route to host
[1458849777] SERVICE ALERT: ADWSUSSERVERS;NSClient++ Version;CRITICAL;SOFT;2;connect to address 164.76.102.133 and port 12489: No route to host
[1458849784] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1458849793] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 30% - OK!
[1458849801] SERVICE ALERT: ADWSUSCLIENTS;Service: NetLogin;CRITICAL;SOFT;2;connect to address 164.76.102.134 and port 12489: No route to host
[1458849825] SERVICE ALERT: ADWSUSSERVERS;E:\ Drive Space;CRITICAL;SOFT;2;connect to address 164.76.102.133 and port 12489: No route to host
[1458849828] SERVICE ALERT: ADWSUSSERVERS;Service: DNS Client;CRITICAL;SOFT;2;connect to address 164.76.102.133 and port 12489: No route to host
[1458849836] SERVICE ALERT: ADWSUSCLIENTS;Service: Firewall;CRITICAL;HARD;3;connect to address 164.76.102.134 and port 12489: No route to host
[1458849838] SERVICE ALERT: ADWSUSSERVERS;Service: NetLogin;CRITICAL;SOFT;2;connect to address 164.76.102.133 and port 12489: No route to host
[1458849846] SERVICE ALERT: ADWSUSCLIENTS;Service: RPC;CRITICAL;SOFT;2;connect to address 164.76.102.134 and port 12489: No route to host
[1458849872] SERVICE ALERT: ADWSUSSERVERS;Service: Firewall;CRITICAL;HARD;3;connect to address 164.76.102.133 and port 12489: No route to host
[1458849881] SERVICE ALERT: ADWSUSCLIENTS;Service: Group Policy;CRITICAL;HARD;3;connect to address 164.76.102.134 and port 12489: No route to host
[1458849885] SERVICE ALERT: ADWSUSSERVERS;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.102.133 and port 12489: No route to host
[1458849885] SERVICE ALERT: ADWSUSSERVERS;VMWare Tools;CRITICAL;HARD;3;connect to address 164.76.102.133 and port 12489: No route to host
[1458849885] SERVICE ALERT: ADWSUSSERVERS;CPU Load;CRITICAL;HARD;3;connect to address 164.76.102.133 and port 12489: No route to host
[1458849900] SERVICE ALERT: ADWSUSSERVERS;NSClient++ Version;CRITICAL;HARD;3;connect to address 164.76.102.133 and port 12489: No route to host
[1458849921] SERVICE ALERT: ADWSUSCLIENTS;Service: NetLogin;CRITICAL;HARD;3;connect to address 164.76.102.134 and port 12489: No route to host
[1458849927] SERVICE ALERT: ADWSUSSERVERS;Service: User Profile;CRITICAL;SOFT;1;connect to address 164.76.102.133 and port 12489: No route to host
[1458849945] SERVICE ALERT: ADWSUSSERVERS;E:\ Drive Space;CRITICAL;HARD;3;connect to address 164.76.102.133 and port 12489: No route to host
[1458849948] SERVICE ALERT: ADWSUSSERVERS;Service: DNS Client;CRITICAL;HARD;3;connect to address 164.76.102.133 and port 12489: No route to host
[1458849958] SERVICE ALERT: ADWSUSSERVERS;Service: NetLogin;CRITICAL;HARD;3;connect to address 164.76.102.133 and port 12489: No route to host
[1458849969] SERVICE ALERT: ADWSUSCLIENTS;Service: RPC;CRITICAL;HARD;3;connect to address 164.76.102.134 and port 12489: No route to host
[1458850007] SERVICE ALERT: ADWSUSSERVERS;Service: RPC;CRITICAL;SOFT;2;connect to address 164.76.102.133 and port 12489: No route to host
[1458850047] SERVICE ALERT: ADWSUSSERVERS;Service: User Profile;CRITICAL;SOFT;2;connect to address 164.76.102.133 and port 12489: No route to host
[1458850084] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1458850127] SERVICE ALERT: ADWSUSSERVERS;Service: RPC;CRITICAL;HARD;3;connect to address 164.76.102.133 and port 12489: No route to host
[1458850145] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1458850158] SERVICE ALERT: prayharroldpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1458850166] SERVICE ALERT: ADWSUSSERVERS;Service: User Profile;CRITICAL;HARD;3;connect to address 164.76.102.133 and port 12489: No route to host
[1458850204] SERVICE ALERT: halleground2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1458850204] SERVICE FLAPPING ALERT: halleground2;Device Status;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1458850232] SERVICE ALERT: prayharroldpk;Device Status;OK;HARD;3;Clearing paper path
[1458850393] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1458850453] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 100% - OK!
[1458850972] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 1 Trays Empty
[1458851032] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 1 Trays Empty
[1458851093] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 1 Trays Empty
[1458851432] SERVICE ALERT: prayharroldpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1458851492] SERVICE ALERT: prayharroldpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1458851553] SERVICE ALERT: prayharroldpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1458851659] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1458851689] HOST ALERT: prayharroldpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1458851719] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1458851749] HOST ALERT: prayharroldpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1458851779] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1458851809] HOST ALERT: prayharroldpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1458851899] HOST ALERT: prayharroldpk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1458851963] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1458851989] HOST ALERT: prayharroldpk;DOWN;SOFT;5;(Host check timed out after 30.00 seconds)
[1458852004] SERVICE ALERT: halleground2;Device Status;OK;HARD;3;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1458852023] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1458852053] HOST ALERT: prayharroldpk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1458852083] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1458852113] HOST ALERT: prayharroldpk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1458852158] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1458852158] SERVICE FLAPPING ALERT: prayharroldpk;Device Status;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1458852184] SERVICE ALERT: studentcenterpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1458852188] HOST ALERT: prayharroldpk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1458852278] HOST ALERT: prayharroldpk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1458852368] HOST ALERT: prayharroldpk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1458856491] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1458857303] SERVICE ALERT: owen101lab1;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1458857363] SERVICE ALERT: owen101lab1;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1458857423] SERVICE ALERT: owen101lab1;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1458858623] SERVICE ALERT: owen101lab1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1458859353] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 55% (5 min average)
[1458859953] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1458860073] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 75% (5 min average)
[1458860673] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1458860793] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1458860913] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1458862204] SERVICE FLAPPING ALERT: halleground2;Device Status;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1458863558] SERVICE FLAPPING ALERT: prayharroldpk;Device Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1458864513] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 75% (5 min average)
[1458865113] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458865233] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 94% (5 min average)
[1458865353] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 95% (5 min average)
[1458865353] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1458865953] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 69% (5 min average)
[1458866553] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1458866608] HOST DOWNTIME ALERT: MAGICAPPTEST;STOPPED; Host has exited from a period of scheduled downtime
[1458866635] HOST DOWNTIME ALERT: MAGICDBPROD;STOPPED; Host has exited from a period of scheduled downtime
[1458866673] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1458866793] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1458867700] SERVICE ALERT: TARGETX-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1458867820] SERVICE ALERT: TARGETX-2012;CPU Load;OK;SOFT;2;CPU Load 60% (5 min average)
[1458875752] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1458875872] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 89% (5 min average)
[1458878193] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)