[1493697989] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 50% (15 min average)
[1493699402] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1493700388] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1493700448] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1493700508] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1493701108] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 67% (15 min average)
[1493701108] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1493701382] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;1;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1493701502] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;2;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1493701622] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;HARD;3;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1493701708] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1493701768] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1493701828] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1493701912] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 134 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493701972] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493702102] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1493702162] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1493702428] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 84% (15 min average)
[1493702762] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1493702822] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;2;Warning:  idm1  are more than 120 seconds behind!
[1493702883] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1493704082] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1493704682] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1493704742] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;2;Warning:  idm1  are more than 120 seconds behind!
[1493704802] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1493704828] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1493704888] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1493705402] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1493705402] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1493705488] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1493705548] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1493705608] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1493706002] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1493706208] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1493706459] SERVICE ALERT: FITNESSGRAM;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1493706519] SERVICE ALERT: FITNESSGRAM;CPU Load;OK;SOFT;2;CPU Load 2% (15 min average)
[1493706602] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1493706662] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;2;Warning:  idm1  are more than 120 seconds behind!
[1493706722] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1493707322] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1493707922] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1493708008] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 88% (15 min average)
[1493708222] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;OK;HARD;3;OK: No entries found
[1493708609] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1493708668] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1493708728] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1493709327] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1493709927] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1493709986] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1493710045] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1493712122] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1493712443] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 56% (15 min average)
[1493713322] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1493713382] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;2;Warning:  idm1  are more than 120 seconds behind!
[1493713442] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1493713800] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service has exited from a period of scheduled downtime
[1493713803] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (43.9% change >= 20.0% threshold)
[1493713803] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (48.9% change >= 20.0% threshold)
[1493713803] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1493714042] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1493715602] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (33.5% change >= 20.0% threshold)
[1493715602] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (47.4% change >= 20.0% threshold)
[1493715602] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1493717043] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1493717102] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;2;Warning:  idm1  are more than 120 seconds behind!
[1493717162] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1493717762] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1493719643] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1493719703] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1493719763] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1493720963] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1493721563] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1493722163] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 64% (15 min average)
[1493722763] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1493722823] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1493722883] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1493725283] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 86% (15 min average)
[1493726483] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1493726543] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1493726603] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1493727203] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 41% (15 min average)
[1493728204] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (42.4% change >= 20.0% threshold)
[1493728204] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1493730202] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Cyan Cartridge HP CE261A'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1493730262] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;OK;SOFT;2;Cyan Cartridge HP CE261A is at 11% - OK!
[1493731802] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (36.1% change >= 20.0% threshold)
[1493731803] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1493733956] SERVICE ALERT: AD-RMAD;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1493733996] SERVICE ALERT: AD-RMAD;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1493735403] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1493736202] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1493736262] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 85% (15 min average)
[1493740973] SERVICE ALERT: CAVA2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1493741033] SERVICE ALERT: CAVA2;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1493741093] SERVICE ALERT: CAVA2;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1493741693] SERVICE ALERT: CAVA2;CPU Load;OK;HARD;3;CPU Load 3% (15 min average)
[1493742493] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 99.90 Gb - used: 90.29 Gb (90%) - free 9.61 Gb (10%)
[1493742553] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 99.90 Gb - used: 90.29 Gb (90%) - free 9.61 Gb (10%)
[1493742613] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 99.90 Gb - used: 90.29 Gb (90%) - free 9.60 Gb (10%)
[1493744402] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1493745877] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1493745937] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1493745997] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1493751611] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;OK;HARD;3;c:\ - total: 99.90 Gb - used: 88.81 Gb (89%) - free 11.09 Gb (11%)
[1493759660] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1493759720] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1493759780] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 77% (15 min average)
[1493760380] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1493760440] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 92% (15 min average)
[1493764040] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1493764100] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1493764700] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1493764760] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 77% (15 min average)
[1493767159] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1493767219] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1493767279] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 94% (15 min average)
[1493773879] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1493773939] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1493775000] HOST DOWNTIME ALERT: IMAGES;STARTED; Host has entered a period of scheduled downtime
[1493776009] HOST ALERT: IMAGES;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1493776099] HOST ALERT: IMAGES;DOWN;HARD;2;PING CRITICAL - Packet loss = 100%
[1493776122] SERVICE ALERT: IMAGES;Service: MDT Monitor;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1493776122] SERVICE ALERT: IMAGES;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1493776122] SERVICE ALERT: IMAGES;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1493776122] SERVICE ALERT: IMAGES;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1493776125] SERVICE ALERT: IMAGES;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1493776143] SERVICE ALERT: IMAGES;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1493776143] SERVICE ALERT: IMAGES;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1493776147] SERVICE ALERT: IMAGES;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1493776156] SERVICE ALERT: IMAGES;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1493776156] SERVICE ALERT: IMAGES;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1493776180] SERVICE ALERT: IMAGES;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1493776180] SERVICE ALERT: IMAGES;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1493776180] SERVICE ALERT: IMAGES;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1493776180] SERVICE ALERT: IMAGES;Service: MDT Monitor;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1493776185] SERVICE ALERT: IMAGES;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1493776203] SERVICE ALERT: IMAGES;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1493776203] SERVICE ALERT: IMAGES;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1493776207] SERVICE ALERT: IMAGES;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1493776216] SERVICE ALERT: IMAGES;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1493776216] SERVICE ALERT: IMAGES;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1493776217] SERVICE ALERT: IMAGES;EMU Star SSL Cert Expirations;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1493776240] SERVICE ALERT: IMAGES;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1493776240] SERVICE ALERT: IMAGES;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1493776240] SERVICE ALERT: IMAGES;Service: MDT Monitor;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1493776240] SERVICE ALERT: IMAGES;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1493776245] SERVICE ALERT: IMAGES;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1493776263] SERVICE ALERT: IMAGES;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1493776263] SERVICE ALERT: IMAGES;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1493776267] SERVICE ALERT: IMAGES;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1493776276] SERVICE ALERT: IMAGES;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1493776276] SERVICE ALERT: IMAGES;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1493776282] SERVICE ALERT: IMAGES;Service: PXE Boot;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1493776285] SERVICE ALERT: IMAGES;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1493776287] SERVICE ALERT: IMAGES;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1493776304] SERVICE ALERT: IMAGES;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1493776342] SERVICE ALERT: IMAGES;Service: PXE Boot;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1493776345] SERVICE ALERT: IMAGES;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1493776347] SERVICE ALERT: IMAGES;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1493776364] SERVICE ALERT: IMAGES;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1493776390] SERVICE ALERT: IMAGES;PXE Boot Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1493776396] SERVICE ALERT: IMAGES;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1493776402] SERVICE ALERT: IMAGES;Service: PXE Boot;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1493776405] SERVICE ALERT: IMAGES;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1493776406] SERVICE ALERT: IMAGES;EMU Star SSL Cert Expirations;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1493776407] SERVICE ALERT: IMAGES;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1493776424] SERVICE ALERT: IMAGES;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1493776450] SERVICE ALERT: IMAGES;PXE Boot Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1493776456] SERVICE ALERT: IMAGES;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1493776510] SERVICE ALERT: IMAGES;PXE Boot Service;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1493776516] SERVICE ALERT: IMAGES;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1493776596] SERVICE ALERT: IMAGES;EMU Star SSL Cert Expirations;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1493776750] HOST ALERT: IMAGES;UP;HARD;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1493776821] SERVICE ALERT: IMAGES;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1493776821] SERVICE ALERT: IMAGES;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1493776821] SERVICE ALERT: IMAGES;Service: MDT Monitor;OK;HARD;3;MDT_Monitor: Started
[1493776821] SERVICE ALERT: IMAGES;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 2 minute(s)
[1493776825] SERVICE ALERT: IMAGES;Service: Windows Time;OK;HARD;3;W32Time: Started
[1493776843] SERVICE ALERT: IMAGES;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1493776843] SERVICE ALERT: IMAGES;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1493776847] SERVICE ALERT: IMAGES;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1493776856] SERVICE ALERT: IMAGES;C:\ Drive Space;OK;HARD;3;c:\ - total: 1157.56 Gb - used: 552.89 Gb (48%) - free 604.67 Gb (52%)
[1493776856] SERVICE ALERT: IMAGES;CPU Load;OK;HARD;3;CPU Load 0% (15 min average)
[1493776982] SERVICE ALERT: IMAGES;Service: PXE Boot;OK;HARD;3;WDSServer: Started
[1493776985] SERVICE ALERT: IMAGES;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1493776987] SERVICE ALERT: IMAGES;Service: RPC;OK;HARD;3;RpcSs: Started
[1493777004] SERVICE ALERT: IMAGES;Service: WMI;OK;HARD;3;Winmgmt: Started
[1493777011] SERVICE ALERT: IMAGES;EMU Star SSL Cert Expirations;OK;HARD;3;OK! All certificates are current.
[1493777090] SERVICE ALERT: IMAGES;PXE Boot Service;OK;HARD;3;WDSServer: Started
[1493777096] SERVICE ALERT: IMAGES;Memory Usage;OK;HARD;3;Memory usage: total:37586.86 MB - used: 1409.50 MB (4%) - free: 36177.37 MB (96%)
[1493780539] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1493780599] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1493780659] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1493781259] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 30% (15 min average)
[1493782142] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 202 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493782200] HOST DOWNTIME ALERT: IMAGES;STOPPED; Host has exited from a period of scheduled downtime
[1493782202] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493782363] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 159 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1493782423] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.