[1482730721] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1482730841] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1482730961] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1482732161] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1482732761] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1482733739] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1482733800] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1482733858] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1482733961] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 35% (15 min average)
[1482733979] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1482734578] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1482735161] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1482735281] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 83% (15 min average)
[1482735778] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1482735881] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1482736001] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1482736121] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1482736378] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1482736378] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1482736683] SERVICE ALERT: BBTS_DB;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1482736721] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1482736803] SERVICE ALERT: BBTS_DB;CPU Load;OK;SOFT;2;CPU Load 32% (15 min average)
[1482737321] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1482737441] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1482737561] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1482739547] SERVICE ALERT: epprogsql2;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1482739667] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1482739787] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1482740578] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1482741161] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 86% (15 min average)
[1482741178] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1482741761] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1482741779] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1482741881] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1482742001] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1482742378] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1482742498] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1482742618] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;3;OK. All replicas within 120 seconds.
[1482743801] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 83% (15 min average)
[1482744113] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 102 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482744232] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482744587] SERVICE ALERT: epprogsql2;CPU Load;OK;HARD;3;CPU Load 39% (15 min average)
[1482744603] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (43.2% change >= 20.0% threshold)
[1482744603] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (33.9% change >= 20.0% threshold)
[1482744603] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1482744603] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1482744603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1482744603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1482744603] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1482744603] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1482744603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1482744603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1482744603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1482744603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1482744603] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1482744603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1482744603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1482744603] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1482744603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1482744603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1482744603] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1482746402] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (32.8% change >= 20.0% threshold)
[1482746402] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (27.5% change >= 20.0% threshold)
[1482746402] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1482746402] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1482746402] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1482746402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1482746402] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1482746402] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1482746402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1482746402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1482746402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1482746402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1482746402] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1482746402] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1482746402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1482746402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1482746402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1482746402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1482746402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1482746507] SERVICE ALERT: EPPROGSQL2;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 249.66 Gb - used: 227.18 Gb (91%) - free 22.47 Gb (9%)
[1482746627] SERVICE ALERT: EPPROGSQL2;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 249.66 Gb - used: 227.19 Gb (91%) - free 22.47 Gb (9%)
[1482746747] SERVICE ALERT: EPPROGSQL2;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 249.66 Gb - used: 227.19 Gb (91%) - free 22.47 Gb (9%)
[1482750401] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1482750521] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1482750641] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1482751241] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1482751800] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service has exited from a period of scheduled downtime
[1482753418] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1482754043] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1271 found since last run.
[1482754163] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482754431] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1432 found since last run.
[1482754551] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482754763] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1424 found since last run.
[1482754883] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482755151] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1579 found since last run.
[1482755271] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482755440] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1482755483] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1623 found since last run.
[1482755603] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482755871] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1624 found since last run.
[1482755991] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482756039] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1482756158] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 74% (15 min average)
[1482756203] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1585 found since last run.
[1482756323] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482756591] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1625 found since last run.
[1482756711] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482756759] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1482756878] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 92% (15 min average)
[1482756923] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1632 found since last run.
[1482757043] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482757311] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1331 found since last run.
[1482757431] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482758078] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1482758197] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1482758316] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 55% (15 min average)
[1482762603] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1482762603] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1482762603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1482762603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1482762603] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1482762603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1482762603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1482762603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1482762603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1482762603] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1482762603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1482762603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1482762603] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1482762603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1482762603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1482762603] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1482764746] SERVICE ALERT: EPPROGSQL2;C:\ Drive Space;OK;HARD;3;c:\ - total: 249.66 Gb - used: 211.52 Gb (85%) - free 38.14 Gb (15%)
[1482773403] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1482773403] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1482773403] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1482773403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1482773403] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1482773403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1482773403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1482773403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1482773403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1482773403] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1482773403] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1482773403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1482773403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1482773403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1482773403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1482773403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1482791915] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1482792035] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 85% (15 min average)
[1482793235] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1482793355] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 31% (15 min average)
[1482793375] SERVICE ALERT: cava2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1482793495] SERVICE ALERT: cava2;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1482793615] SERVICE ALERT: cava2;CPU Load;OK;SOFT;3;CPU Load 0% (15 min average)
[1482793955] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1482794075] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1482794195] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1482795395] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 77% (15 min average)
[1482795995] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1482796115] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 86% (15 min average)
[1482798230] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 117 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482798350] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482799041] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 100 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482799161] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482799715] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1482799835] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1482799955] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 94% (15 min average)
[1482800555] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1482800675] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 74% (15 min average)
[1482804649] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;SOFT;1;e:\ - total: 799.87 Gb - used: 658.50 Gb (82%) - free 141.37 Gb (18%)
[1482804768] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;SOFT;2;e:\ - total: 799.87 Gb - used: 658.50 Gb (82%) - free 141.37 Gb (18%)
[1482804889] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;HARD;3;e:\ - total: 799.87 Gb - used: 658.50 Gb (82%) - free 141.37 Gb (18%)
[1482805111] SERVICE ALERT: OFFICEKMS;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1482805231] SERVICE ALERT: OFFICEKMS;CPU Load;OK;SOFT;2;CPU Load 1% (15 min average)
[1482807288] SERVICE ALERT: AD-RMAD;E:\ Drive Space;OK;HARD;3;e:\ - total: 799.87 Gb - used: 635.03 Gb (79%) - free 164.84 Gb (21%)
[1482807631] SERVICE ALERT: OFFICEKMS;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1482807751] SERVICE ALERT: OFFICEKMS;CPU Load;OK;SOFT;2;CPU Load 2% (15 min average)
[1482811474] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1482811594] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1482811714] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1482814714] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 70% (15 min average)