[1485495992] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1485496112] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1485496712] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1485496832] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1485496951] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1485497464] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1485497550] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1485497585] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1485498150] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1485498186] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1485498304] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1485498425] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1485498601] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1485498749] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1485498869] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1485498989] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1485499026] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1485500224] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1485500275] SERVICE ALERT: SANDBOX-DB;C:\ Drive Space;OK;HARD;3;c:\ - total: 39.66 Gb - used: 35.24 Gb (89%) - free 4.42 Gb (11%)
[1485500314] SERVICE ALERT: NETENG-PA;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1485500344] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;2;Warning:  idm1  are more than 120 seconds behind!
[1485500433] SERVICE ALERT: NETENG-PA;CPU Load;OK;SOFT;2;CPU Load 43% (15 min average)
[1485500464] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;3;OK. All replicas within 120 seconds.
[1485501064] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1485501143] SERVICE ALERT: BBTS_DB;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1485501184] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1485501263] SERVICE ALERT: BBTS_DB;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1485501304] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1485501389] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1485501990] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1485503189] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 90% (15 min average)
[1485503789] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1485503909] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 87% (15 min average)
[1485504033] SERVICE ALERT: NETENG-PA;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1485504153] SERVICE ALERT: NETENG-PA;CPU Load;OK;SOFT;2;CPU Load 8% (15 min average)
[1485504304] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1485504304] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1485504509] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1485504629] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1485504749] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1485504904] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1485505349] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1485506105] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1485506704] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1485507304] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1485508346] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 61% (15 min average)
[1485508504] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1485509104] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1485509402] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (47.2% change >= 20.0% threshold)
[1485509402] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (52.8% change >= 20.0% threshold)
[1485509402] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1485509402] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1485509402] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1485509402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1485509402] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1485509402] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1485509402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1485509402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1485509402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1485509402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1485509402] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1485509402] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1485509402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1485509402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1485509402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1485509402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1485509402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1485511202] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (32.8% change >= 20.0% threshold)
[1485511203] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (45.2% change >= 20.0% threshold)
[1485511203] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1485511203] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1485511203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1485511203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1485511203] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1485511203] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1485511203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1485511203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1485511203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1485511203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1485511203] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1485511203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1485511203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1485511203] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1485511203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1485511203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1485511203] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1485511346] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1485511464] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 92% (15 min average)
[1485515664] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1485515784] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1485515904] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1485516601] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service has exited from a period of scheduled downtime
[1485517704] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1485518304] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1485518424] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1485518544] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1485519744] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (15 min average)
[1485520344] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1485520464] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1485520503] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1485520584] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1485521154] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1485521213] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1485521274] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1485521784] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1485522801] SERVICE ALERT: IMAGES;Service: PXE Boot;OK;HARD;3;WDSServer: Started
[1485522865] SERVICE ALERT: IMAGES;PXE Boot Service;OK;HARD;3;WDSServer: Started
[1485523703] SERVICE ALERT: R25-WEBVIEWER;Service: Firewall;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1485523720] SERVICE ALERT: R25-WEBVIEWER;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1485523823] SERVICE ALERT: R25-WEBVIEWER;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1485523831] SERVICE ALERT: R25-WEBVIEWER;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1485523963] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at WARNING level!
[1485524023] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at WARNING level!
[1485524083] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at WARNING level!
[1485524184] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1485524304] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1485525523] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1485525584] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1485525644] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1485525883] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1485526243] SERVICE ALERT: marshallpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1485526673] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1485527403] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (27.1% change >= 20.0% threshold)
[1485527404] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1485527404] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1485527404] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1485527404] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1485527404] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1485527404] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1485527404] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1485527404] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1485527404] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1485527404] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1485527404] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1485527404] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1485527404] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1485527404] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1485527404] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1485527404] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1485529483] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at WARNING level!
[1485529543] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at WARNING level!
[1485529603] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at WARNING level!
[1485531227] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5076 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1485531347] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1485532704] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1485532802] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1485532802] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1485532802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1485532802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1485532802] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1485532802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1485532802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1485532802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1485532802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1485532802] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1485532802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1485532802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1485532802] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1485532802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1485532802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1485532802] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1485533007] SERVICE ALERT: ADDC5;GPO Replication;WARNING;SOFT;1;Warning! 1 policies are not replicating properly!
[1485533204] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1485533204] SERVICE FLAPPING ALERT: prayharroldpk;Black Print Cartridge HP C8543X;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1485533306] SERVICE ALERT: ADDC5;GPO Replication;WARNING;SOFT;2;Warning! 1 policies are not replicating properly!
[1485533606] SERVICE ALERT: ADDC5;GPO Replication;WARNING;HARD;3;Warning! 2 policies are not replicating properly!
[1485533813] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5384 found since last run. oneLevelSearchOps OK. searchOps is too high! 5502 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1485533933] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1485534521] SERVICE ALERT: CAVA2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1485534641] SERVICE ALERT: CAVA2;CPU Load;OK;SOFT;2;CPU Load 0% (15 min average)
[1485534947] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5015 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1485535067] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1485535133] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5196 found since last run. oneLevelSearchOps OK. searchOps is too high! 5253 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1485535253] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1485535373] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1485535433] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1485535493] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1485535853] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7502 found since last run. oneLevelSearchOps OK. searchOps is too high! 7671 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1485535973] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1485536441] SERVICE ALERT: CAVA2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1485536561] SERVICE ALERT: CAVA2;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1485536573] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6657 found since last run. oneLevelSearchOps OK. searchOps is too high! 6849 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1485536681] SERVICE ALERT: CAVA2;CPU Load;OK;SOFT;3;CPU Load 0% (15 min average)
[1485536692] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1485537292] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5750 found since last run. oneLevelSearchOps OK. searchOps is too high! 5894 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1485537412] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1485537892] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;There is no help available for this item. 54 6F 20 65 78 69 74 20 70 72 65 73 73 20 1F 2E
[1485538002] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at WARNING level!
[1485538062] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at WARNING level!
[1485538122] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at WARNING level!
[1485538203] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1485538203] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1485538203] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1485538203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1485538203] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1485538203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1485538203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1485538203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1485538203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1485538203] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1485538203] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1485538203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1485538203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1485538203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1485538203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1485538203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1485538492] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1485538552] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1485538612] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5145 found since last run. oneLevelSearchOps OK. searchOps is too high! 5410 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1485538612] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1485538732] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1485538977] SERVICE ALERT: cava2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1485539097] SERVICE ALERT: cava2;CPU Load;OK;SOFT;2;CPU Load 3% (15 min average)
[1485539332] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5917 found since last run. oneLevelSearchOps OK. searchOps is too high! 6166 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1485539452] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1485539812] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Tray 4 PLAIN LETTER To change size or 74 79 70 65 20 70 72 65 73 73 20 1E  Tray 4 PLAIN LETTER To accept settings 70 72 65 73 73 20 18
[1485539812] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1485539904] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1485540024] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1485540052] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7825 found since last run. oneLevelSearchOps OK. searchOps is too high! 8046 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1485540144] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1485540172] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1485540744] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1485541372] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5444 found since last run. oneLevelSearchOps OK. searchOps is too high! 5735 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1485541492] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1485541803] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1485541803] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1485541803] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1485541803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1485541803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1485541803] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1485541803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1485541803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1485541803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1485541803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1485541803] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1485541803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1485541803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1485541803] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1485541803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1485541803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1485541803] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1485543144] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1485543264] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 66% (15 min average)
[1485543892] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6189 found since last run. oneLevelSearchOps OK. searchOps is too high! 6295 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1485544012] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1485544338] SERVICE ALERT: XDM;Service: WMI;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1485544458] SERVICE ALERT: XDM;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1485544612] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5653 found since last run. oneLevelSearchOps OK. searchOps is too high! 5817 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1485544732] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1485545461] SERVICE ALERT: MBAMDB;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1485545561] SERVICE ALERT: MBAMDB;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1485547205] SERVICE ALERT: BBTS_DB;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.227.4 and port 12489: Connection refused
[1485547325] SERVICE ALERT: BBTS_DB;C:\ Drive Space;OK;SOFT;2;c:\ - total: 144.98 Gb - used: 87.30 Gb (60%) - free 57.68 Gb (40%)
[1485547667] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6049 found since last run. oneLevelSearchOps OK. searchOps is too high! 6146 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1485547787] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1485548010] SERVICE ALERT: ADDC5;GPO Replication;OK;HARD;3;OK! All GPOs are properly replicated.
[1485548932] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5089 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1485549052] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1485550252] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5223 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1485550372] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1485550463] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1485550583] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1485550610] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1485550703] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1485550787] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5305 found since last run. oneLevelSearchOps OK. searchOps is too high! 5311 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1485550907] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1485551303] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1485551303] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1485553703] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1485553823] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 55% (15 min average)
[1485554423] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1485554543] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1485554663] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1485555263] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 88% (15 min average)
[1485555863] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1485555983] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 88% (15 min average)
[1485557183] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1485557303] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1485557423] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1485558393] SERVICE ALERT: ADRODC-DET;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1485558398] SERVICE ALERT: ADRODC-LIV;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1485558404] HOST ALERT: Google;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1485558423] HOST ALERT: ADRODC-DET;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1485558428] HOST ALERT: ADRODC-LIV;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1485558494] HOST ALERT: Google;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1485558513] HOST ALERT: ADRODC-DET;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1485558513] SERVICE ALERT: ADRODC-DET;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1485558518] SERVICE ALERT: ADRODC-LIV;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1485558518] HOST ALERT: ADRODC-LIV;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1485558543] HOST ALERT: ADRODC-DET;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1485558573] SERVICE ALERT: ADRODC-DET;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1485558584] HOST ALERT: Google;DOWN;HARD;3;PING CRITICAL - Packet loss = 100%
[1485558603] HOST ALERT: ADRODC-DET;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1485558608] HOST ALERT: ADRODC-LIV;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1485558623] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1485558633] SERVICE ALERT: ADRODC-DET;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1485558636] SERVICE ALERT: ADRODC-DET;Service: Intersite Messaging;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1485558636] SERVICE ALERT: ADRODC-DET;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1485558637] SERVICE ALERT: ADRODC-DET;Service: PXE Boot;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1485558638] SERVICE ALERT: ADRODC-LIV;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1485558638] SERVICE ALERT: ADRODC-DET;Service: AD Domain Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1485558639] SERVICE ALERT: ADRODC-DET;Service: Remote Management;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1485558643] SERVICE ALERT: ADRODC-DET;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1485558643] SERVICE ALERT: ADRODC-DET;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1485558644] SERVICE ALERT: ADRODC-LIV;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1485558645] SERVICE ALERT: ADRODC-DET;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1485558646] SERVICE ALERT: ADRODC-DET;Service: DNS Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1485558661] SERVICE ALERT: ADRODC-DET;Service: Distributed File System Replication;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1485558663] HOST ALERT: ADRODC-DET;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1485558668] HOST ALERT: ADRODC-LIV;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1485558668] SERVICE ALERT: ADRODC-DET;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1485558673] SERVICE ALERT: ADRODC-DET;Service: AD Web Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1485558676] SERVICE ALERT: ADRODC-DET;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1485558676] SERVICE ALERT: ADRODC-DET;Service: MDT Monitor;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1485558683] SERVICE ALERT: ADRODC-DET;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1485558693] SERVICE ALERT: ADRODC-DET;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1485558698] HOST ALERT: ADRODC-DET;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1485558706] SERVICE ALERT: ADRODC-LIV;Service: PXE Boot;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1485558708] SERVICE ALERT: ADRODC-LIV;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1485558711] SERVICE ALERT: ADRODC-DET;Service: DFS Namespace;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1485558713] SERVICE ALERT: ADRODC-LIV;Service: MDT Monitor;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1485558733] SERVICE ALERT: ADRODC-LIV;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1485558733] SERVICE ALERT: ADRODC-LIV;Service: DFS Namespace;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1485558735] SERVICE ALERT: ADRODC-LIV;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1485558736] HOST ALERT: ADRODC-LIV;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1485558741] HOST ALERT: ADRODC-DET;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1485558743] SERVICE ALERT: ADRODC-LIV;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1485558743] SERVICE ALERT: ADRODC-DET;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1485558743] SERVICE ALERT: ADRODC-DET;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1485558743] SERVICE ALERT: ADRODC-LIV;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1485558751] SERVICE ALERT: ADRODC-DET;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1485558756] SERVICE ALERT: ADRODC-DET;Service: Intersite Messaging;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1485558756] SERVICE ALERT: ADRODC-DET;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1485558757] SERVICE ALERT: ADRODC-DET;Service: PXE Boot;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1485558758] SERVICE ALERT: ADRODC-LIV;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1485558758] SERVICE ALERT: ADRODC-DET;Service: AD Domain Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1485558759] SERVICE ALERT: ADRODC-DET;Service: Remote Management;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1485558763] SERVICE ALERT: ADRODC-DET;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1485558763] SERVICE ALERT: ADRODC-DET;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1485558764] SERVICE ALERT: ADRODC-LIV;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1485558765] SERVICE ALERT: ADRODC-DET;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1485558766] SERVICE ALERT: ADRODC-DET;Service: DNS Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1485558767] SERVICE ALERT: ADRODC-LIV;Service: Distributed File System Replication;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1485558769] SERVICE ALERT: ADRODC-LIV;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1485558770] SERVICE ALERT: ADRODC-LIV;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1485558772] SERVICE ALERT: ADRODC-DET;AD Replication;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 30 seconds.
[1485558773] HOST ALERT: ADRODC-LIV;DOWN;SOFT;6;(Host check timed out after 30.00 seconds)
[1485558773] HOST ALERT: ADRODC-DET;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1485558778] SERVICE ALERT: ADRODC-LIV;AD Replication;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 30 seconds.
[1485558781] SERVICE ALERT: ADRODC-DET;Service: Distributed File System Replication;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1485558786] SERVICE ALERT: ADRODC-LIV;Service: AD Web Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1485558788] SERVICE ALERT: ADRODC-DET;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1485558793] SERVICE ALERT: ADRODC-DET;Service: AD Web Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1485558796] SERVICE ALERT: ADRODC-DET;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1485558796] SERVICE ALERT: ADRODC-DET;Service: MDT Monitor;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1485558803] SERVICE ALERT: ADRODC-DET;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1485558808] HOST ALERT: ADRODC-LIV;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1485558811] HOST ALERT: ADRODC-DET;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1485558813] SERVICE ALERT: ADRODC-DET;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1485558826] SERVICE ALERT: ADRODC-LIV;Service: PXE Boot;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1485558828] SERVICE ALERT: ADRODC-LIV;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1485558828] SERVICE ALERT: ADRODC-LIV;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1485558831] SERVICE ALERT: ADRODC-DET;Service: DFS Namespace;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1485558832] SERVICE ALERT: ADRODC-LIV;Service: MDT Monitor;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1485558843] SERVICE ALERT: ADRODC-LIV;Service: AD Domain Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1485558843] HOST ALERT: ADRODC-DET;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1485558852] SERVICE ALERT: ADRODC-LIV;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1485558853] SERVICE ALERT: ADRODC-LIV;Service: DFS Namespace;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1485558854] SERVICE ALERT: ADRODC-LIV;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1485558856] HOST ALERT: ADRODC-LIV;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1485558863] SERVICE ALERT: ADRODC-LIV;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1485558863] SERVICE ALERT: ADRODC-DET;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1485558863] SERVICE ALERT: ADRODC-DET;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1485558863] SERVICE ALERT: ADRODC-DET;Service: Key Distribution Center;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1485558863] SERVICE ALERT: ADRODC-LIV;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1485558864] SERVICE ALERT: ADRODC-LIV;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1485558871] SERVICE ALERT: ADRODC-DET;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1485558875] SERVICE ALERT: ADRODC-DET;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1485558876] SERVICE ALERT: ADRODC-DET;Service: Intersite Messaging;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1485558876] SERVICE ALERT: ADRODC-DET;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1485558877] SERVICE ALERT: ADRODC-DET;Service: PXE Boot;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1485558878] SERVICE ALERT: ADRODC-LIV;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1485558878] SERVICE ALERT: ADRODC-DET;Service: AD Domain Service;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1485558879] SERVICE ALERT: ADRODC-DET;Service: Remote Management;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1485558883] SERVICE ALERT: ADRODC-DET;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1485558883] SERVICE ALERT: ADRODC-DET;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1485558884] SERVICE ALERT: ADRODC-LIV;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1485558885] SERVICE ALERT: ADRODC-DET;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1485558886] SERVICE ALERT: ADRODC-DET;Service: DNS Service;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1485558887] SERVICE ALERT: ADRODC-LIV;Service: Key Distribution Center;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1485558887] SERVICE ALERT: ADRODC-LIV;Service: Distributed File System Replication;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1485558889] SERVICE ALERT: ADRODC-LIV;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1485558890] SERVICE ALERT: ADRODC-LIV;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1485558892] SERVICE ALERT: ADRODC-DET;AD Replication;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 30 seconds.
[1485558893] HOST ALERT: ADRODC-LIV;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1485558898] SERVICE ALERT: ADRODC-LIV;AD Replication;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 30 seconds.
[1485558901] SERVICE ALERT: ADRODC-DET;Service: Distributed File System Replication;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1485558901] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1485558905] SERVICE ALERT: ADRODC-LIV;Service: DNS Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1485558905] SERVICE ALERT: ADRODC-LIV;Service: Intersite Messaging;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1485558905] SERVICE ALERT: ADRODC-LIV;Service: Remote Management;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1485558906] SERVICE ALERT: ADRODC-LIV;Service: AD Web Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1485558908] SERVICE ALERT: ADRODC-DET;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1485558913] SERVICE ALERT: ADRODC-DET;Service: AD Web Service;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1485558916] SERVICE ALERT: ADRODC-DET;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1485558916] SERVICE ALERT: ADRODC-DET;Service: MDT Monitor;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1485558923] SERVICE ALERT: ADRODC-DET;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1485558923] SERVICE ALERT: ADRODC-LIV;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1485558928] HOST ALERT: ADRODC-LIV;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1485558935] SERVICE ALERT: ADRODC-LIV;Service: PXE Boot;OK;SOFT;3;WDSServer: Started
[1485558936] SERVICE ALERT: ADRODC-LIV;Service: MDT Monitor;OK;SOFT;3;MDT_Monitor: Started
[1485558937] SERVICE ALERT: ADRODC-LIV;Service: Firewall;OK;SOFT;3;MpsSvc: Started
[1485558937] SERVICE ALERT: ADRODC-LIV;CPU Load;OK;SOFT;2;CPU Load 0% (15 min average)
[1485558939] HOST ALERT: ADRODC-LIV;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 6.00 ms
[1485558940] SERVICE ALERT: ADRODC-DET;Service: DFS Namespace;OK;SOFT;3;Dfs: Started
[1485558940] HOST ALERT: Google;UP;HARD;3;PING WARNING - Packet loss = 81%, RTA = 16.61 ms
[1485558943] SERVICE ALERT: ADRODC-LIV;Service: AD Domain Service;OK;SOFT;2;NTDS: Started
[1485558944] HOST ALERT: ADRODC-DET;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 8.80 ms
[1485558951] SERVICE ALERT: ADRODC-LIV;Service: Windows Time;OK;SOFT;3;W32Time: Started
[1485558953] SERVICE ALERT: ADRODC-LIV;Service: DFS Namespace;OK;SOFT;3;Dfs: Started
[1485558954] SERVICE ALERT: ADRODC-LIV;Uptime;OK;SOFT;3;System Uptime - 38 day(s) 10 hour(s) 13 minute(s)
[1485558963] SERVICE ALERT: ADRODC-LIV;Service: WMI;OK;SOFT;3;Winmgmt: Started
[1485558963] SERVICE ALERT: ADRODC-DET;Service: NetLogin;OK;SOFT;3;Netlogon: Started
[1485558963] SERVICE ALERT: ADRODC-DET;Service: Firewall;OK;SOFT;3;MpsSvc: Started
[1485558963] SERVICE ALERT: ADRODC-DET;Service: Key Distribution Center;OK;SOFT;2;KDC: Started
[1485558963] SERVICE ALERT: ADRODC-LIV;Service: User Profile;OK;SOFT;3;ProfSvc: Started
[1485558964] SERVICE ALERT: ADRODC-LIV;Memory Usage;OK;SOFT;2;Memory usage: total:37586.86 MB - used: 4624.12 MB (12%) - free: 32962.74 MB (88%)
[1485558971] SERVICE ALERT: ADRODC-DET;Service: Windows Time;OK;SOFT;3;W32Time: Started
[1485558975] SERVICE ALERT: ADRODC-DET;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1485558978] SERVICE ALERT: ADRODC-LIV;C:\ Drive Space;OK;SOFT;3;c:\ - total: 1115.16 Gb - used: 559.93 Gb (50%) - free 555.23 Gb (50%)
[1485558987] SERVICE ALERT: ADRODC-LIV;Service: Key Distribution Center;OK;SOFT;2;KDC: Started
[1485558987] SERVICE ALERT: ADRODC-LIV;Service: Distributed File System Replication;OK;SOFT;3;DFSR: Started
[1485558988] SERVICE ALERT: ADRODC-DET;AD Replication;OK;SOFT;3;OK: replication is up and running
[1485558989] SERVICE ALERT: ADRODC-LIV;NSClient++ Version;OK;SOFT;3;NSClient++ 0.4.4.15 2015-11-25
[1485558990] SERVICE ALERT: ADRODC-LIV;Service: DNS Client;OK;SOFT;3;Dnscache: Started
[1485558994] SERVICE ALERT: ADRODC-LIV;AD Replication;OK;SOFT;3;OK: replication is up and running
[1485559005] SERVICE ALERT: ADRODC-LIV;Service: DNS Service;OK;SOFT;2;DNS: Started
[1485559005] SERVICE ALERT: ADRODC-LIV;Service: Intersite Messaging;OK;SOFT;2;IsmServ: Started
[1485559005] SERVICE ALERT: ADRODC-LIV;Service: Remote Management;OK;SOFT;2;WinRM: Started
[1485559006] SERVICE ALERT: ADRODC-LIV;Service: AD Web Service;OK;SOFT;3;ADWS: Started
[1485559021] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1485559023] SERVICE ALERT: ADRODC-LIV;Service: RPC;OK;SOFT;2;RpcSs: Started
[1485559213] SERVICE ALERT: ADRODC-DET;Uptime;OK;HARD;3;System Uptime - 43 day(s) 6 hour(s) 18 minute(s)
[1485559218] SERVICE ALERT: ADRODC-LIV;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1485559223] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 75% (15 min average)
[1485559393] SERVICE ALERT: ADRODC-DET;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1485559443] SERVICE ALERT: idmrbs.emich.edu;LDAP Stats;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1485559443] SERVICE ALERT: idm4.emich.edu;eDirectory Service;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1485559456] SERVICE ALERT: ADRODC-DET;Service: Intersite Messaging;OK;HARD;3;IsmServ: Started
[1485559456] SERVICE ALERT: ADRODC-DET;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1485559457] SERVICE ALERT: ADRODC-DET;Service: PXE Boot;OK;HARD;3;WDSServer: Started
[1485559458] SERVICE ALERT: ADRODC-DET;Service: AD Domain Service;OK;HARD;3;NTDS: Started
[1485559459] SERVICE ALERT: ADRODC-DET;Service: Remote Management;OK;HARD;3;WinRM: Started
[1485559460] SERVICE ALERT: idm2.emich.edu;LDAP Stats;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1485559463] SERVICE ALERT: ADRODC-DET;Service: RPC;OK;HARD;3;RpcSs: Started
[1485559463] SERVICE ALERT: ADRODC-DET;CPU Load;OK;HARD;3;CPU Load 0% (15 min average)
[1485559463] SERVICE ALERT: idm1.emich.edu;eDirectory Service;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1485559463] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 364;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 30 seconds.
[1485559464] SERVICE ALERT: ADRODC-LIV;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1485559465] SERVICE ALERT: ADRODC-DET;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1485559466] SERVICE ALERT: ADRODC-DET;Service: DNS Service;OK;HARD;3;DNS: Started
[1485559481] SERVICE ALERT: ADRODC-DET;Service: Distributed File System Replication;OK;HARD;3;DFSR: Started
[1485559482] SERVICE ALERT: ldap.emich.edu;LDAP Stats;CRITICAL;SOFT;1;18:24:42  - (FAILURE!) LDAP Failed
[1485559488] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10042;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 30 seconds.
[1485559488] SERVICE ALERT: ADRODC-DET;Memory Usage;OK;HARD;3;Memory usage: total:37586.86 MB - used: 3912.26 MB (10%) - free: 33674.60 MB (90%)
[1485559493] SERVICE ALERT: ADRODC-DET;Service: AD Web Service;OK;HARD;3;ADWS: Started
[1485559496] SERVICE ALERT: ADRODC-DET;Service: WMI;OK;HARD;3;Winmgmt: Started
[1485559496] SERVICE ALERT: ADRODC-DET;Service: MDT Monitor;OK;HARD;3;MDT_Monitor: Started
[1485559499] SERVICE ALERT: LANSWEEPER;SSL Certificate 14 Days With Auto Renew;WARNING;SOFT;1;(No output on stdout) stderr: Could not resolve hostname LANSWEEPER: Name or service not known
[1485559499] SERVICE ALERT: idm1.emich.edu;OfficeAD Driver;WARNING;SOFT;1;(No output on stdout) stderr: Could not resolve hostname idm1.emich.edu: Name or service not known
[1485559503] SERVICE ALERT: ADRODC-DET;C:\ Drive Space;OK;HARD;3;c:\ - total: 1115.16 Gb - used: 279.90 Gb (25%) - free 835.26 Gb (75%)
[1485559503] SERVICE ALERT: idmrbs.emich.edu;eDirectory Service;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1485559508] SERVICE ALERT: ADDC4;AD Replication;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 30 seconds.
[1485559524] SERVICE ALERT: ADDC5;AD Group Changes;WARNING;SOFT;1;(No output on stdout) stderr: Could not resolve hostname ADDC5: Name or service not known
[1485559542] SERVICE ALERT: ldap.emich.edu;LDAP Stats;CRITICAL;HARD;2;18:25:42  - (FAILURE!) LDAP Failed
[1485559561] SERVICE ALERT: idmrbs.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1485559561] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 364;OK;SOFT;2;Eventlog check ok
[1485559562] SERVICE ALERT: idm4.emich.edu;eDirectory Service;OK;SOFT;2;OK: ndsd status is good.
[1485559564] SERVICE ALERT: LANSWEEPER;SSL Certificate 14 Days With Auto Renew;OK;SOFT;2;OK! All certificates are current.
[1485559564] SERVICE ALERT: idm1.emich.edu;OfficeAD Driver;OK;SOFT;2;OK: cn=OfficeAD,cn=driverset1,o=system is running.
[1485559570] SERVICE ALERT: idm2.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1485559573] SERVICE ALERT: idm1.emich.edu;eDirectory Service;OK;SOFT;2;OK: ndsd status is good.
[1485559578] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10042;OK;SOFT;2;Eventlog check ok
[1485559604] SERVICE ALERT: ADDC4;AD Replication;OK;SOFT;2;OK: replication is up and running
[1485559606] SERVICE ALERT: ldap.emich.edu;LDAP Stats;OK;HARD;2;18:26:46  - Destination LDAP online and available.
[1485559613] SERVICE ALERT: idmrbs.emich.edu;eDirectory Service;OK;SOFT;2;OK: ndsd status is good.
[1485559697] SERVICE ALERT: ADDC5;AD Group Changes;OK;SOFT;2;WARNING:  0  groups have changed!
[1485559823] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1485559943] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1485560063] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 93% (15 min average)
[1485561263] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1485561383] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1485561503] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 54% (15 min average)
[1485561849] SERVICE ALERT: hallegroundcolor;Black Cartridge HP CE260X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1485561903] SERVICE ALERT: hallegroundcolor;Black Cartridge HP CE260X;OK;SOFT;2;Black Cartridge HP CE260X is at 73% - OK!
[1485563903] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1485564023] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1485564143] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 70% (15 min average)
[1485566543] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1485566663] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1485566783] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 90% (15 min average)
[1485567383] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1485567503] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 92% (15 min average)
[1485571101] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1485571222] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1485571340] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1485571660] HOST ALERT: IMAGES;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1485571724] HOST ALERT: IMAGES;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.18 ms
[1485571743] SERVICE ALERT: strongpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1485571804] SERVICE ALERT: strongpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1485571863] SERVICE ALERT: strongpk;Device Status;OK;SOFT;3;50 72 65 73 73 20 17 20 74 6F 20 73 65 65 20 73  74 65 70  by step information
[1485571940] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 71% (15 min average)
[1485573140] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1485573260] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1485573380] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1485575180] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1485576980] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1485577100] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1485577220] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1485577820] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1485579056] SERVICE ALERT: hallegroundcolor;Yellow Cartridge HP CE262A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1485579110] SERVICE ALERT: hallegroundcolor;Yellow Cartridge HP CE262A;OK;SOFT;2;Yellow Cartridge HP CE262A is at 97% - OK!