[1459744055] HOST ALERT: halleground1;UP;HARD;10;PING WARNING - Packet loss = 96%, RTA = 0.46 ms
[1459744159] HOST ALERT: halleground2;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.56 ms
[1459744201] SERVICE ALERT: NETIDADPROXY;Service: DNS Client;CRITICAL;SOFT;1;Dnscache: Stopped
[1459744321] SERVICE ALERT: NETIDADPROXY;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1459748756] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1459749193] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1459749313] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 92% (5 min average)
[1459749433] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 72% (5 min average)
[1459750033] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1459750153] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 87% (5 min average)
[1459750800] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;SOFT;1;WARNING: Replication with ADDC5 failed!
[1459750919] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;SOFT;2;WARNING: Replication with ADDC5 failed!
[1459751040] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;HARD;3;WARNING: Replication with ADDC5 failed!
[1459751353] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1459751473] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 85% (5 min average)
[1459751756] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1459752356] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1459752673] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 94% (5 min average)
[1459752793] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 76% (5 min average)
[1459753393] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1459753513] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 88% (5 min average)
[1459754113] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1459755313] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1459755433] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 77% (5 min average)
[1459756033] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1459756153] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 76% (5 min average)
[1459757756] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 68% (5 min average)
[1459760156] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459760276] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1459760396] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 93% (5 min average)
[1459760396] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1459760996] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 96% (5 min average)
[1459761244] SERVICE ALERT: ADRODC-LIV;AD Replication;OK;HARD;3;OK: replication is up and running
[1459761596] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 62% (5 min average)
[1459763994] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1459764114] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 79% (5 min average)
[1459764714] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459764834] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1459764954] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1459765805] HOST ALERT: halleground1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459765865] HOST ALERT: halleground1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1459765925] HOST ALERT: halleground1;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1459766015] HOST ALERT: halleground1;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1459766105] HOST ALERT: halleground1;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1459766194] HOST ALERT: halleground1;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1459766284] HOST ALERT: halleground1;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1459766374] HOST ALERT: halleground1;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1459766405] HOST ALERT: halleground1;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1459766439] HOST ALERT: halleground1;UP;SOFT;10;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1459766525] HOST ALERT: halleground1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459766615] HOST ALERT: halleground1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1459766705] HOST ALERT: halleground1;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1459766794] HOST ALERT: halleground1;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1459766884] HOST ALERT: halleground1;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1459766974] HOST ALERT: halleground1;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1459766993] HOST ALERT: halleground1;UP;SOFT;7;PING OK - Packet loss = 73%, RTA = 0.57 ms
[1459767065] HOST ALERT: halleground1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459767065] HOST FLAPPING ALERT: halleground1;STARTED; Host appears to have started flapping (22.1% change > 20.0% threshold)
[1459767125] HOST ALERT: halleground1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1459767167] HOST ALERT: halleground2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459767215] HOST ALERT: halleground1;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1459767255] HOST ALERT: halleground2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1459767305] HOST ALERT: halleground1;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1459767345] HOST ALERT: halleground2;DOWN;SOFT;3;(Host check timed out after 30.00 seconds)
[1459767354] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1459767394] HOST ALERT: halleground1;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1459767435] HOST ALERT: halleground2;DOWN;SOFT;4;(Host check timed out after 30.00 seconds)
[1459767484] HOST ALERT: halleground1;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1459767525] HOST ALERT: halleground2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1459767548] HOST ALERT: halleground1;UP;SOFT;7;PING OK - Packet loss = 0%, RTA = 0.55 ms
[1459767605] HOST ALERT: halleground1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459767614] HOST ALERT: halleground2;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1459767665] HOST ALERT: halleground1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1459767704] HOST ALERT: halleground2;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1459767725] HOST ALERT: halleground1;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1459767767] HOST ALERT: halleground2;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1459767815] HOST ALERT: halleground1;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1459767830] HOST ALERT: halleground2;UP;SOFT;9;PING OK - Packet loss = 16%, RTA = 0.85 ms
[1459767905] HOST ALERT: halleground1;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1459767954] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 87% (5 min average)
[1459767995] HOST ALERT: halleground1;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1459768085] HOST ALERT: halleground1;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1459768160] HOST ALERT: halleground2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459768175] HOST ALERT: halleground1;DOWN;SOFT;8;(Host check timed out after 30.00 seconds)
[1459768205] HOST ALERT: halleground1;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1459768250] HOST ALERT: halleground2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1459768265] HOST ALERT: halleground1;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1459768340] HOST ALERT: halleground2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1459768375] HOST ALERT: halleground2;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1459768455] HOST ALERT: halleground2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1459768545] HOST ALERT: halleground2;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1459768635] HOST ALERT: halleground2;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1459768724] HOST ALERT: halleground2;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1459768805] HOST FLAPPING ALERT: halleground1;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1459768814] HOST ALERT: halleground2;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1459768904] HOST ALERT: halleground2;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1459769154] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459769274] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1459769394] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 97% (5 min average)
[1459769999] SERVICE ALERT: MBAM;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459770103] SERVICE ALERT: VCENTERSSO;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459770119] SERVICE ALERT: MBAM;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1459770223] SERVICE ALERT: VCENTERSSO;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1459771390] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1459771449] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1459771510] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1459771779] HOST ALERT: halleground1;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.60 ms
[1459771785] HOST ALERT: halleground2;UP;HARD;10;PING WARNING - Packet loss = 96%, RTA = 25.77 ms
[1459772109] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1459772891] SERVICE ALERT: owen101lab3;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1459772921] HOST ALERT: owen101lab3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459772923] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1459772951] SERVICE ALERT: owen101lab3;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1459772953] HOST ALERT: owen101lab3;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1459772983] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1459772994] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (5 min average)
[1459773005] HOST ALERT: halleground1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459773011] SERVICE ALERT: owen101lab3;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1459773013] HOST ALERT: owen101lab3;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1459773039] HOST ALERT: halleground1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.54 ms
[1459773044] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1459773074] HOST ALERT: owen101lab3;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1459773089] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 39% - OK!
[1459773129] SERVICE ALERT: halleground2;Device Status;OK;HARD;3;Ready
[1459773139] SERVICE ALERT: halleground2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1459773164] HOST ALERT: owen101lab3;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1459773219] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 62% - OK!
[1459773253] HOST ALERT: owen101lab3;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1459773343] HOST ALERT: owen101lab3;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1459773355] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1459773385] HOST ALERT: owen101lab3;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1459773415] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1459773445] HOST ALERT: owen101lab3;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1459773475] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1459773505] HOST ALERT: owen101lab3;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1459773569] SERVICE ALERT: halleground1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1459773594] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 91% (5 min average)
[1459773629] SERVICE ALERT: halleground1;Device Status;OK;HARD;3;Ready
[1459773714] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 65% (5 min average)
[1459774365] SERVICE ALERT: alexanderpk;Black Cartridge HP CC364X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.142.38. Is the printer off?  Is the network down?
[1459774395] HOST ALERT: alexanderpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459774420] SERVICE ALERT: alexanderpk;Black Cartridge HP CC364X;OK;SOFT;2;Black Cartridge HP CC364X is at 34% - OK!
[1459774424] HOST ALERT: alexanderpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.49 ms
[1459774802] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (25.2% change >= 20.0% threshold)
[1459774802] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1459774802] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1459775830] SERVICE ALERT: ADRODC-WEMU;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459775873] SERVICE ALERT: ADRODC-WEMU;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459775950] SERVICE ALERT: ADRODC-WEMU;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1459775973] SERVICE ALERT: ADRODC-WEMU;Service: RPC;OK;SOFT;2;RpcSs: Started
[1459776050] SERVICE ALERT: ADRODC-WEMU;Service: NetLogin;OK;SOFT;3;Netlogon: Started
[1459776114] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1459776123] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459776148] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1459776153] HOST ALERT: hallefirst2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459776163] HOST ALERT: hallefirst1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459776183] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459776199] HOST ALERT: hallefirst2;UP;SOFT;2;PING OK - Packet loss = 70%, RTA = 0.65 ms
[1459776208] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1459776214] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1459776234] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1459776238] HOST ALERT: hallefirst1;UP;SOFT;2;PING WARNING - Packet loss = 93%, RTA = 0.54 ms
[1459776243] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459776262] SERVICE ALERT: hallefirst1;Device Status;OK;SOFT;3;Processing job
[1459776268] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 55% - OK!
[1459776315] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459776354] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1459776375] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459776435] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459776455] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459776465] HOST ALERT: hallefirst2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459776515] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459776545] HOST ALERT: hallefirst2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1459776575] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459776605] HOST ALERT: hallefirst2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1459776672] HOST ALERT: hallefirst2;UP;SOFT;4;PING OK - Packet loss = 37%, RTA = 0.63 ms
[1459776672] HOST FLAPPING ALERT: hallefirst2;STARTED; Host appears to have started flapping (20.4% change > 20.0% threshold)
[1459777030] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1459777169] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Ready
[1459777437] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 36% - OK!
[1459778369] HOST FLAPPING ALERT: hallefirst2;STOPPED; Host appears to have stopped flapping (4.3% change < 5.0% threshold)
[1459778689] SERVICE ALERT: hallesouthcommons;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1459778913] SERVICE ALERT: EMU-SCCM01;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459778942] SERVICE ALERT: EMU-SCCM01;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459778943] HOST ALERT: EMU-SCCM01;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459779029] SERVICE ALERT: EMU-SCCM01;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459779033] SERVICE ALERT: EMU-SCCM01;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1459779033] HOST ALERT: EMU-SCCM01;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1459779062] SERVICE ALERT: EMU-SCCM01;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1459779079] SERVICE ALERT: EMU-SCCM01;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459779089] SERVICE ALERT: EMU-SCCM01;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459779092] HOST ALERT: EMU-SCCM01;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1459779125] SERVICE ALERT: EMU-SCCM01;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459779145] SERVICE ALERT: EMU-SCCM01;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459779149] SERVICE ALERT: EMU-SCCM01;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1459779153] SERVICE ALERT: EMU-SCCM01;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1459779155] HOST ALERT: EMU-SCCM01;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1459779169] SERVICE ALERT: EMU-SCCM01;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459779182] SERVICE ALERT: EMU-SCCM01;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1459779199] SERVICE ALERT: EMU-SCCM01;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1459779199] HOST ALERT: EMU-SCCM01;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1459779209] SERVICE ALERT: EMU-SCCM01;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1459779239] HOST ALERT: EMU-SCCM01;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1459779245] SERVICE ALERT: EMU-SCCM01;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1459779257] SERVICE ALERT: EMU-SCCM01;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459779265] SERVICE ALERT: EMU-SCCM01;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1459779269] SERVICE ALERT: EMU-SCCM01;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1459779275] HOST ALERT: EMU-SCCM01;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1459779282] SERVICE ALERT: EMU-SCCM01;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459779289] SERVICE ALERT: EMU-SCCM01;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459779289] SERVICE ALERT: EMU-SCCM01;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1459779312] HOST ALERT: FTK;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.240)
[1459779312] HOST ALERT: EMU-SCCM01;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1459779319] SERVICE ALERT: EMU-SCCM01;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1459779329] SERVICE ALERT: EMU-SCCM01;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1459779349] HOST ALERT: EMU-SCCM01;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1459779354] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (5 min average)
[1459779359] SERVICE ALERT: EMU-SCCM01;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459779364] SERVICE ALERT: FTK;CPU Load;CRITICAL;SOFT;1;connect to address 164.76.102.240 and port 12489: No route to host
[1459779364] SERVICE ALERT: FTK;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.102.240 and port 12489: No route to host
[1459779365] SERVICE ALERT: EMU-SCCM01;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1459779367] HOST ALERT: FTK;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.240)
[1459779377] SERVICE ALERT: EMU-SCCM01;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1459779385] SERVICE ALERT: EMU-SCCM01;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1459779389] HOST ALERT: EMU-SCCM01;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1459779392] SERVICE ALERT: FTK;Uptime;CRITICAL;SOFT;1;connect to address 164.76.102.240 and port 12489: No route to host
[1459779395] HOST ALERT: FTK;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.102.240)
[1459779402] SERVICE ALERT: EMU-SCCM01;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459779402] SERVICE ALERT: EMU-SCCM01;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1459779409] SERVICE ALERT: EMU-SCCM01;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1459779409] SERVICE ALERT: EMU-SCCM01;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1459779427] SERVICE ALERT: FTK;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.102.240 and port 12489: No route to host
[1459779430] HOST ALERT: FTK;DOWN;SOFT;4;CRITICAL - Host Unreachable (164.76.102.240)
[1459779452] SERVICE ALERT: FTK;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.102.240 and port 12489: No route to host
[1459779455] HOST ALERT: FTK;DOWN;SOFT;5;CRITICAL - Host Unreachable (164.76.102.240)
[1459779463] SERVICE ALERT: FTK;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.102.240 and port 12489: No route to host
[1459779469] SERVICE ALERT: EMU-SCCM01;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459779474] SERVICE ALERT: FTK;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.102.240 and port 12489: No route to host
[1459779477] SERVICE ALERT: FTK;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.102.240 and port 12489: No route to host
[1459779477] HOST ALERT: FTK;DOWN;SOFT;6;CRITICAL - Host Unreachable (164.76.102.240)
[1459779479] SERVICE ALERT: EMU-SCCM01;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1459779484] SERVICE ALERT: FTK;CPU Load;CRITICAL;SOFT;2;connect to address 164.76.102.240 and port 12489: No route to host
[1459779484] SERVICE ALERT: FTK;Memory Usage;CRITICAL;SOFT;2;connect to address 164.76.102.240 and port 12489: No route to host
[1459779497] SERVICE ALERT: EMU-SCCM01;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1459779507] SERVICE ALERT: FTK;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.102.240 and port 12489: No route to host
[1459779510] SERVICE ALERT: FTK;Uptime;CRITICAL;SOFT;2;connect to address 164.76.102.240 and port 12489: No route to host
[1459779510] HOST ALERT: FTK;DOWN;SOFT;7;CRITICAL - Host Unreachable (164.76.102.240)
[1459779522] SERVICE ALERT: EMU-SCCM01;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1459779522] SERVICE ALERT: EMU-SCCM01;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1459779529] SERVICE ALERT: EMU-SCCM01;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1459779547] SERVICE ALERT: FTK;Service: WMI;CRITICAL;SOFT;2;connect to address 164.76.102.240 and port 12489: No route to host
[1459779550] HOST ALERT: FTK;DOWN;SOFT;8;CRITICAL - Host Unreachable (164.76.102.240)
[1459779572] SERVICE ALERT: FTK;C:\ Drive Space;CRITICAL;SOFT;2;connect to address 164.76.102.240 and port 12489: No route to host
[1459779575] HOST ALERT: FTK;DOWN;SOFT;9;CRITICAL - Host Unreachable (164.76.102.240)
[1459779583] SERVICE ALERT: FTK;NSClient++ Version;CRITICAL;SOFT;2;connect to address 164.76.102.240 and port 12489: No route to host
[1459779589] SERVICE ALERT: EMU-SCCM01;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1459779594] SERVICE ALERT: FTK;Service: Windows Time;CRITICAL;SOFT;2;connect to address 164.76.102.240 and port 12489: No route to host
[1459779597] HOST ALERT: FTK;DOWN;HARD;10;CRITICAL - Host Unreachable (164.76.102.240)
[1459779599] SERVICE ALERT: EMU-SCCM01;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1459779600] SERVICE ALERT: FTK;Service: Firewall;CRITICAL;SOFT;2;connect to address 164.76.102.240 and port 12489: No route to host
[1459779604] SERVICE ALERT: FTK;CPU Load;CRITICAL;HARD;3;connect to address 164.76.102.240 and port 12489: No route to host
[1459779604] SERVICE ALERT: FTK;Memory Usage;CRITICAL;HARD;3;connect to address 164.76.102.240 and port 12489: No route to host
[1459779627] SERVICE ALERT: FTK;Service: DNS Client;CRITICAL;SOFT;2;connect to address 164.76.102.240 and port 12489: No route to host
[1459779630] SERVICE ALERT: FTK;Uptime;CRITICAL;HARD;3;connect to address 164.76.102.240 and port 12489: No route to host
[1459779642] SERVICE ALERT: EMU-SCCM01;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1459779644] SERVICE ALERT: FTK;Service: NetLogin;CRITICAL;SOFT;1;connect to address 164.76.102.240 and port 12489: No route to host
[1459779667] SERVICE ALERT: FTK;Service: WMI;CRITICAL;HARD;3;connect to address 164.76.102.240 and port 12489: No route to host
[1459779681] SERVICE ALERT: FTK;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.102.240 and port 12489: No route to host
[1459779692] SERVICE ALERT: FTK;C:\ Drive Space;CRITICAL;HARD;3;connect to address 164.76.102.240 and port 12489: No route to host
[1459779703] SERVICE ALERT: FTK;NSClient++ Version;CRITICAL;HARD;3;connect to address 164.76.102.240 and port 12489: No route to host
[1459779709] SERVICE ALERT: EMU-SCCM01;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1459779714] SERVICE ALERT: FTK;Service: Windows Time;CRITICAL;HARD;3;connect to address 164.76.102.240 and port 12489: No route to host
[1459779720] SERVICE ALERT: FTK;Service: Firewall;CRITICAL;HARD;3;connect to address 164.76.102.240 and port 12489: No route to host
[1459779728] SERVICE ALERT: FTK;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.102.240 and port 12489: No route to host
[1459779728] SERVICE ALERT: FTK;Service: User Profile;CRITICAL;SOFT;1;connect to address 164.76.102.240 and port 12489: No route to host
[1459779747] SERVICE ALERT: FTK;Service: DNS Client;CRITICAL;HARD;3;connect to address 164.76.102.240 and port 12489: No route to host
[1459779764] SERVICE ALERT: FTK;Service: NetLogin;CRITICAL;SOFT;2;connect to address 164.76.102.240 and port 12489: No route to host
[1459779801] SERVICE ALERT: FTK;Service: RPC;CRITICAL;SOFT;2;connect to address 164.76.102.240 and port 12489: No route to host
[1459779848] SERVICE ALERT: FTK;Service: User Profile;CRITICAL;SOFT;2;connect to address 164.76.102.240 and port 12489: No route to host
[1459779848] SERVICE ALERT: FTK;Service: Group Policy;CRITICAL;SOFT;2;connect to address 164.76.102.240 and port 12489: No route to host
[1459779884] SERVICE ALERT: FTK;Service: NetLogin;CRITICAL;HARD;3;connect to address 164.76.102.240 and port 12489: No route to host
[1459779921] SERVICE ALERT: FTK;Service: RPC;CRITICAL;HARD;3;connect to address 164.76.102.240 and port 12489: No route to host
[1459779954] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1459779968] SERVICE ALERT: FTK;Service: Group Policy;CRITICAL;HARD;3;connect to address 164.76.102.240 and port 12489: No route to host
[1459779968] SERVICE ALERT: FTK;Service: User Profile;CRITICAL;HARD;3;connect to address 164.76.102.240 and port 12489: No route to host
[1459780074] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1459780194] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1459780203] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (30.4% change >= 20.0% threshold)
[1459780203] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1459780203] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1459780207] HOST ALERT: EMU-SCCM01;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1459780222] SERVICE ALERT: EMU-SCCM01;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1459780229] SERVICE ALERT: FTK;Uptime;OK;HARD;3;System Uptime - 18 day(s) 11 hour(s) 23 minute(s)
[1459780233] HOST ALERT: FTK;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.22 ms
[1459780264] SERVICE ALERT: FTK;Service: WMI;OK;HARD;3;Winmgmt: Started
[1459780289] SERVICE ALERT: FTK;C:\ Drive Space;OK;HARD;3;c:\ - total: 408.28 Gb - used: 138.51 Gb (34%) - free 269.77 Gb (66%)
[1459780289] SERVICE ALERT: EMU-SCCM01;Service: WMI;OK;HARD;3;Winmgmt: Started
[1459780296] SERVICE ALERT: FTK;CPU Load;OK;HARD;3;CPU Load 0% (5 min average)
[1459780300] SERVICE ALERT: FTK;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1459780311] SERVICE ALERT: FTK;Service: Windows Time;OK;HARD;3;W32Time: Started
[1459780317] SERVICE ALERT: FTK;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1459780333] SERVICE ALERT: EMU-SCCM01;CPU Load;OK;HARD;3;CPU Load 0% (5 min average)
[1459780343] SERVICE ALERT: FTK;Memory Usage;OK;HARD;3;Memory usage: total:65524.64 MB - used: 19569.46 MB (30%) - free: 45955.19 MB (70%)
[1459780344] SERVICE ALERT: FTK;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1459780362] SERVICE ALERT: EMU-SCCM01;Memory Usage;OK;HARD;3;Memory usage: total:21759.74 MB - used: 8376.72 MB (38%) - free: 13383.02 MB (62%)
[1459780449] SERVICE ALERT: EMU-SCCM01;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1459780481] SERVICE ALERT: FTK;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1459780499] SERVICE ALERT: EMU-SCCM01;Service: Windows Time;OK;HARD;3;W32Time: Started
[1459780509] SERVICE ALERT: EMU-SCCM01;VMWare Tools;OK;HARD;3;VMTools: Started
[1459780518] SERVICE ALERT: FTK;Service: RPC;OK;HARD;3;RpcSs: Started
[1459780545] SERVICE ALERT: EMU-SCCM01;Uptime;OK;HARD;3;System Uptime - 26 day(s) 4 hour(s) 43 minute(s)
[1459780565] SERVICE ALERT: EMU-SCCM01;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1459780565] SERVICE ALERT: FTK;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1459780567] SERVICE ALERT: FTK;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1459780589] SERVICE ALERT: EMU-SCCM01;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1459780677] SERVICE ALERT: EMU-SCCM01;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1459780702] SERVICE ALERT: EMU-SCCM01;C:\ Drive Space;OK;HARD;3;c:\ - total: 79.66 Gb - used: 31.45 Gb (39%) - free 48.21 Gb (61%)
[1459780709] SERVICE ALERT: EMU-SCCM01;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1459780779] SERVICE ALERT: EMU-SCCM01;Service: RPC;OK;HARD;3;RpcSs: Started
[1459781708] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1459781769] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1459781829] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1459782428] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1459782428] SERVICE FLAPPING ALERT: alexanderpk;Paper Tray Status;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1459782569] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459782630] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459782689] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459783288] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Processing job from tray 4
[1459783288] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1459784304] SERVICE ALERT: INTLTESTAPP;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459784324] SERVICE ALERT: INTLTESTAPP;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459784325] SERVICE ALERT: SANDBOX-DB;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459784334] HOST ALERT: INTLTESTAPP;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459784398] HOST ALERT: INTLTESTAPP;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.36 ms
[1459784404] SERVICE ALERT: INTLTESTAPP;Uptime;OK;SOFT;2;System Uptime - 2 day(s) 20 hour(s) 4 minute(s)
[1459784424] SERVICE ALERT: INTLTESTAPP;Service: RPC;OK;SOFT;2;RpcSs: Started
[1459784425] SERVICE ALERT: SANDBOX-DB;VMWare Tools;OK;SOFT;2;VMTools: Started
[1459784790] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1459784849] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1459784910] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1459785027] HOST ALERT: SNAP;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1459785594] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 83% (5 min average)
[1459786194] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459786314] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1459786434] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1459787309] SERVICE ALERT: crossroadspk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1459788509] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1459788570] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1459788629] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1459788772] SERVICE ALERT: C4SAS-SE;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459788891] SERVICE ALERT: C4SAS-SE;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1459790034] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 84% (5 min average)
[1459790208] SERVICE ALERT: eateriespk;Black Cartridge HP CC364X;WARNING;SOFT;1;Black Cartridge HP CC364X is at 2% - WARNING!
[1459790269] SERVICE ALERT: eateriespk;Black Cartridge HP CC364X;WARNING;SOFT;2;Black Cartridge HP CC364X is at 2% - WARNING!
[1459790328] SERVICE ALERT: eateriespk;Black Cartridge HP CC364X;WARNING;HARD;3;Black Cartridge HP CC364X is at 2% - WARNING!
[1459790634] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1459790754] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1459790874] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1459792540] SERVICE ALERT: studentcenterpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.150.252. Is the printer off?  Is the network down?
[1459792594] SERVICE ALERT: studentcenterpk;Device Status;OK;SOFT;2;Tray 4 ANY TYPE LETTER To change size or 74 79 70 65 20 70 72 65 73 73 20 1E  Tray 4 ANY TYPE LETTER To accept settings 70 72 65 73 73 20 18
[1459792934] SERVICE ALERT: INTLTESTAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459793054] SERVICE ALERT: INTLTESTAPP;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1459793283] SERVICE ALERT: markjeffersonpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1459793313] HOST ALERT: markjeffersonpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459793343] SERVICE ALERT: markjeffersonpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1459793373] HOST ALERT: markjeffersonpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1459793403] SERVICE ALERT: markjeffersonpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1459793433] HOST ALERT: markjeffersonpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1459793517] SERVICE ALERT: MBAM;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459793523] HOST ALERT: markjeffersonpk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1459793584] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1459793613] HOST ALERT: markjeffersonpk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1459793637] SERVICE ALERT: MBAM;CPU Load;OK;SOFT;2;CPU Load 1% (5 min average)
[1459793638] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;SOFT;1;WARNING: Replication with ADDC5 failed!
[1459793644] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1459793674] HOST ALERT: markjeffersonpk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1459793704] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1459793734] HOST ALERT: markjeffersonpk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1459793759] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;SOFT;2;WARNING: Replication with ADDC5 failed!
[1459793804] SERVICE ALERT: markjeffersonpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1459793824] HOST ALERT: markjeffersonpk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1459793828] SERVICE FLAPPING ALERT: alexanderpk;Paper Tray Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1459793864] SERVICE ALERT: markjeffersonpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1459793874] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 94% (5 min average)
[1459793879] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;HARD;3;WARNING: Replication with ADDC5 failed!
[1459793894] HOST ALERT: markjeffersonpk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1459793924] SERVICE ALERT: markjeffersonpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1459793954] HOST ALERT: markjeffersonpk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1459794328] SERVICE ALERT: ADCA1;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1459794448] SERVICE ALERT: ADCA1;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1459794474] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1459794631] HOST ALERT: markjeffersonpk;UP;HARD;10;PING WARNING - Packet loss = 82%, RTA = 0.69 ms
[1459794688] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1459794888] SERVICE ALERT: SP2010DB;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459794898] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 52% - OK!
[1459795007] SERVICE ALERT: SP2010DB;CPU Load;OK;SOFT;2;CPU Load 6% (5 min average)
[1459795118] SERVICE ALERT: markjeffersonpk;Device Status;OK;HARD;3;Ready
[1459795197] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1459795648] SERVICE ALERT: ADCA1;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459795768] SERVICE ALERT: ADCA1;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1459796177] HOST ALERT: markjeffersonpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459796241] HOST ALERT: markjeffersonpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.67 ms
[1459797303] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1459797363] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1459797423] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1459797757] HOST ALERT: hallefirst1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459797821] HOST ALERT: hallefirst1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.55 ms
[1459797866] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1459797874] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1459797926] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1459797934] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1459797956] HOST ALERT: hallefirst1;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1459797980] SERVICE ALERT: hallefirst1;Device Status;OK;SOFT;3;Processing job
[1459797984] HOST ALERT: hallefirst1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.58 ms
[1459797984] HOST FLAPPING ALERT: hallefirst1;STARTED; Host appears to have started flapping (22.5% change > 20.0% threshold)
[1459797988] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;SOFT;3;Black Print Cartridge HP C8543X is at 43% - OK!
[1459798018] SERVICE ALERT: hallefirst1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1459798034] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459798094] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459798154] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459798292] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459798326] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1459798346] SERVICE ALERT: hallefirst2;Device Status;OK;SOFT;2;Processing job from tray 4
[1459798387] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1459798446] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1459798586] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1459798594] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1459798623] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1459798646] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1459798654] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1459798683] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1459798706] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1459798714] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1459798743] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1459798747] HOST FLAPPING ALERT: hallefirst1;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1459798749] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1459799300] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;CONFIGURE DEVICE  PRINTING  PRINT QUALITY  SYSTEM SETUP
[1459799308] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 42% - OK!
[1459799337] SERVICE ALERT: hallefirst1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1459799337] SERVICE FLAPPING ALERT: hallefirst1;Paper Tray Status;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1459800002] SERVICE FLAPPING ALERT: hallefirst1;Paper Tray Status;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1459800002] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1459800002] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1459800736] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 1 Trays Empty
[1459800797] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 1 Trays Empty
[1459800856] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 1 Trays Empty
[1459801074] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 94% (5 min average)
[1459801074] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1459801674] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1459802056] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1459802546] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459802605] SERVICE ALERT: hallefirst2;Device Status;OK;SOFT;2;Clearing paper path
[1459803149] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459803208] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459803268] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459803868] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;Ready
[1459804683] SERVICE ALERT: ADRODC-LIV;AD Replication;OK;HARD;3;OK: replication is up and running
[1459805288] SERVICE ALERT: mckennypk;Black Cartridge HP CC364X;WARNING;SOFT;1;Black Cartridge HP CC364X is at 2% - WARNING!
[1459805348] SERVICE ALERT: mckennypk;Black Cartridge HP CC364X;WARNING;SOFT;2;Black Cartridge HP CC364X is at 2% - WARNING!
[1459805408] SERVICE ALERT: mckennypk;Black Cartridge HP CC364X;WARNING;HARD;3;Black Cartridge HP CC364X is at 2% - WARNING!
[1459806245] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1459806845] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1459810737] SERVICE FLAPPING ALERT: hallefirst1;Paper Tray Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1459811604] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459811664] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459811724] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459813074] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1459814299] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459814359] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459814420] SERVICE ALERT: hallefirst1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459814429] SERVICE ALERT: crossroadspk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1459814662] SERVICE ALERT: hallefirstcolor;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1459814722] SERVICE ALERT: hallefirstcolor;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1459814763] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1459814782] SERVICE ALERT: hallefirstcolor;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1459814803] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1459814812] HOST ALERT: hallefirstcolor;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459814823] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1459814853] HOST ALERT: hallefirstcolor;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1459814857] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Yellow Cartridge HP CE262A'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1459814862] HOST ALERT: hallefirstcolor;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1459814877] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;OK;SOFT;3;Black Cartridge HP CE260X is at 31% - OK!
[1459814918] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;OK;SOFT;3;Yellow Cartridge HP CE262A is at 12% - OK!
[1459815019] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;Ready
[1459815376] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1459816523] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Warming up
[1459817124] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459817183] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459817242] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459817419] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459817479] SERVICE ALERT: hallefirst1;Device Status;OK;SOFT;2;Warming up
[1459818043] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459818103] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459818163] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459818764] SERVICE ALERT: hallesecondsouth;Device Status;OK;HARD;3;50 72 65 73 73 20 17 20 74 6F 20 73 65 65 20 73  74 65 70  by step information
[1459818787] SERVICE ALERT: SANDBOX-DB;CPU Load;WARNING;SOFT;1;CPU Load 91% (5 min average)
[1459818907] SERVICE ALERT: SANDBOX-DB;CPU Load;WARNING;SOFT;2;CPU Load 92% (5 min average)
[1459819027] SERVICE ALERT: SANDBOX-DB;CPU Load;OK;SOFT;3;CPU Load 82% (5 min average)
[1459819042] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Processing job from tray 4
[1459819042] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1459819642] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459819702] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459819763] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459819963] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459820023] SERVICE ALERT: hallesecondsouth;Device Status;OK;SOFT;2;Checking printer
[1459820362] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Clearing paper path
[1459821223] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459821283] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459821344] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459823951] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459823968] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459823981] HOST ALERT: hallefirst2;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1459824011] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459824028] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459824042] HOST ALERT: hallefirst2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1459824072] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459824088] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459824102] HOST ALERT: hallefirst2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1459824192] HOST ALERT: hallefirst2;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1459824239] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459824269] HOST ALERT: hallefirst2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1459824299] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459824329] HOST ALERT: hallefirst2;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1459824359] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459824389] HOST ALERT: hallefirst2;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1459824479] HOST ALERT: hallefirst2;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1459824569] HOST ALERT: hallefirst2;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1459824659] HOST ALERT: hallefirst2;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1459826349] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 91% (5 min average)
[1459826469] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 76% (5 min average)
[1459826926] SERVICE ALERT: eateriespk;Black Cartridge HP CC364X;CRITICAL;HARD;3;Black Cartridge HP CC364X is at 1% - CRITICAL!