[1462853827] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462853945] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1462854064] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462854664] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462855265] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462855864] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462856941] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462857061] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 91% (15 min average)
[1462858264] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (15 min average)
[1462858861] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1462858864] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462858981] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 89% (15 min average)
[1462858984] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1462859104] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 87% (15 min average)
[1462859581] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462859700] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1462860300] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1462860304] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462860420] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 74% (15 min average)
[1462860424] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462860544] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1462862220] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462862340] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1462862460] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1462863060] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462864260] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462864860] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462865460] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462866661] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462867260] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 78% (15 min average)
[1462867860] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1462867980] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1462868580] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1462868700] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1462868820] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1462868944] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 88% (15 min average)
[1462869544] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462869664] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462869784] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462870620] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1462871220] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1462871340] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1462871460] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 92% (15 min average)
[1462871584] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 59% (15 min average)
[1462872060] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1462872180] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1462872300] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1462873384] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462873504] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1462873624] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462874224] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1462874824] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462874944] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462875064] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462875302] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462875899] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1462876264] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 90% (15 min average)
[1462876499] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1462876619] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1462876864] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462876984] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1462877104] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1462877219] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1462877339] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 90% (15 min average)
[1462877344] SERVICE ALERT: c137.emich.edu;Service: HAProxy Status;WARNING;HARD;2;(No output on stdout) stderr: Could not resolve hostname c137.emich.edu: Name or service not known
[1462877374] HOST ALERT: c137.emich.edu;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462877464] HOST ALERT: c137.emich.edu;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1462877554] HOST ALERT: c137.emich.edu;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1462877644] HOST ALERT: c137.emich.edu;DOWN;SOFT;4;(Host check timed out after 30.00 seconds)
[1462877674] HOST ALERT: c137.emich.edu;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1462877764] HOST ALERT: c137.emich.edu;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1462877854] HOST ALERT: c137.emich.edu;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1462877918] HOST ALERT: c137.emich.edu;UP;SOFT;8;PING OK - Packet loss = 0%, RTA = 0.41 ms
[1462879139] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462879259] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 74% (15 min average)
[1462879958] SERVICE ALERT: halleground1;Device Status;OK;HARD;3;Warming up
[1462879962] HOST ALERT: halleground1;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.57 ms
[1462880142] SERVICE ALERT: halleground1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1462880250] SERVICE ALERT: hallefirst2;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1462880315] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1462880321] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1462880345] HOST ALERT: hallefirst2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462880370] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;SOFT;3;Printer Paper Is OK! All Trays Still Full
[1462880374] HOST ALERT: hallefirst2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.57 ms
[1462880449] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 63% - OK!
[1462880916] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Ready
[1462881274] HOST ALERT: c137.emich.edu;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462881304] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 86% (15 min average)
[1462881358] HOST ALERT: c137.emich.edu;UP;SOFT;2;PING WARNING - Packet loss = 80%, RTA = 0.45 ms
[1462881904] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462882024] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1462882144] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 65% (15 min average)
[1462882204] SERVICE ALERT: mckennypk;Device Status;OK;HARD;3;Warming up
[1462882539] SERVICE ALERT: idm1.emich.edu;Loopback Driver;WARNING;SOFT;1;0
[1462882660] SERVICE ALERT: idm1.emich.edu;Loopback Driver;WARNING;SOFT;2;0
[1462882744] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462882779] SERVICE ALERT: idm1.emich.edu;Loopback Driver;WARNING;HARD;3;0
[1462882864] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462882964] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;WARNING;SOFT;1;0
[1462882984] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462883084] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;WARNING;SOFT;2;0
[1462883199] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1462883204] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;WARNING;HARD;3;0
[1462883258] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1462883318] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1462883584] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 51% (15 min average)
[1462884252] HOST ALERT: mckennypk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462884323] HOST ALERT: mckennypk;UP;SOFT;2;PING OK - Packet loss = 58%, RTA = 0.79 ms
[1462884403] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;OK;HARD;3;0
[1462884518] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 42% - OK!
[1462884578] SERVICE ALERT: idm1.emich.edu;Loopback Driver;OK;HARD;3;0
[1462884604] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462884664] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462884724] SERVICE ALERT: mckennypk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462884784] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1462884904] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 88% (15 min average)
[1462885504] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1462885624] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 52% (15 min average)
[1462885924] SERVICE ALERT: mckennypk;Device Status;OK;HARD;3;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1462885968] SERVICE ALERT: R25-ALPHA2012;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462885998] HOST ALERT: R25-ALPHA2012;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462886020] SERVICE ALERT: R25-ALPHA2012;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462886050] HOST ALERT: R25-ALPHA2012;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1462886070] SERVICE ALERT: R25-ALPHA2012;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462886074] SERVICE ALERT: R25-ALPHA2012;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462886075] SERVICE ALERT: R25-ALPHA2012;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462886088] SERVICE ALERT: R25-ALPHA2012;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462886096] SERVICE ALERT: R25-ALPHA2012;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462886097] SERVICE ALERT: R25-ALPHA2012;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462886100] HOST ALERT: R25-ALPHA2012;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1462886140] SERVICE ALERT: R25-ALPHA2012;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462886170] HOST ALERT: R25-ALPHA2012;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1462886181] SERVICE ALERT: R25-DELTA2012;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462886184] SERVICE ALERT: R25-DELTA2012;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462886190] SERVICE ALERT: R25-ALPHA2012;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462886194] SERVICE ALERT: R25-ALPHA2012;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462886195] SERVICE ALERT: R25-ALPHA2012;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462886208] SERVICE ALERT: R25-ALPHA2012;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462886211] HOST ALERT: R25-DELTA2012;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462886216] SERVICE ALERT: R25-ALPHA2012;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462886217] SERVICE ALERT: R25-ALPHA2012;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462886220] HOST ALERT: R25-ALPHA2012;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1462886221] SERVICE ALERT: R25-DELTA2012;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462886222] SERVICE ALERT: R25-ALPHA2012;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462886250] SERVICE ALERT: R25-DELTA2012;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462886250] SERVICE ALERT: R25-ALPHA2012;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462886251] HOST ALERT: R25-DELTA2012;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1462886252] HOST ALERT: R25-ALPHA2012;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1462886260] SERVICE ALERT: R25-ALPHA2012;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462886262] SERVICE ALERT: R25-DELTA2012;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462886276] SERVICE ALERT: R25-DELTA2012;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462886276] SERVICE ALERT: R25-ALPHA2012;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462886277] SERVICE ALERT: R25-ALPHA2012;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462886290] HOST ALERT: R25-ALPHA2012;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1462886292] HOST ALERT: R25-DELTA2012;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1462886301] SERVICE ALERT: R25-DELTA2012;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462886304] SERVICE ALERT: R25-DELTA2012;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462886307] SERVICE ALERT: R25-DELTA2012;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462886310] SERVICE ALERT: R25-ALPHA2012;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462886314] SERVICE ALERT: R25-ALPHA2012;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462886315] SERVICE ALERT: R25-ALPHA2012;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462886318] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1462886331] HOST ALERT: R25-DELTA2012;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1462886336] SERVICE ALERT: R25-ALPHA2012;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462886337] SERVICE ALERT: R25-ALPHA2012;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462886340] HOST ALERT: R25-ALPHA2012;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1462886341] SERVICE ALERT: R25-DELTA2012;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462886342] SERVICE ALERT: R25-ALPHA2012;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462886350] SERVICE ALERT: R25-ZETA2012;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462886370] SERVICE ALERT: R25-DELTA2012;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462886370] SERVICE ALERT: R25-ALPHA2012;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462886371] HOST ALERT: R25-DELTA2012;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1462886372] HOST ALERT: R25-ALPHA2012;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1462886378] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1462886380] HOST ALERT: R25-ZETA2012;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462886382] SERVICE ALERT: R25-DELTA2012;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462886394] SERVICE ALERT: R25-ALPHA2012;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462886396] SERVICE ALERT: R25-DELTA2012;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462886396] SERVICE ALERT: R25-ALPHA2012;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462886397] SERVICE ALERT: R25-ALPHA2012;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462886412] HOST ALERT: R25-DELTA2012;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1462886421] SERVICE ALERT: R25-DELTA2012;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462886424] SERVICE ALERT: R25-DELTA2012;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462886424] HOST ALERT: R25-ALPHA2012;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1462886427] SERVICE ALERT: R25-DELTA2012;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462886438] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1462886451] HOST ALERT: R25-DELTA2012;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1462886458] SERVICE ALERT: R25-ALPHA2012;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462886461] SERVICE ALERT: R25-DELTA2012;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462886462] SERVICE ALERT: R25-ALPHA2012;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462886470] SERVICE ALERT: R25-ZETA2012;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462886470] HOST ALERT: R25-ZETA2012;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1462886490] SERVICE ALERT: R25-DELTA2012;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462886490] SERVICE ALERT: R25-ALPHA2012;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462886491] HOST ALERT: R25-DELTA2012;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1462886491] HOST DOWNTIME ALERT: R25-ALPHA2012;STARTED; Host has entered a period of scheduled downtime
[1462886494] SERVICE ALERT: R25-DELTA2012;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462886500] SERVICE ALERT: R25-DELTA2012;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462886500] SERVICE ALERT: R25-ALPHA2012;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462886502] SERVICE ALERT: R25-DELTA2012;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462886502] HOST DOWNTIME ALERT: R25-DELTA2012;STARTED; Host has entered a period of scheduled downtime
[1462886502] SERVICE ALERT: R25-ZETA2012;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462886514] SERVICE ALERT: R25-ALPHA2012;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462886516] SERVICE ALERT: R25-DELTA2012;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462886516] SERVICE ALERT: R25-ALPHA2012;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462886517] SERVICE ALERT: R25-ALPHA2012;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462886522] HOST DOWNTIME ALERT: R25-ZETA2012;STARTED; Host has entered a period of scheduled downtime
[1462886524] HOST ALERT: R25-DELTA2012;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1462886532] HOST ALERT: R25-ZETA2012;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1462886547] SERVICE ALERT: R25-DELTA2012;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462886555] SERVICE ALERT: R25-ZETA2012;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462886558] SERVICE ALERT: R25-DELTA2012;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462886569] SERVICE ALERT: R25-ZETA2012;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462886576] SERVICE ALERT: R25-DELTA2012;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462886577] HOST ALERT: R25-DELTA2012;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1462886578] SERVICE ALERT: R25-ALPHA2012;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462886584] SERVICE ALERT: R25-ZETA2012;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462886585] HOST ALERT: R25-ZETA2012;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1462886590] SERVICE ALERT: R25-ZETA2012;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462886607] SERVICE ALERT: R25-DELTA2012;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462886614] SERVICE ALERT: R25-DELTA2012;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462886616] SERVICE ALERT: R25-ZETA2012;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462886620] SERVICE ALERT: R25-DELTA2012;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462886620] SERVICE ALERT: R25-ALPHA2012;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462886620] HOST ALERT: R25-ZETA2012;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1462886622] SERVICE ALERT: R25-ZETA2012;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462886634] SERVICE ALERT: R25-ALPHA2012;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462886652] HOST ALERT: R25-ZETA2012;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1462886675] SERVICE ALERT: R25-ZETA2012;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462886678] SERVICE ALERT: R25-DELTA2012;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462886689] SERVICE ALERT: R25-ZETA2012;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462886695] SERVICE ALERT: R25-DELTA2012;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462886696] SERVICE ALERT: R25-DELTA2012;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462886698] SERVICE ALERT: R25-ALPHA2012;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462886704] SERVICE ALERT: R25-ZETA2012;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462886705] HOST ALERT: R25-ZETA2012;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1462886720] SERVICE ALERT: R25-DELTA2012;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462886727] SERVICE ALERT: R25-ZETA2012;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462886727] SERVICE ALERT: R25-DELTA2012;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462886734] SERVICE ALERT: R25-DELTA2012;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462886736] SERVICE ALERT: R25-ZETA2012;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462886740] SERVICE ALERT: R25-DELTA2012;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462886740] SERVICE ALERT: R25-ALPHA2012;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462886742] SERVICE ALERT: R25-ZETA2012;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462886747] SERVICE ALERT: R25-ZETA2012;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462886751] SERVICE ALERT: R25-ZETA2012;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462886757] HOST ALERT: R25-ZETA2012;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1462886795] SERVICE ALERT: R25-ZETA2012;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462886798] SERVICE ALERT: R25-DELTA2012;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462886809] SERVICE ALERT: R25-ZETA2012;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462886815] SERVICE ALERT: R25-DELTA2012;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462886816] SERVICE ALERT: R25-DELTA2012;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462886819] SERVICE ALERT: R25-ZETA2012;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462886824] SERVICE ALERT: R25-ZETA2012;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462886825] HOST ALERT: R25-ZETA2012;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1462886840] SERVICE ALERT: R25-DELTA2012;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462886846] SERVICE ALERT: R25-ZETA2012;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462886847] SERVICE ALERT: R25-DELTA2012;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462886856] SERVICE ALERT: R25-ZETA2012;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462886867] SERVICE ALERT: R25-ZETA2012;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462886871] SERVICE ALERT: R25-ZETA2012;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462886876] HOST ALERT: R25-ZETA2012;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1462886878] SERVICE ALERT: R25-ZETA2012;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462886878] SERVICE ALERT: R25-ZETA2012;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462886907] SERVICE ALERT: R25-ZETA2012;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462886908] SERVICE ALERT: R25-ZETA2012;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462886935] SERVICE ALERT: R25-DELTA2012;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462886939] SERVICE ALERT: R25-ZETA2012;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462886960] SERVICE ALERT: R25-DELTA2012;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462886966] SERVICE ALERT: R25-ZETA2012;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462886987] SERVICE ALERT: R25-ZETA2012;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462886991] SERVICE ALERT: R25-ZETA2012;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462886998] SERVICE ALERT: R25-ZETA2012;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462886998] SERVICE ALERT: R25-ZETA2012;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462887010] SERVICE ALERT: ADDC5;GPO Replication;WARNING;SOFT;1;Warning! 1 policies are not replicating properly!
[1462887027] SERVICE ALERT: R25-ZETA2012;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462887028] SERVICE ALERT: R25-ZETA2012;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462887059] SERVICE ALERT: R25-ZETA2012;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462887118] SERVICE ALERT: R25-ZETA2012;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462887118] SERVICE ALERT: R25-ZETA2012;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462887147] SERVICE ALERT: R25-ZETA2012;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462887148] SERVICE ALERT: R25-ZETA2012;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462887309] SERVICE ALERT: ADDC5;GPO Replication;WARNING;SOFT;2;Warning! 1 policies are not replicating properly!
[1462887424] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1462887544] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 65% (15 min average)
[1462887611] SERVICE ALERT: ADDC5;GPO Replication;WARNING;HARD;3;Warning! 1 policies are not replicating properly!
[1462887659] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1462888804] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (22.7% change >= 20.0% threshold)
[1462888804] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1462888804] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1462888804] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1462888804] HOST DOWNTIME ALERT: ADAUTHDC3;STARTED; Host has entered a period of scheduled downtime
[1462888804] HOST DOWNTIME ALERT: R25-ALPHA2012;STARTED; Host has entered a period of scheduled downtime
[1462888804] SERVICE DOWNTIME ALERT: NETIDADPROXY;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1462888804] HOST DOWNTIME ALERT: R25-DELTA2012;STARTED; Host has entered a period of scheduled downtime
[1462888804] HOST DOWNTIME ALERT: R25-ZETA2012;STARTED; Host has entered a period of scheduled downtime
[1462888804] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1462888804] HOST DOWNTIME ALERT: DIRECTACCESS;STARTED; Host has entered a period of scheduled downtime
[1462888804] HOST DOWNTIME ALERT: FITNESSGRAM;STARTED; Host has entered a period of scheduled downtime
[1462888839] SERVICE ALERT: R25-ZETA2012;Service: Windows Time;OK;HARD;3;W32Time: Started
[1462888841] SERVICE ALERT: R25-DELTA2012;CPU Load;OK;HARD;3;CPU Load 0% (15 min average)
[1462888842] SERVICE ALERT: R25-ALPHA2012;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1462888843] HOST ALERT: R25-ZETA2012;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1462888845] HOST ALERT: R25-DELTA2012;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1462888846] HOST ALERT: R25-ALPHA2012;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1462888870] SERVICE ALERT: R25-ALPHA2012;CPU Load;OK;HARD;3;CPU Load 8% (15 min average)
[1462888870] SERVICE ALERT: R25-DELTA2012;Uptime;OK;HARD;3;System Uptime - 0 day(s) 23 hour(s) 21 minute(s)
[1462888881] SERVICE ALERT: R25-DELTA2012;VMWare Tools;OK;HARD;3;VMTools: Started
[1462888882] SERVICE ALERT: R25-DELTA2012;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1462888896] SERVICE ALERT: R25-DELTA2012;Memory Usage;OK;HARD;3;Memory usage: total:10943.64 MB - used: 1476.34 MB (13%) - free: 9467.31 MB (87%)
[1462888896] SERVICE ALERT: R25-ALPHA2012;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1462888897] SERVICE ALERT: R25-ALPHA2012;Memory Usage;OK;HARD;3;Memory usage: total:10943.64 MB - used: 1362.55 MB (12%) - free: 9581.09 MB (88%)
[1462888898] SERVICE ALERT: R25-ZETA2012;Uptime;OK;HARD;3;System Uptime - 10 day(s) 22 hour(s) 27 minute(s)
[1462888898] SERVICE ALERT: R25-ZETA2012;C:\ Drive Space;OK;HARD;3;c:\ - total: 39.66 Gb - used: 16.49 Gb (42%) - free 23.17 Gb (58%)
[1462888927] SERVICE ALERT: R25-DELTA2012;Service: WMI;OK;HARD;3;Winmgmt: Started
[1462888927] SERVICE ALERT: R25-ZETA2012;Service: WMI;OK;HARD;3;Winmgmt: Started
[1462888928] SERVICE ALERT: R25-ZETA2012;VMWare Tools;OK;HARD;3;VMTools: Started
[1462888970] SERVICE ALERT: R25-ZETA2012;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1462889014] SERVICE ALERT: R25-ALPHA2012;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1462889078] SERVICE ALERT: R25-ALPHA2012;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1462889114] SERVICE ALERT: R25-DELTA2012;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1462889120] SERVICE ALERT: R25-ALPHA2012;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1462889120] SERVICE ALERT: R25-DELTA2012;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1462889122] SERVICE ALERT: R25-ZETA2012;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1462889175] SERVICE ALERT: R25-ZETA2012;Memory Usage;OK;HARD;3;Memory usage: total:10943.64 MB - used: 1611.99 MB (15%) - free: 9331.65 MB (85%)
[1462889178] SERVICE ALERT: R25-DELTA2012;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1462889188] SERVICE ALERT: R25-ALPHA2012;Service: RPC;OK;HARD;3;RpcSs: Started
[1462889189] SERVICE ALERT: R25-ZETA2012;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1462889196] SERVICE ALERT: R25-DELTA2012;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1462889204] SERVICE ALERT: R25-ZETA2012;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1462889227] SERVICE ALERT: R25-DELTA2012;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1462889235] SERVICE ALERT: R25-ZETA2012;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1462889240] SERVICE ALERT: R25-ALPHA2012;Uptime;OK;HARD;3;System Uptime - 0 day(s) 18 hour(s) 15 minute(s)
[1462889290] SERVICE ALERT: R25-ALPHA2012;C:\ Drive Space;OK;HARD;3;c:\ - total: 39.66 Gb - used: 16.41 Gb (41%) - free 23.24 Gb (59%)
[1462889294] SERVICE ALERT: R25-ALPHA2012;VMWare Tools;OK;HARD;3;VMTools: Started
[1462889295] SERVICE ALERT: R25-ALPHA2012;Service: WMI;OK;HARD;3;Winmgmt: Started
[1462889315] SERVICE ALERT: R25-DELTA2012;Service: Windows Time;OK;HARD;3;W32Time: Started
[1462889316] SERVICE ALERT: R25-ALPHA2012;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1462889317] SERVICE ALERT: R25-ALPHA2012;Service: Windows Time;OK;HARD;3;W32Time: Started
[1462889335] SERVICE ALERT: R25-DELTA2012;C:\ Drive Space;OK;HARD;3;c:\ - total: 39.66 Gb - used: 19.72 Gb (50%) - free 19.94 Gb (50%)
[1462889340] SERVICE ALERT: R25-DELTA2012;Service: RPC;OK;HARD;3;RpcSs: Started
[1462889346] SERVICE ALERT: R25-ZETA2012;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1462889367] SERVICE ALERT: R25-ZETA2012;Service: RPC;OK;HARD;3;RpcSs: Started
[1462889371] SERVICE ALERT: R25-ZETA2012;CPU Load;OK;HARD;3;CPU Load 1% (15 min average)
[1462889944] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462890059] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462890064] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462890179] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1462890184] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1462890299] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462890784] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1462890899] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 84% (15 min average)
[1462891380] SERVICE ALERT: owen101lab3;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462891439] SERVICE ALERT: owen101lab3;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462891499] SERVICE ALERT: owen101lab3;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462891984] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462892099] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1462892099] SERVICE ALERT: owen101lab3;Device Status;OK;HARD;3;41.3 UNEXPECTED SIZE IN TRAY 4 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   LOAD TRAY 4 PLAIN LETTER To use another 74 72 61 79 20 70 72 65 73 73 20 1E
[1462892104] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1462892219] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 75% (15 min average)
[1462892224] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462892312] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1462892366] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 22% - OK!
[1462892403] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (31.2% change >= 20.0% threshold)
[1462892403] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1462892403] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1462892403] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1462892403] HOST DOWNTIME ALERT: R25-ALPHA2012;STARTED; Host has entered a period of scheduled downtime
[1462892403] HOST DOWNTIME ALERT: R25-DELTA2012;STARTED; Host has entered a period of scheduled downtime
[1462892403] SERVICE DOWNTIME ALERT: NETIDADPROXY;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1462892403] HOST DOWNTIME ALERT: R25-ZETA2012;STARTED; Host has entered a period of scheduled downtime
[1462892403] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1462892403] HOST DOWNTIME ALERT: DIRECTACCESS;STARTED; Host has entered a period of scheduled downtime
[1462892403] HOST DOWNTIME ALERT: FITNESSGRAM;STARTED; Host has entered a period of scheduled downtime
[1462892815] HOST DOWNTIME ALERT: ADDC66;STARTED; Host has entered a period of scheduled downtime
[1462892824] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 84% (15 min average)
[1462892846] SERVICE ALERT: ADDC66;Service: Firewall;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1462892966] SERVICE ALERT: ADDC66;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1462894019] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1462894024] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462894137] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 65% (15 min average)
[1462894144] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 71% (15 min average)
[1462894737] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462894843] SERVICE ALERT: ADDC5;GPO Replication;OK;HARD;3;OK! All GPOs are properly replicated.
[1462894843] SERVICE FLAPPING ALERT: ADDC5;GPO Replication;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1462894857] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1462894977] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 68% (15 min average)
[1462895343] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1462895463] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 67% (15 min average)
[1462896987] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462897047] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462897107] SERVICE ALERT: hallefirst1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462897263] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462897383] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462897503] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462897707] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;Ready
[1462897923] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462897983] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462898043] SERVICE ALERT: mckennypk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462898103] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462898438] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 86% - OK!
[1462899303] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462899372] SERVICE ALERT: marshallpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.137.49. Is the printer off?  Is the network down?
[1462899402] HOST ALERT: marshallpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462899432] SERVICE ALERT: marshallpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.137.49. Is the printer off?  Is the network down?
[1462899457] HOST ALERT: marshallpk;UP;SOFT;2;PING WARNING - Packet loss = 80%, RTA = 0.71 ms
[1462899486] SERVICE ALERT: marshallpk;Device Status;OK;SOFT;3;Processing job from tray 3
[1462899778] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462899897] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1462900008] HOST DOWNTIME ALERT: ADDC66;STOPPED; Host has exited from a period of scheduled downtime
[1462900503] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 42% (15 min average)
[1462901097] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462901103] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462901118] SERVICE ALERT: NETENG-PA;Uptime;CRITICAL;SOFT;1;connect to address 164.76.102.208 and port 12489: No route to host
[1462901121] HOST ALERT: NETENG-PA;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.208)
[1462901133] SERVICE ALERT: NETENG-PA;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.102.208 and port 12489: No route to host
[1462901136] HOST ALERT: NETENG-PA;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.208)
[1462901146] SERVICE ALERT: NETENG-PA;VMWare Tools;CRITICAL;SOFT;1;connect to address 164.76.102.208 and port 12489: No route to host
[1462901161] SERVICE ALERT: NETENG-PA;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.102.208 and port 12489: No route to host
[1462901164] HOST ALERT: NETENG-PA;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.102.208)
[1462901196] SERVICE ALERT: NETENG-PA;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.102.208 and port 12489: No route to host
[1462901199] HOST ALERT: NETENG-PA;DOWN;SOFT;4;CRITICAL - Host Unreachable (164.76.102.208)
[1462901216] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 80% (15 min average)
[1462901217] HOST DOWNTIME ALERT: NETENG-PA;STARTED; Host has entered a period of scheduled downtime
[1462901223] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462901229] SERVICE ALERT: NETENG-PA;Uptime;CRITICAL;SOFT;2;connect to address 164.76.102.208 and port 12489: No route to host
[1462901232] HOST ALERT: NETENG-PA;DOWN;SOFT;5;CRITICAL - Host Unreachable (164.76.102.208)
[1462901253] SERVICE ALERT: NETENG-PA;C:\ Drive Space;CRITICAL;SOFT;2;connect to address 164.76.102.208 and port 12489: No route to host
[1462901253] SERVICE ALERT: NETENG-PA;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.102.208 and port 12489: No route to host
[1462901256] HOST ALERT: NETENG-PA;DOWN;SOFT;6;CRITICAL - Host Unreachable (164.76.102.208)
[1462901266] SERVICE ALERT: NETENG-PA;VMWare Tools;CRITICAL;SOFT;2;connect to address 164.76.102.208 and port 12489: No route to host
[1462901281] SERVICE ALERT: NETENG-PA;Service: Windows Time;CRITICAL;SOFT;2;connect to address 164.76.102.208 and port 12489: No route to host
[1462901284] HOST ALERT: NETENG-PA;DOWN;SOFT;7;CRITICAL - Host Unreachable (164.76.102.208)
[1462901316] SERVICE ALERT: NETENG-PA;CPU Load;CRITICAL;SOFT;1;connect to address 164.76.102.208 and port 12489: No route to host
[1462901316] SERVICE ALERT: NETENG-PA;Service: RPC;CRITICAL;SOFT;2;connect to address 164.76.102.208 and port 12489: No route to host
[1462901316] SERVICE ALERT: NETENG-PA;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.102.208 and port 12489: No route to host
[1462901319] HOST ALERT: NETENG-PA;DOWN;SOFT;8;CRITICAL - Host Unreachable (164.76.102.208)
[1462901343] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462901349] SERVICE ALERT: NETENG-PA;Uptime;CRITICAL;HARD;3;connect to address 164.76.102.208 and port 12489: No route to host
[1462901352] HOST ALERT: NETENG-PA;DOWN;SOFT;9;CRITICAL - Host Unreachable (164.76.102.208)
[1462901373] SERVICE ALERT: NETENG-PA;Memory Usage;CRITICAL;SOFT;2;connect to address 164.76.102.208 and port 12489: No route to host
[1462901373] SERVICE ALERT: NETENG-PA;C:\ Drive Space;CRITICAL;HARD;3;connect to address 164.76.102.208 and port 12489: No route to host
[1462901376] HOST ALERT: NETENG-PA;DOWN;HARD;10;CRITICAL - Host Unreachable (164.76.102.208)
[1462901386] SERVICE ALERT: NETENG-PA;VMWare Tools;CRITICAL;HARD;3;connect to address 164.76.102.208 and port 12489: No route to host
[1462901401] SERVICE ALERT: NETENG-PA;Service: Windows Time;CRITICAL;HARD;3;connect to address 164.76.102.208 and port 12489: No route to host
[1462901411] SERVICE ALERT: NETENG-PA;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.102.208 and port 12489: No route to host
[1462901431] SERVICE ALERT: NETENG-PA;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.102.208 and port 12489: No route to host
[1462901434] SERVICE ALERT: NETENG-PA;Service: RPC;CRITICAL;HARD;3;connect to address 164.76.102.208 and port 12489: No route to host
[1462901434] SERVICE ALERT: NETENG-PA;Service: Firewall;CRITICAL;SOFT;2;connect to address 164.76.102.208 and port 12489: No route to host
[1462901437] SERVICE ALERT: NETENG-PA;CPU Load;CRITICAL;SOFT;2;connect to address 164.76.102.208 and port 12489: No route to host
[1462901469] SERVICE ALERT: NETENG-PA;Service: NetLogin;CRITICAL;SOFT;1;connect to address 164.76.102.208 and port 12489: No route to host
[1462901494] SERVICE ALERT: NETENG-PA;Memory Usage;CRITICAL;HARD;3;connect to address 164.76.102.208 and port 12489: No route to host
[1462901501] SERVICE ALERT: NETENG-PA;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.102.208 and port 12489: No route to host
[1462901531] SERVICE ALERT: NETENG-PA;NSClient++ Version;CRITICAL;SOFT;2;connect to address 164.76.102.208 and port 12489: No route to host
[1462901551] SERVICE ALERT: NETENG-PA;Service: Group Policy;CRITICAL;SOFT;2;connect to address 164.76.102.208 and port 12489: No route to host
[1462901554] SERVICE ALERT: NETENG-PA;Service: Firewall;CRITICAL;HARD;3;connect to address 164.76.102.208 and port 12489: No route to host
[1462901554] SERVICE ALERT: NETENG-PA;Service: User Profile;CRITICAL;SOFT;1;connect to address 164.76.102.208 and port 12489: No route to host
[1462901558] SERVICE ALERT: NETENG-PA;CPU Load;CRITICAL;HARD;3;connect to address 164.76.102.208 and port 12489: No route to host
[1462901589] SERVICE ALERT: NETENG-PA;Service: NetLogin;CRITICAL;SOFT;2;connect to address 164.76.102.208 and port 12489: No route to host
[1462901616] SERVICE ALERT: NETENG-PA;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.102.208 and port 12489: No route to host
[1462901619] SERVICE ALERT: NETENG-PA;Service: DNS Client;CRITICAL;SOFT;2;connect to address 164.76.102.208 and port 12489: No route to host
[1462901651] SERVICE ALERT: NETENG-PA;NSClient++ Version;CRITICAL;HARD;3;connect to address 164.76.102.208 and port 12489: No route to host
[1462901671] SERVICE ALERT: NETENG-PA;Service: Group Policy;CRITICAL;HARD;3;connect to address 164.76.102.208 and port 12489: No route to host
[1462901674] SERVICE ALERT: NETENG-PA;Service: User Profile;CRITICAL;SOFT;2;connect to address 164.76.102.208 and port 12489: No route to host
[1462901709] SERVICE ALERT: NETENG-PA;Service: NetLogin;CRITICAL;HARD;3;connect to address 164.76.102.208 and port 12489: No route to host
[1462901736] SERVICE ALERT: NETENG-PA;Service: WMI;CRITICAL;SOFT;2;connect to address 164.76.102.208 and port 12489: No route to host
[1462901739] SERVICE ALERT: NETENG-PA;Service: DNS Client;CRITICAL;HARD;3;connect to address 164.76.102.208 and port 12489: No route to host
[1462901796] SERVICE ALERT: NETENG-PA;Service: User Profile;CRITICAL;HARD;3;connect to address 164.76.102.208 and port 12489: No route to host
[1462901856] SERVICE ALERT: NETENG-PA;Service: WMI;CRITICAL;HARD;3;connect to address 164.76.102.208 and port 12489: No route to host
[1462901943] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 88% (15 min average)
[1462902415] SERVICE ALERT: mckennypk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1462902475] SERVICE ALERT: mckennypk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1462902505] HOST ALERT: mckennypk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462902535] SERVICE ALERT: mckennypk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1462902542] HOST ALERT: mckennypk;UP;SOFT;2;PING OK - Packet loss = 37%, RTA = 0.62 ms
[1462902543] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462902638] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;CRITICAL;SOFT;1;Black Print Cartridge HP C8543X is at CRITICAL level!
[1462902663] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 66% (15 min average)
[1462902699] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;CRITICAL;SOFT;2;Black Print Cartridge HP C8543X is at CRITICAL level!
[1462902758] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;CRITICAL;HARD;3;Black Print Cartridge HP C8543X is at CRITICAL level!
[1462902843] SERVICE ALERT: mckennypk;Device Status;OK;HARD;3;CLOSE FRONT OR LEFT DOOR To continue turn off then on CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1462903130] SERVICE ALERT: mckennypk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1462903863] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462903983] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1462904103] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1462905416] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462905536] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1462905678] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1462905739] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1462905799] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1462905903] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1462906217] HOST DOWNTIME ALERT: TABLEAU;STARTED; Host has entered a period of scheduled downtime
[1462906224] HOST ALERT: TABLEAU;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462906234] SERVICE ALERT: TABLEAU;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462906264] SERVICE ALERT: TABLEAU;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462906264] HOST ALERT: TABLEAU;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1462906354] SERVICE ALERT: TABLEAU;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462906354] HOST ALERT: TABLEAU;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1462906384] SERVICE ALERT: TABLEAU;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462906404] SERVICE ALERT: TABLEAU;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462906414] HOST ALERT: TABLEAU;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1462906428] SERVICE ALERT: TABLEAU;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462906433] SERVICE ALERT: TABLEAU;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462906433] SERVICE ALERT: TABLEAU;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462906436] SERVICE ALERT: TABLEAU;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462906458] HOST ALERT: TABLEAU;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1462906459] SERVICE ALERT: TABLEAU;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462906459] SERVICE ALERT: TABLEAU;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462906474] SERVICE ALERT: TABLEAU;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462906489] HOST ALERT: TABLEAU;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1462906497] SERVICE ALERT: TABLEAU;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462906504] SERVICE ALERT: TABLEAU;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462906524] SERVICE ALERT: TABLEAU;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462906527] HOST ALERT: TABLEAU;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1462906538] SERVICE ALERT: TABLEAU;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462906548] SERVICE ALERT: TABLEAU;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462906553] SERVICE ALERT: TABLEAU;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462906553] SERVICE ALERT: TABLEAU;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462906556] SERVICE ALERT: TABLEAU;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462906568] HOST ALERT: TABLEAU;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1462906579] SERVICE ALERT: TABLEAU;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462906579] SERVICE ALERT: TABLEAU;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462906579] SERVICE ALERT: TABLEAU;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462906609] HOST ALERT: TABLEAU;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1462906617] SERVICE ALERT: TABLEAU;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462906623] SERVICE ALERT: TABLEAU;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462906644] SERVICE ALERT: TABLEAU;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462906647] HOST ALERT: TABLEAU;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1462906658] SERVICE ALERT: TABLEAU;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462906666] SERVICE ALERT: TABLEAU;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462906668] SERVICE ALERT: TABLEAU;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462906673] SERVICE ALERT: TABLEAU;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462906673] SERVICE ALERT: TABLEAU;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462906676] SERVICE ALERT: TABLEAU;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462906699] SERVICE ALERT: TABLEAU;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462906699] SERVICE ALERT: TABLEAU;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462906699] SERVICE ALERT: TABLEAU;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462906737] SERVICE ALERT: TABLEAU;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462906743] SERVICE ALERT: TABLEAU;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462906778] SERVICE ALERT: TABLEAU;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462906786] SERVICE ALERT: TABLEAU;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462906819] SERVICE ALERT: TABLEAU;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462906863] SERVICE ALERT: TABLEAU;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462906906] SERVICE ALERT: TABLEAU;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462906932] SERVICE ALERT: ADDC44;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1462906939] SERVICE ALERT: ADDC44;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462906939] SERVICE ALERT: ADDC44;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462906943] SERVICE ALERT: ADDC44;Service: AD Web Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462906977] HOST DOWNTIME ALERT: DIRECTACCESS;STOPPED; Host has exited from a period of scheduled downtime
[1462907039] SERVICE ALERT: ADDC44;VMWare Tools;OK;SOFT;2;VMTools: Started
[1462907039] SERVICE ALERT: ADDC44;C:\ Drive Space;OK;SOFT;2;c:\ - total: 79.66 Gb - used: 25.42 Gb (32%) - free 54.24 Gb (68%)
[1462907043] SERVICE ALERT: ADDC44;Service: AD Web Service;OK;SOFT;2;ADWS: Started
[1462907049] SERVICE ALERT: ADDC44;Service: Windows Time;CRITICAL;SOFT;2;W32Time: Stopped
[1462907103] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462907169] SERVICE ALERT: ADDC44;Service: Windows Time;OK;SOFT;3;W32Time: Started
[1462907223] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462907336] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1462907343] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462907456] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 87% (15 min average)
[1462907559] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1462907943] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 90% (15 min average)
[1462908056] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1462908158] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1462908176] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1462908198] SERVICE ALERT: marshallpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1462908218] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 100% - OK!
[1462908296] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1462908390] SERVICE FLAPPING ALERT: ADDC5;GPO Replication;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1462908390] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (51.4% change >= 20.0% threshold)
[1462908390] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1462908390] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1462908390] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1462908390] HOST DOWNTIME ALERT: R25-DELTA2012;STARTED; Host has entered a period of scheduled downtime
[1462908390] HOST DOWNTIME ALERT: R25-ZETA2012;STARTED; Host has entered a period of scheduled downtime
[1462908390] SERVICE DOWNTIME ALERT: NETIDADPROXY;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1462908390] HOST DOWNTIME ALERT: NETENG-PA;STARTED; Host has entered a period of scheduled downtime
[1462908390] HOST DOWNTIME ALERT: TABLEAU;STARTED; Host has entered a period of scheduled downtime
[1462908390] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1462908390] HOST DOWNTIME ALERT: FITNESSGRAM;STARTED; Host has entered a period of scheduled downtime
[1462908390] HOST DOWNTIME ALERT: R25-ALPHA2012;STARTED; Host has entered a period of scheduled downtime
[1462908602] SERVICE FLAPPING ALERT: ADDC5;GPO Replication;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1462908603] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (50.1% change >= 20.0% threshold)
[1462908603] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1462908603] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1462908603] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1462908603] HOST DOWNTIME ALERT: R25-DELTA2012;STARTED; Host has entered a period of scheduled downtime
[1462908603] HOST DOWNTIME ALERT: R25-ZETA2012;STARTED; Host has entered a period of scheduled downtime
[1462908603] SERVICE DOWNTIME ALERT: NETIDADPROXY;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1462908603] HOST DOWNTIME ALERT: NETENG-PA;STARTED; Host has entered a period of scheduled downtime
[1462908603] HOST DOWNTIME ALERT: TABLEAU;STARTED; Host has entered a period of scheduled downtime
[1462908603] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1462908603] HOST DOWNTIME ALERT: FITNESSGRAM;STARTED; Host has entered a period of scheduled downtime
[1462908603] HOST DOWNTIME ALERT: R25-ALPHA2012;STARTED; Host has entered a period of scheduled downtime
[1462908896] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1462908900] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462908960] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462909020] SERVICE ALERT: sillpk;Device Status;OK;SOFT;3;Open tray indicated.
[1462909143] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1462909263] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462909282] SERVICE ALERT: ADDC5;GPO Replication;UNKNOWN;SOFT;1;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1462909383] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462909526] SERVICE ALERT: ADDC5;GPO Replication;OK;SOFT;2;OK! All GPOs are properly replicated.
[1462909543] HOST DOWNTIME ALERT: FITNESSGRAM;STOPPED; Host has exited from a period of scheduled downtime
[1462910095] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1462910215] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 83% (15 min average)
[1462910220] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462910280] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462910308] SERVICE ALERT: prayharroldpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462910341] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462910368] SERVICE ALERT: prayharroldpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462910428] SERVICE ALERT: prayharroldpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462910940] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Open tray indicated.
[1462911028] SERVICE ALERT: prayharroldpk;Device Status;OK;HARD;3;Ready
[1462911783] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 59% (15 min average)
[1462912140] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462912200] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462912260] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462912444] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462912503] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462912563] SERVICE ALERT: mckennypk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462912860] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Open tray indicated.
[1462912860] SERVICE FLAPPING ALERT: sillpk;Device Status;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1462914061] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462914120] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462914180] SERVICE ALERT: sillpk;Device Status;OK;SOFT;3;Open tray indicated.
[1462915015] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1462915135] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1462915380] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462915440] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462915500] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462916100] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Open tray indicated.
[1462916335] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1462916455] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 65% (15 min average)
[1462916583] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462916703] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462916823] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462917055] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1462917175] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 77% (15 min average)
[1462917300] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462917360] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462917420] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462917423] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462917776] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462917895] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1462918020] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Open tray indicated.
[1462918495] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462918615] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1462918735] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1462919220] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462919281] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462919335] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 82% (15 min average)
[1462919335] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1462919340] SERVICE ALERT: sillpk;Device Status;OK;SOFT;3;Open tray indicated.
[1462920540] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462920600] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462920660] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462921260] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Open tray indicated.
[1462922460] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462922520] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462922580] SERVICE ALERT: sillpk;Device Status;OK;SOFT;3;Open tray indicated. TRAY 2
[1462923423] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (15 min average)
[1462924023] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462924143] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462924263] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1462924380] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462924441] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462924500] SERVICE ALERT: sillpk;Device Status;OK;SOFT;3;Open tray indicated.
[1462925700] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462925760] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462925820] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462926420] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Open tray indicated.
[1462926663] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462927263] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 77% (15 min average)
[1462927620] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462927679] SERVICE ALERT: sillpk;Device Status;OK;SOFT;2;Open tray indicated.
[1462927863] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462927983] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462928103] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1462928879] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462928939] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462928999] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462929599] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Open tray indicated.
[1462929698] SERVICE ALERT: SANDBOX-DB;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1462929818] SERVICE ALERT: SANDBOX-DB;CPU Load;OK;SOFT;2;CPU Load 86% (15 min average)
[1462930734] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1462930799] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462930859] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462930919] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462931519] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Open tray indicated.
[1462932303] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462932719] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462932779] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462932840] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462932903] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462933439] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Open tray indicated.
[1462934039] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462934099] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462934159] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462934536] SERVICE ALERT: VC;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462934636] SERVICE ALERT: VC;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1462934759] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Open tray indicated.
[1462935960] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462936020] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462936079] SERVICE ALERT: sillpk;Device Status;OK;SOFT;3;Open tray indicated.
[1462936134] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462936254] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1462936374] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1462937103] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462937278] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462937338] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462937398] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462937703] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1462937998] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Open tray indicated.
[1462938174] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 81% (15 min average)
[1462938303] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 90% (15 min average)
[1462938903] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462939023] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462939143] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462939198] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!