[1453583059] SERVICE ALERT: halleground1;Paper Status Tray 3;CRITICAL;SOFT;1;Tray 3 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453583118] SERVICE ALERT: halleground1;Paper Status Tray 3;CRITICAL;SOFT;2;Tray 3 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453583179] SERVICE ALERT: halleground1;Paper Status Tray 3;CRITICAL;HARD;3;Tray 3 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453594960] SERVICE ALERT: mckennypk;Paper Status Tray 3;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.152.26! Make sure host is up and SNMP is configured properly.
[1453594990] HOST ALERT: mckennypk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1453595020] SERVICE ALERT: mckennypk;Paper Status Tray 3;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.152.26! Make sure host is up and SNMP is configured properly.
[1453595050] HOST ALERT: mckennypk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1453595080] SERVICE ALERT: mckennypk;Paper Status Tray 3;WARNING;HARD;3;WARNING: No SNMP response from 164.76.152.26! Make sure host is up and SNMP is configured properly.
[1453595110] HOST ALERT: mckennypk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1453595200] HOST ALERT: mckennypk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1453595290] HOST ALERT: mckennypk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1453595380] HOST ALERT: mckennypk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1453595447] SERVICE ALERT: mckennypk;Maintenance Kit HP 110V-CB388A;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.152.26! Make sure host is up and SNMP is configured properly.
[1453595470] HOST ALERT: mckennypk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1453595503] SERVICE ALERT: mckennypk;Paper Status Tray 2;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.152.26! Make sure host is up and SNMP is configured properly.
[1453595507] SERVICE ALERT: mckennypk;Maintenance Kit HP 110V-CB388A;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.152.26! Make sure host is up and SNMP is configured properly.
[1453595533] HOST ALERT: mckennypk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1453595534] SERVICE ALERT: mckennypk;Black Cartridge HP CC364X;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.152.26! Make sure host is up and SNMP is configured properly.
[1453595563] SERVICE ALERT: mckennypk;Paper Status Tray 2;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.152.26! Make sure host is up and SNMP is configured properly.
[1453595564] HOST ALERT: mckennypk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1453595567] SERVICE ALERT: mckennypk;Maintenance Kit HP 110V-CB388A;WARNING;HARD;3;WARNING: No SNMP response from 164.76.152.26! Make sure host is up and SNMP is configured properly.
[1453595594] SERVICE ALERT: mckennypk;Black Cartridge HP CC364X;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.152.26! Make sure host is up and SNMP is configured properly.
[1453595597] HOST ALERT: mckennypk;DOWN;HARD;10;(Host check timed out after 30.01 seconds)
[1453595623] SERVICE ALERT: mckennypk;Paper Status Tray 2;WARNING;HARD;3;WARNING: No SNMP response from 164.76.152.26! Make sure host is up and SNMP is configured properly.
[1453595654] SERVICE ALERT: mckennypk;Black Cartridge HP CC364X;WARNING;HARD;3;WARNING: No SNMP response from 164.76.152.26! Make sure host is up and SNMP is configured properly.
[1453600939] SERVICE ALERT: MBAM;CPU Load;WARNING;SOFT;1;CPU Load 83% (5 min average)
[1453601059] SERVICE ALERT: MBAM;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1453608371] SERVICE ALERT: halleground2;Paper Status Tray 5;CRITICAL;SOFT;1;Tray 5 is at 5% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453608410] SERVICE ALERT: halleground2;Paper Status Tray 2;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.106.52! Make sure host is up and SNMP is configured properly.
[1453608425] SERVICE ALERT: halleground2;Maintenance Kit HP 110V-C9152A;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.106.52! Make sure host is up and SNMP is configured properly.
[1453608436] SERVICE ALERT: halleground2;Paper Status Tray 5;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.106.52! Make sure host is up and SNMP is configured properly.
[1453608470] SERVICE ALERT: halleground2;Paper Status Tray 2;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.106.52! Make sure host is up and SNMP is configured properly.
[1453608485] SERVICE ALERT: halleground2;Maintenance Kit HP 110V-C9152A;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.106.52! Make sure host is up and SNMP is configured properly.
[1453608496] SERVICE ALERT: halleground2;Paper Status Tray 5;WARNING;HARD;3;WARNING: No SNMP response from 164.76.106.52! Make sure host is up and SNMP is configured properly.
[1453608530] SERVICE ALERT: halleground2;Paper Status Tray 2;WARNING;HARD;3;WARNING: No SNMP response from 164.76.106.52! Make sure host is up and SNMP is configured properly.
[1453608545] SERVICE ALERT: halleground2;Maintenance Kit HP 110V-C9152A;WARNING;HARD;3;WARNING: No SNMP response from 164.76.106.52! Make sure host is up and SNMP is configured properly.
[1453608701] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.106.51! Make sure host is up and SNMP is configured properly.
[1453608731] HOST ALERT: halleground1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1453608754] SERVICE ALERT: halleground1;Maintenance Kit HP 110V-C9152A;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.106.51! Make sure host is up and SNMP is configured properly.
[1453608761] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.106.51! Make sure host is up and SNMP is configured properly.
[1453608771] HOST ALERT: halleground1;UP;SOFT;2;PING OK - Packet loss = 72%, RTA = 0.65 ms
[1453608808] SERVICE ALERT: halleground1;Maintenance Kit HP 110V-C9152A;OK;SOFT;2;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 51% - OK!
[1453608815] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;OK;SOFT;3;Black Print Cartridge HP C8543X is at 31% - OK!
[1453608917] SERVICE ALERT: halleground1;Paper Status Tray 5;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.106.51! Make sure host is up and SNMP is configured properly.
[1453608971] SERVICE ALERT: halleground1;Paper Status Tray 5;OK;SOFT;2;Tray 5 is at 40% - OK!
[1453609090] SERVICE ALERT: halleground2;Paper Status Tray 5;CRITICAL;HARD;3;Tray 5 is at CRITICAL level - please refill with more 11.0 x 8.5 paper.
[1453609125] SERVICE ALERT: halleground2;Paper Status Tray 2;OK;HARD;3;Tray 2 is at 40% - OK!
[1453609139] SERVICE ALERT: halleground2;Maintenance Kit HP 110V-C9152A;OK;HARD;3;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 62% - OK!
[1453609407] SERVICE ALERT: halleground1;Paper Status Tray 2;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.106.51! Make sure host is up and SNMP is configured properly.
[1453609414] SERVICE ALERT: halleground1;Maintenance Kit HP 110V-C9152A;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.106.51! Make sure host is up and SNMP is configured properly.
[1453609421] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.106.51! Make sure host is up and SNMP is configured properly.
[1453609467] SERVICE ALERT: halleground1;Paper Status Tray 2;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.106.51! Make sure host is up and SNMP is configured properly.
[1453609474] SERVICE ALERT: halleground1;Maintenance Kit HP 110V-C9152A;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.106.51! Make sure host is up and SNMP is configured properly.
[1453609481] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.106.51! Make sure host is up and SNMP is configured properly.
[1453609527] SERVICE ALERT: halleground1;Paper Status Tray 2;WARNING;HARD;3;WARNING: No SNMP response from 164.76.106.51! Make sure host is up and SNMP is configured properly.
[1453609534] SERVICE ALERT: halleground1;Maintenance Kit HP 110V-C9152A;WARNING;HARD;3;WARNING: No SNMP response from 164.76.106.51! Make sure host is up and SNMP is configured properly.
[1453609541] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;WARNING;HARD;3;WARNING: No SNMP response from 164.76.106.51! Make sure host is up and SNMP is configured properly.
[1453609577] SERVICE ALERT: halleground1;Paper Status Tray 5;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.106.51! Make sure host is up and SNMP is configured properly.
[1453609584] SERVICE ALERT: halleground1;Paper Status Tray 3;WARNING;HARD;3;WARNING: No SNMP response from 164.76.106.51! Make sure host is up and SNMP is configured properly.
[1453609637] SERVICE ALERT: halleground1;Paper Status Tray 5;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.106.51! Make sure host is up and SNMP is configured properly.
[1453609697] SERVICE ALERT: halleground1;Paper Status Tray 5;WARNING;HARD;3;WARNING: No SNMP response from 164.76.106.51! Make sure host is up and SNMP is configured properly.
[1453610757] HOST ALERT: halleground1;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1453610783] SERVICE ALERT: halleground1;Paper Status Tray 3;CRITICAL;HARD;3;Tray 3 is at CRITICAL level - please refill with more 11.0 x 8.5 paper.
[1453610787] HOST ALERT: halleground1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.57 ms
[1453611384] SERVICE ALERT: halleground1;Paper Status Tray 3;WARNING;HARD;3;WARNING: No SNMP response from 164.76.106.51! Make sure host is up and SNMP is configured properly.
[1453611491] SERVICE ALERT: halleground2;Paper Status Tray 5;OK;HARD;3;Tray 5 is at 40% - OK!