[1456650187] SERVICE FLAPPING ALERT: ADDC5;GPO Replication;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1456690651] SERVICE ALERT: DB01-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1456690771] SERVICE ALERT: DB01-2012;CPU Load;OK;SOFT;2;CPU Load 3% (5 min average)
[1456703068] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1456703127] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1456703188] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1456706787] SERVICE ALERT: crossroadspk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1456712203] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1456712263] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1456712323] SERVICE ALERT: hallefirst1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1456712372] SERVICE ALERT: SANDBOX-DB;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1456712412] SERVICE ALERT: EVENTSTEST-2012;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1456712492] SERVICE ALERT: SANDBOX-DB;CPU Load;OK;SOFT;2;CPU Load 33% (5 min average)
[1456712532] SERVICE ALERT: EVENTSTEST-2012;CPU Load;OK;SOFT;2;CPU Load 85% (5 min average)
[1456713523] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;Processing job from tray 4