[1463889623] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463889691] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 219.90 Gb - used: 201.22 Gb (92%) - free 18.68 Gb (8%)
[1463889811] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 201.91 Gb (92%) - free 17.99 Gb (8%)
[1463890223] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1463890823] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463890943] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463891011] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 209.60 Gb (95%) - free 10.30 Gb (5%)
[1463891063] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463892811] SERVICE ALERT: BOEAPP;C:\ Drive Space;OK;HARD;3;c:\ - total: 219.90 Gb - used: 189.04 Gb (86%) - free 30.86 Gb (14%)
[1463892863] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463893463] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463894063] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 56% (15 min average)
[1463894643] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463894663] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463894783] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 89% (15 min average)
[1463895211] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 219.90 Gb - used: 199.49 Gb (91%) - free 20.41 Gb (9%)
[1463895243] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 290 found since last run. oneLevelSearchOps OK. searchOps is too high! 305 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 616 found since last run.
[1463895314] SERVICE ALERT: idm2.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1463895331] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 219.90 Gb - used: 200.19 Gb (91%) - free 19.70 Gb (9%)
[1463895356] SERVICE ALERT: idm3.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1463895363] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463895383] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1463895412] SERVICE ALERT: idm1.emich.edu;NSDS Errors;OK;HARD;3;OK: 0 Errors Reported
[1463895451] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 200.89 Gb (91%) - free 19.01 Gb (9%)
[1463895503] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463895584] SERVICE ALERT: idm4.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1463895623] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463895963] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 602 found since last run.
[1463896083] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463896683] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 330 found since last run. oneLevelSearchOps OK. searchOps is too high! 337 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 711 found since last run.
[1463896803] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463896823] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463897251] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 212.44 Gb (97%) - free 7.46 Gb (3%)
[1463897403] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 427 found since last run. oneLevelSearchOps OK. searchOps is too high! 433 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 654 found since last run.
[1463897523] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463898124] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 386 found since last run. oneLevelSearchOps OK. searchOps is too high! 398 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 552 found since last run.
[1463898243] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463898355] SERVICE ALERT: idm2.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 8
[1463898356] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 3
[1463898412] SERVICE ALERT: idm1.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 1
[1463898433] SERVICE ALERT: idm2.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 1
[1463898451] SERVICE ALERT: BOEAPP;C:\ Drive Space;OK;HARD;3;c:\ - total: 219.90 Gb - used: 189.27 Gb (86%) - free 30.63 Gb (14%)
[1463898477] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 1
[1463898532] SERVICE ALERT: idm1.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 1
[1463898553] SERVICE ALERT: idm2.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 2
[1463898584] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 2
[1463898596] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 7
[1463898652] SERVICE ALERT: idm1.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 1
[1463898705] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 1
[1463898824] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 1
[1463898843] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 506 found since last run.
[1463898963] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463899823] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463900353] SERVICE ALERT: idm2.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1463900396] SERVICE ALERT: idm3.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1463900452] SERVICE ALERT: idm1.emich.edu;NSDS Errors;OK;HARD;3;OK: 0 Errors Reported
[1463900623] SERVICE ALERT: idm4.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1463900851] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 219.90 Gb - used: 199.33 Gb (91%) - free 20.56 Gb (9%)
[1463900971] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 219.90 Gb - used: 200.73 Gb (91%) - free 19.17 Gb (9%)
[1463901091] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 201.01 Gb (91%) - free 18.88 Gb (9%)
[1463901363] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 553 found since last run.
[1463901483] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463902890] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 212.22 Gb (97%) - free 7.68 Gb (3%)
[1463903284] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 265 found since last run. oneLevelSearchOps OK. searchOps is too high! 268 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463903403] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463904003] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 236 found since last run. oneLevelSearchOps OK. searchOps is too high! 238 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 556 found since last run.
[1463904089] SERVICE ALERT: BOEAPP;C:\ Drive Space;OK;HARD;3;c:\ - total: 219.90 Gb - used: 189.27 Gb (86%) - free 30.63 Gb (14%)
[1463904123] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463904623] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 48% (15 min average)
[1463904723] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 299 found since last run. oneLevelSearchOps OK. searchOps is too high! 300 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 662 found since last run.
[1463904843] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463905443] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 447 found since last run. oneLevelSearchOps OK. searchOps is too high! 449 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 679 found since last run.
[1463905563] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463906163] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 585 found since last run. oneLevelSearchOps OK. searchOps is too high! 588 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 777 found since last run.
[1463906283] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463906423] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463906543] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1463906663] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1463906883] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 371 found since last run. oneLevelSearchOps OK. searchOps is too high! 372 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 691 found since last run.
[1463907003] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463907263] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 76% (15 min average)
[1463907603] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 319 found since last run. oneLevelSearchOps OK. searchOps is too high! 321 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 698 found since last run.
[1463907689] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 219.90 Gb - used: 201.22 Gb (92%) - free 18.68 Gb (8%)
[1463907723] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463907809] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 219.90 Gb - used: 201.89 Gb (92%) - free 18.01 Gb (8%)
[1463907929] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 202.69 Gb (92%) - free 17.21 Gb (8%)
[1463908323] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 312 found since last run. oneLevelSearchOps OK. searchOps is too high! 313 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 676 found since last run.
[1463908443] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463908463] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1463908578] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463908583] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463908703] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463908991] SERVICE ALERT: EVENTS-2012;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1463909043] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 430 found since last run. oneLevelSearchOps OK. searchOps is too high! 435 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 735 found since last run.
[1463909111] SERVICE ALERT: EVENTS-2012;CPU Load;OK;SOFT;2;CPU Load 55% (15 min average)
[1463909129] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 210.55 Gb (96%) - free 9.35 Gb (4%)
[1463909163] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463909178] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 218 found since last run. oneLevelSearchOps OK. searchOps is too high! 227 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463909299] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463909303] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 91% (15 min average)
[1463909763] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 412 found since last run. oneLevelSearchOps OK. searchOps is too high! 414 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 772 found since last run.
[1463909883] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463909898] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 210 found since last run. oneLevelSearchOps OK. searchOps is too high! 215 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463909903] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463910018] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463910023] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463910143] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 77% (15 min average)
[1463910484] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 542 found since last run. oneLevelSearchOps OK. searchOps is too high! 545 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 757 found since last run.
[1463910603] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463910618] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 200 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463910738] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463910743] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463910863] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463910929] SERVICE ALERT: BOEAPP;C:\ Drive Space;OK;HARD;3;c:\ - total: 219.90 Gb - used: 189.64 Gb (86%) - free 30.26 Gb (14%)
[1463910983] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463911203] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 520 found since last run. oneLevelSearchOps OK. searchOps is too high! 521 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 672 found since last run.
[1463911297] SERVICE ALERT: cava2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1463911323] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463911417] SERVICE ALERT: cava2;CPU Load;OK;SOFT;2;CPU Load 0% (15 min average)
[1463911923] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 448 found since last run. oneLevelSearchOps OK. searchOps is too high! 450 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 743 found since last run.
[1463912043] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463912538] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 214 found since last run. oneLevelSearchOps OK. searchOps is too high! 215 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463912643] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 361 found since last run. oneLevelSearchOps OK. searchOps is too high! 365 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 686 found since last run.
[1463912658] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463912729] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 219.90 Gb - used: 198.54 Gb (90%) - free 21.36 Gb (10%)
[1463912763] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463912783] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 82% (15 min average)
[1463912849] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 219.90 Gb - used: 198.49 Gb (90%) - free 21.41 Gb (10%)
[1463912969] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 199.22 Gb (91%) - free 20.68 Gb (9%)
[1463913363] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 299 found since last run. oneLevelSearchOps OK. searchOps is too high! 300 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 763 found since last run.
[1463913483] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463913983] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463914083] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 485 found since last run. oneLevelSearchOps OK. searchOps is too high! 488 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 720 found since last run.
[1463914103] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463914203] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463914223] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 87% (15 min average)
[1463914260] SERVICE ALERT: ADWU1;Service: DNS Client;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1463914261] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;UNKNOWN;SOFT;1;Failed to process command : EvtQuery failed: 1717: The interface is unknown.
[1463914380] SERVICE ALERT: ADWU1;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1463914381] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;2;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1463914458] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 220 found since last run. oneLevelSearchOps OK. searchOps is too high! 224 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463914501] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;HARD;3;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1463914578] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463914769] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 210.55 Gb (96%) - free 9.35 Gb (4%)
[1463914803] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 410 found since last run. oneLevelSearchOps OK. searchOps is too high! 413 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 774 found since last run.
[1463914923] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463915178] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 284 found since last run. oneLevelSearchOps OK. searchOps is too high! 341 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463915298] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463915423] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463915523] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 300 found since last run. oneLevelSearchOps OK. searchOps is too high! 309 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 747 found since last run.
[1463915543] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1463915643] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463915663] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 79% (15 min average)
[1463915898] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 341 found since last run. oneLevelSearchOps OK. searchOps is too high! 348 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463916018] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463916243] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 448 found since last run. oneLevelSearchOps OK. searchOps is too high! 450 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 774 found since last run.
[1463916263] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463916363] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463916383] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463916503] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463916569] SERVICE ALERT: BOEAPP;C:\ Drive Space;OK;HARD;3;c:\ - total: 219.90 Gb - used: 189.78 Gb (86%) - free 30.12 Gb (14%)
[1463916618] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 275 found since last run. oneLevelSearchOps OK. searchOps is too high! 290 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463916738] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463916963] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 107 found since last run. wholeSubtreeSearchOps is too high! 607 found since last run. oneLevelSearchOps OK. searchOps is too high! 613 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 820 found since last run.
[1463917083] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463917103] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1463917338] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 266 found since last run. oneLevelSearchOps OK. searchOps is too high! 285 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463917459] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463917683] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 475 found since last run. oneLevelSearchOps OK. searchOps is too high! 480 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 759 found since last run.
[1463917703] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463917803] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463917823] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463917943] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1463918058] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 235 found since last run. oneLevelSearchOps OK. searchOps is too high! 239 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463918178] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463918369] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 219.90 Gb - used: 198.30 Gb (90%) - free 21.60 Gb (10%)
[1463918403] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 760 found since last run. oneLevelSearchOps OK. searchOps is too high! 764 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 763 found since last run.
[1463918489] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 219.90 Gb - used: 199.12 Gb (91%) - free 20.78 Gb (9%)
[1463918523] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463918609] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 199.82 Gb (91%) - free 20.08 Gb (9%)
[1463918778] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 200 found since last run. oneLevelSearchOps OK. searchOps is too high! 217 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463918898] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463919123] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 553 found since last run. oneLevelSearchOps OK. searchOps is too high! 558 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 825 found since last run.
[1463919223] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 8
[1463919253] SERVICE ALERT: idm3.emich.edu;LDAP Stats;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1463919287] SERVICE ALERT: idm3.emich.edu;eDirectory Service;CRITICAL;SOFT;1;CRITICAL: Failed to run the ndsd status command successfully!
[1463919343] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 13
[1463919373] SERVICE ALERT: idm3.emich.edu;LDAP Stats;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1463919403] HOST ALERT: idm3.emich.edu;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1463919416] SERVICE ALERT: idm3.emich.edu;eDirectory Service;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1463919446] HOST ALERT: idm3.emich.edu;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1463919463] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 13
[1463919498] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 308 found since last run. oneLevelSearchOps OK. searchOps is too high! 361 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463919510] HOST ALERT: idm3.emich.edu;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1463919526] SERVICE ALERT: idm3.emich.edu;eDirectory Service;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1463919554] SERVICE ALERT: idm2.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 13
[1463919619] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463919651] SERVICE ALERT: idm1.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 13
[1463919654] SERVICE ALERT: idm3.emich.edu;NSDS Errors;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 60 seconds.
[1463919673] SERVICE ALERT: idm2.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 12
[1463919743] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 56% (15 min average)
[1463919771] SERVICE ALERT: idm1.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 7
[1463919774] SERVICE ALERT: idm3.emich.edu;NSDS Errors;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 60 seconds.
[1463919793] SERVICE ALERT: idm2.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 7
[1463919809] SERVICE ALERT: BOEAPP;C:\ Drive Space;OK;HARD;3;c:\ - total: 219.90 Gb - used: 189.95 Gb (86%) - free 29.95 Gb (14%)
[1463919844] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 6
[1463919891] SERVICE ALERT: idm1.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 6
[1463919963] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 296 found since last run. oneLevelSearchOps OK. searchOps is too high! 298 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 824 found since last run.
[1463920126] SERVICE ALERT: idm3.emich.edu;eDirectory Service;OK;HARD;3;OK: ndsd status is good.
[1463920218] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 254 found since last run. oneLevelSearchOps OK. searchOps is too high! 322 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463920338] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463920938] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 242 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463921058] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463921263] SERVICE ALERT: idm4.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1463921543] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463921658] SERVICE ALERT: idm4.emich.edu;LDAP Stats;CRITICAL;SOFT;1;CRITICAL: Failed to connect to LDAP!
[1463921662] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 90% (15 min average)
[1463921701] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;OK;HARD;3;OK: No entries found
[1463921778] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463921863] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 2
[1463921983] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 2
[1463922103] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 1
[1463922209] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 219.90 Gb - used: 198.48 Gb (90%) - free 21.42 Gb (10%)
[1463922261] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463922310] HOST DOWNTIME ALERT: idm1.emich.edu;STARTED; Host has entered a period of scheduled downtime
[1463922310] HOST DOWNTIME ALERT: idm2.emich.edu;STARTED; Host has entered a period of scheduled downtime
[1463922310] HOST DOWNTIME ALERT: idm3.emich.edu;STARTED; Host has entered a period of scheduled downtime
[1463922310] HOST DOWNTIME ALERT: idm4.emich.edu;STARTED; Host has entered a period of scheduled downtime
[1463922310] HOST DOWNTIME ALERT: idmrbs.emich.edu;STARTED; Host has entered a period of scheduled downtime
[1463922310] HOST DOWNTIME ALERT: idmrbstest.emich.edu;STARTED; Host has entered a period of scheduled downtime
[1463922329] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 219.90 Gb - used: 199.27 Gb (91%) - free 20.63 Gb (9%)
[1463922381] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1463922448] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 200.06 Gb (91%) - free 19.84 Gb (9%)
[1463922501] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 86% (15 min average)
[1463922793] SERVICE ALERT: idm2.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1463922834] SERVICE ALERT: idm3.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1463922891] SERVICE ALERT: idm1.emich.edu;NSDS Errors;OK;HARD;3;OK: 0 Errors Reported
[1463923101] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463923221] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463923341] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463923393] SERVICE ALERT: idm2.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 1
[1463923434] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 1
[1463923491] SERVICE ALERT: idm1.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 1
[1463923513] SERVICE ALERT: idm2.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 1
[1463923554] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 1
[1463923611] SERVICE ALERT: idm1.emich.edu;NSDS Errors;OK;SOFT;2;OK: 0 Errors Reported
[1463923633] SERVICE ALERT: idm2.emich.edu;NSDS Errors;OK;SOFT;3;OK: Total errors: 0
[1463923674] SERVICE ALERT: idm3.emich.edu;NSDS Errors;OK;SOFT;3;OK: Total errors: 0
[1463923823] SERVICE ALERT: idm1.emich.edu;OfficeAD Driver;OK;HARD;3;OK: cn=OfficeAD,cn=driverset1,o=system is running.
[1463923903] SERVICE ALERT: idm4.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1463923903] SERVICE FLAPPING ALERT: idm4.emich.edu;NSDS Errors;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1463924247] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 211.67 Gb (96%) - free 8.23 Gb (4%)
[1463925141] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 40% (15 min average)
[1463925549] SERVICE ALERT: idm1.emich.edu;AD Driver;OK;HARD;3;OK: cn=Active Directory,cn=driverset1,o=system is running.
[1463925741] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1463925861] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1463925981] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463926047] SERVICE ALERT: BOEAPP;C:\ Drive Space;OK;HARD;3;c:\ - total: 219.90 Gb - used: 190.07 Gb (86%) - free 29.83 Gb (14%)
[1463926150] SERVICE ALERT: idm1.emich.edu;AD Driver;CRITICAL;SOFT;1;CRITICAL: cn=Active Directory,cn=driverset1,o=system is stopped!
[1463926269] SERVICE ALERT: idm1.emich.edu;AD Driver;CRITICAL;SOFT;2;Unknown: Error reported?
[1463926299] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;CRITICAL;SOFT;1;Unknown: Error reported?
[1463926303] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 13
[1463926336] SERVICE ALERT: idm1.emich.edu;eDirectory Service;CRITICAL;SOFT;1;CRITICAL: Failed to run the ndsd status command successfully!
[1463926390] SERVICE ALERT: idm1.emich.edu;AD Driver;CRITICAL;HARD;3;Unknown: Error reported?
[1463926395] SERVICE ALERT: idm1.emich.edu;Google Driver;CRITICAL;SOFT;1;Unknown: Error reported?
[1463926418] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;CRITICAL;SOFT;2;Unknown: Error reported?
[1463926423] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 13
[1463926456] SERVICE ALERT: idm1.emich.edu;eDirectory Service;OK;SOFT;2;OK: ndsd status is good.
[1463926509] SERVICE ALERT: idm1.emich.edu;Loopback Driver;CRITICAL;SOFT;1;Unknown: Error reported?
[1463926516] SERVICE ALERT: idm1.emich.edu;Google Driver;CRITICAL;SOFT;2;Unknown: Error reported?
[1463926538] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;CRITICAL;HARD;3;Unknown: Error reported?
[1463926543] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 13
[1463926554] SERVICE ALERT: idm1.emich.edu;LDAP Stats;CRITICAL;SOFT;1;CRITICAL: Failed to connect to LDAP!
[1463926630] SERVICE ALERT: idm1.emich.edu;Loopback Driver;CRITICAL;SOFT;2;Unknown: Error reported?
[1463926633] SERVICE ALERT: idm2.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 13
[1463926635] SERVICE ALERT: idm1.emich.edu;Google Driver;CRITICAL;HARD;3;Unknown: Error reported?
[1463926671] SERVICE ALERT: idm1.emich.edu;NSDS Errors;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 60 seconds.
[1463926674] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463926674] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 1
[1463926749] SERVICE ALERT: idm1.emich.edu;Loopback Driver;CRITICAL;HARD;3;CRITICAL: cn=Generic Loopback,cn=driverset1,o=system is stopped!
[1463926753] SERVICE ALERT: idm2.emich.edu;NSDS Errors;OK;SOFT;2;OK: Total errors: 0
[1463926791] SERVICE ALERT: idm1.emich.edu;NSDS Errors;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 60 seconds.
[1463926794] SERVICE ALERT: idm3.emich.edu;NSDS Errors;OK;SOFT;2;OK: Total errors: 0
[1463926823] SERVICE ALERT: idm1.emich.edu;OfficeAD Driver;CRITICAL;SOFT;1;CRITICAL: cn=OfficeAD,cn=driverset1,o=system is stopped!
[1463926851] SERVICE ALERT: idm1.emich.edu;NSDS Errors;OK;SOFT;3;OK: 0 Errors Reported
[1463926944] SERVICE ALERT: idm1.emich.edu;OfficeAD Driver;OK;SOFT;2;OK: cn=OfficeAD,cn=driverset1,o=system is running.
[1463927138] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;OK;HARD;3;OK: cn=ADAUTH Driver,cn=driverset1,o=system is running.
[1463927143] SERVICE ALERT: idm4.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1463927236] SERVICE ALERT: idm1.emich.edu;Google Driver;OK;HARD;3;OK: cn=gdap,cn=driverset1,o=system is running.
[1463927350] SERVICE ALERT: idm1.emich.edu;Loopback Driver;OK;HARD;3;OK: cn=Generic Loopback,cn=driverset1,o=system is running.
[1463927847] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 219.90 Gb - used: 198.99 Gb (90%) - free 20.91 Gb (10%)
[1463927967] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 219.90 Gb - used: 200.71 Gb (91%) - free 19.19 Gb (9%)
[1463928087] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 200.71 Gb (91%) - free 19.19 Gb (9%)
[1463928378] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 222 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463928498] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463928981] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 77% (15 min average)
[1463929098] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 286 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463929218] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463929581] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463929701] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1463929821] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1463929887] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 212.13 Gb (96%) - free 7.77 Gb (4%)
[1463931687] SERVICE ALERT: BOEAPP;C:\ Drive Space;OK;HARD;3;c:\ - total: 219.90 Gb - used: 190.24 Gb (87%) - free 29.66 Gb (13%)
[1463932818] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 268 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463932821] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463932938] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463933421] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463933487] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 219.90 Gb - used: 206.76 Gb (94%) - free 13.14 Gb (6%)
[1463933538] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 200 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463933607] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 219.90 Gb - used: 206.76 Gb (94%) - free 13.14 Gb (6%)
[1463933658] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463933727] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 206.76 Gb (94%) - free 13.14 Gb (6%)
[1463934021] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 86% (15 min average)
[1463934188] SERVICE ALERT: idm1.emich.edu;AD Driver;OK;HARD;3;OK: cn=Campus AD,cn=driverset1,o=system is running.
[1463934258] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 256 found since last run. oneLevelSearchOps OK. searchOps is too high! 376 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463934378] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463935221] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463935341] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463935461] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463935578] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 105 found since last run. wholeSubtreeSearchOps is too high! 533 found since last run. oneLevelSearchOps OK. searchOps is too high! 619 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 758 found since last run.
[1463935698] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 207 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463935818] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463936418] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 337 found since last run. oneLevelSearchOps OK. searchOps is too high! 429 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463936538] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463936699] HOST DOWNTIME ALERT: idm1.emich.edu;STOPPED; Host has exited from a period of scheduled downtime
[1463936699] HOST DOWNTIME ALERT: idm3.emich.edu;STOPPED; Host has exited from a period of scheduled downtime
[1463936699] HOST DOWNTIME ALERT: idm4.emich.edu;STOPPED; Host has exited from a period of scheduled downtime
[1463936699] HOST DOWNTIME ALERT: idm2.emich.edu;STOPPED; Host has exited from a period of scheduled downtime
[1463936699] HOST DOWNTIME ALERT: idmrbs.emich.edu;STOPPED; Host has exited from a period of scheduled downtime
[1463936699] HOST DOWNTIME ALERT: idmrbstest.emich.edu;STOPPED; Host has exited from a period of scheduled downtime
[1463936727] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 212.62 Gb (97%) - free 7.28 Gb (3%)
[1463937138] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 359 found since last run. oneLevelSearchOps OK. searchOps is too high! 496 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463937258] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463937261] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1463937858] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 289 found since last run. oneLevelSearchOps OK. searchOps is too high! 300 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463937861] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1463937927] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 198.25 Gb (90%) - free 21.65 Gb (10%)
[1463937978] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463937981] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 66% (15 min average)
[1463938543] SERVICE FLAPPING ALERT: idm4.emich.edu;NSDS Errors;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1463938578] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 360 found since last run. oneLevelSearchOps OK. searchOps is too high! 434 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463938581] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1463938698] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463938701] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463938821] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1463939298] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 526 found since last run. oneLevelSearchOps OK. searchOps is too high! 641 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463939418] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463940018] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 256 found since last run. oneLevelSearchOps OK. searchOps is too high! 388 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 510 found since last run.
[1463940039] SERVICE ALERT: ADRODC-DET;AD Replication;WARNING;SOFT;1;WARNING: Replication with ADDC5 failed!
[1463940138] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463940163] SERVICE ALERT: ADRODC-DET;AD Replication;OK;SOFT;2;OK: replication is up and running
[1463940175] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;SOFT;1;WARNING: Replication with ADDC5-OLD failed!
[1463940299] SERVICE ALERT: ADRODC-LIV;AD Replication;OK;SOFT;2;OK: replication is up and running
[1463940738] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 267 found since last run. oneLevelSearchOps OK. searchOps is too high! 442 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463940858] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463941221] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 66% (15 min average)
[1463941458] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 311 found since last run. oneLevelSearchOps OK. searchOps is too high! 437 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463941578] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463941821] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1463941941] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1463942061] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463942178] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 235 found since last run. oneLevelSearchOps OK. searchOps is too high! 453 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463942298] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463942727] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 210.29 Gb (96%) - free 9.61 Gb (4%)
[1463942898] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 316 found since last run. oneLevelSearchOps OK. searchOps is too high! 507 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463943018] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463943618] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 279 found since last run. oneLevelSearchOps OK. searchOps is too high! 418 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 532 found since last run.
[1463943738] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463944338] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 210 found since last run. oneLevelSearchOps OK. searchOps is too high! 298 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463944458] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463944461] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 31% (15 min average)
[1463944527] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 198.45 Gb (90%) - free 21.45 Gb (10%)
[1463945058] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 281 found since last run. oneLevelSearchOps OK. searchOps is too high! 341 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 557 found since last run.
[1463945178] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463945661] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463945778] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 432 found since last run. oneLevelSearchOps OK. searchOps is too high! 595 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 576 found since last run.
[1463945781] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463945898] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463945901] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463946498] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 265 found since last run. oneLevelSearchOps OK. searchOps is too high! 304 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463946618] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463946926] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 210.68 Gb (96%) - free 9.22 Gb (4%)
[1463947217] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 251 found since last run. oneLevelSearchOps OK. searchOps is too high! 304 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 507 found since last run.
[1463947336] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463947936] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 323 found since last run. oneLevelSearchOps OK. searchOps is too high! 576 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 549 found since last run.
[1463948055] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463948301] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1463948655] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 482 found since last run. oneLevelSearchOps OK. searchOps is too high! 496 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 531 found since last run.
[1463948726] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 198.51 Gb (90%) - free 21.39 Gb (10%)
[1463948775] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463948901] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463949021] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1463949141] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463949375] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 245 found since last run. oneLevelSearchOps OK. searchOps is too high! 294 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 522 found since last run.
[1463949495] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463949741] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 90% (15 min average)
[1463950095] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 287 found since last run. oneLevelSearchOps OK. searchOps is too high! 289 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 512 found since last run.
[1463950215] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463950341] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463950461] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1463950581] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1463950815] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 377 found since last run. oneLevelSearchOps OK. searchOps is too high! 430 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 532 found since last run.
[1463950935] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463951535] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 511 found since last run. oneLevelSearchOps OK. searchOps is too high! 625 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 575 found since last run.
[1463951655] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463952255] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 564 found since last run. oneLevelSearchOps OK. searchOps is too high! 571 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 552 found since last run.
[1463952375] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463952926] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 210.38 Gb (96%) - free 9.52 Gb (4%)
[1463952975] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 513 found since last run. oneLevelSearchOps OK. searchOps is too high! 648 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 565 found since last run.
[1463953095] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463953695] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 524 found since last run. oneLevelSearchOps OK. searchOps is too high! 715 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 598 found since last run.
[1463953815] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 240 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463953935] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463954535] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 412 found since last run. oneLevelSearchOps OK. searchOps is too high! 414 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 518 found since last run.
[1463954655] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463954726] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 202.62 Gb (92%) - free 17.28 Gb (8%)
[1463955255] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 469 found since last run. oneLevelSearchOps OK. searchOps is too high! 509 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463955375] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463956526] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 210.80 Gb (96%) - free 9.10 Gb (4%)
[1463957775] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 239 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463957895] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463958326] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 202.68 Gb (92%) - free 17.22 Gb (8%)
[1463959095] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 208 found since last run. oneLevelSearchOps OK. searchOps is too high! 254 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463959215] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463959526] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 214.34 Gb (97%) - free 5.56 Gb (3%)
[1463959815] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 395 found since last run. oneLevelSearchOps OK. searchOps is too high! 502 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463959935] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463960126] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 202.72 Gb (92%) - free 17.18 Gb (8%)
[1463960535] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 306 found since last run. oneLevelSearchOps OK. searchOps is too high! 503 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463960655] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463960781] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463961255] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 318 found since last run. oneLevelSearchOps OK. searchOps is too high! 428 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463961375] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463961926] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 212.72 Gb (97%) - free 7.18 Gb (3%)
[1463961975] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 276 found since last run. oneLevelSearchOps OK. searchOps is too high! 390 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463961981] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463962095] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463962105] SERVICE ALERT: PNCDB-2012;D:\ Drive Space;CRITICAL;HARD;3;d:\ - total: 450.00 Gb - used: 424.75 Gb (94%) - free 25.25 Gb (6%)
[1463962695] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 306 found since last run. oneLevelSearchOps OK. searchOps is too high! 461 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463962815] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463963126] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 202.77 Gb (92%) - free 17.13 Gb (8%)
[1463963415] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 232 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463963535] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463963781] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 87% (15 min average)
[1463964135] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 340 found since last run. oneLevelSearchOps OK. searchOps is too high! 391 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463964255] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463964381] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463964501] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1463964621] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463964855] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 433 found since last run. oneLevelSearchOps OK. searchOps is too high! 561 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463964975] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463965221] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463965526] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 215.76 Gb (98%) - free 4.14 Gb (2%)
[1463965575] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 367 found since last run. oneLevelSearchOps OK. searchOps is too high! 454 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463965695] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463965821] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463966295] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 516 found since last run. oneLevelSearchOps OK. searchOps is too high! 665 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463966415] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463966421] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463966726] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 206.07 Gb (94%) - free 13.83 Gb (6%)
[1463967015] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 470 found since last run. oneLevelSearchOps OK. searchOps is too high! 524 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 500 found since last run.
[1463967021] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1463967135] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463967735] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 447 found since last run. oneLevelSearchOps OK. searchOps is too high! 639 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463967855] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463968455] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 722 found since last run. oneLevelSearchOps OK. searchOps is too high! 765 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 565 found since last run.
[1463968526] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 212.20 Gb (96%) - free 7.70 Gb (4%)
[1463968575] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463969175] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 363 found since last run. oneLevelSearchOps OK. searchOps is too high! 495 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 524 found since last run.
[1463969295] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463969895] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 380 found since last run. oneLevelSearchOps OK. searchOps is too high! 529 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463970015] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463970325] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 202.87 Gb (92%) - free 17.03 Gb (8%)
[1463970615] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 492 found since last run. oneLevelSearchOps OK. searchOps is too high! 536 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 527 found since last run.
[1463970735] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463971335] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 717 found since last run. oneLevelSearchOps OK. searchOps is too high! 799 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 549 found since last run.
[1463971455] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463971821] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (15 min average)
[1463972055] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 592 found since last run. oneLevelSearchOps OK. searchOps is too high! 822 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 501 found since last run.
[1463972175] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463972421] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1463972541] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1463972661] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463972725] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 210.21 Gb (96%) - free 9.69 Gb (4%)
[1463972775] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 394 found since last run. oneLevelSearchOps OK. searchOps is too high! 436 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463972895] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463973495] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 267 found since last run. oneLevelSearchOps OK. searchOps is too high! 337 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463973615] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463974215] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 239 found since last run. oneLevelSearchOps OK. searchOps is too high! 304 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463974335] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463974461] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 65% (15 min average)
[1463974525] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 202.87 Gb (92%) - free 17.02 Gb (8%)
[1463974935] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 239 found since last run. oneLevelSearchOps OK. searchOps is too high! 258 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463975055] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463975061] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1463975181] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463975301] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463975655] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 224 found since last run. oneLevelSearchOps OK. searchOps is too high! 240 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463975775] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463975901] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)