[1481519216] SERVICE ALERT: epprogsql2;C:\ Drive Space;OK;HARD;3;c:\ - total: 249.66 Gb - used: 217.82 Gb (87%) - free 31.83 Gb (13%)
[1481519669] SERVICE ALERT: eateriespk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481519729] SERVICE ALERT: eateriespk;Device Status;OK;SOFT;2;MENUS  RETRIEVE JOB  INFORMATION  PAPER HANDLING
[1481520330] SERVICE ALERT: eateriespk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481520389] SERVICE ALERT: eateriespk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481520449] SERVICE ALERT: eateriespk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1481521049] SERVICE ALERT: eateriespk;Device Status;OK;HARD;3;Clearing paper path
[1481521649] SERVICE ALERT: eateriespk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481521709] SERVICE ALERT: eateriespk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481521769] SERVICE ALERT: eateriespk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1481522370] SERVICE ALERT: eateriespk;Device Status;OK;HARD;3;Open front door.
[1481522370] SERVICE FLAPPING ALERT: eateriespk;Device Status;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1481522969] SERVICE ALERT: eateriespk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481523029] SERVICE ALERT: eateriespk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481523084] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 85% (15 min average)
[1481523089] SERVICE ALERT: eateriespk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1481523463] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1481523583] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1481523684] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1481523689] SERVICE ALERT: eateriespk;Device Status;OK;HARD;3;DUPLEXER ERROR REMOVE DUPLEXER Install duplexer with power off
[1481523706] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1481523804] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1481523924] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1481524200] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1481524281] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1481524341] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1481524401] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1481524524] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 80% (15 min average)
[1481524903] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1481525124] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1481525244] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 91% (15 min average)
[1481525503] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1481525844] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1481525964] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1481526084] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1481526103] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1481526223] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1481526343] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1481526343] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1481526843] SERVICE ALERT: FITNESSGRAM;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1481526963] SERVICE ALERT: FITNESSGRAM;CPU Load;OK;SOFT;2;CPU Load 6% (15 min average)
[1481527543] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1481528143] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1481529343] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1481529463] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1481529584] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1481529684] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 84% (15 min average)
[1481530183] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1481530284] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1481530404] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1481530524] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 94% (15 min average)
[1481530783] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1481531124] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1481531244] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1481531364] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1481532260] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1481532380] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1481532500] SERVICE ALERT: epprogsql2;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1481532583] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1481533100] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1481533183] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1481533764] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1481534900] SERVICE ALERT: epprogsql2;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1481535002] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (33.1% change >= 20.0% threshold)
[1481535002] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (53.4% change >= 20.0% threshold)
[1481535002] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1481535002] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1481535002] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1481535002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1481535002] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1481535002] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1481535002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1481535002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1481535002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1481535002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1481535002] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1481535002] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1481535002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1481535002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1481535002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1481535002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1481535002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1481535500] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1481535500] SERVICE FLAPPING ALERT: epprogsql2;CPU Load;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1481536100] SERVICE ALERT: epprogsql2;CPU Load;OK;HARD;3;CPU Load 73% (15 min average)
[1481536802] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1481536802] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (45.5% change >= 20.0% threshold)
[1481536802] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1481536802] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1481536802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1481536802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1481536802] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1481536802] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1481536802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1481536802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1481536802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1481536802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1481536802] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1481536802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1481536802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1481536802] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1481536802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1481536802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1481536802] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1481540964] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1481541084] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 73% (15 min average)
[1481541684] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1481541804] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1481541924] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1481542200] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service has exited from a period of scheduled downtime
[1481542524] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1481543124] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 88% (15 min average)
[1481543724] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1481543844] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1481543964] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 93% (15 min average)
[1481544282] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1012 found since last run.
[1481544402] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481544583] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1481544607] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1558 found since last run.
[1481544727] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481545002] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1651 found since last run.
[1481545053] SERVICE ALERT: prayharroldpk;Device Status;OK;HARD;3;PERFORM PRINTER MAINTENANCE 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1481545122] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481545164] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1481545284] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1481545327] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1685 found since last run.
[1481545404] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1481545447] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481545722] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1789 found since last run.
[1481545842] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481546047] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1778 found since last run.
[1481546167] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481546442] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1797 found since last run.
[1481546562] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481546767] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1773 found since last run.
[1481546887] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481547162] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1853 found since last run.
[1481547282] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481547422] SERVICE ALERT: rooseveltpk;Device Status;OK;HARD;3;Warming up
[1481547426] HOST ALERT: rooseveltpk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.54 ms
[1481547463] SERVICE ALERT: rooseveltpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 8% - OK!
[1481547487] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1818 found since last run.
[1481547607] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481547804] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1481547859] SERVICE ALERT: rooseveltpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1481547882] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1265 found since last run.
[1481548002] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481548207] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5432 found since last run. oneLevelSearchOps OK. searchOps is too high! 5565 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481548327] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481548404] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1481550204] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 49% (15 min average)
[1481550727] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5017 found since last run. oneLevelSearchOps OK. searchOps is too high! 5207 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481550847] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481551447] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5754 found since last run. oneLevelSearchOps OK. searchOps is too high! 5988 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481551567] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481551661] SERVICE ALERT: sillpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1481551720] SERVICE ALERT: sillpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1481551781] SERVICE ALERT: sillpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1481552004] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1481552124] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 90% (15 min average)
[1481552167] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5797 found since last run. oneLevelSearchOps OK. searchOps is too high! 5998 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481552202] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6101 found since last run. oneLevelSearchOps OK. searchOps is too high! 6435 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481552252] SERVICE ALERT: prayharroldpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481552287] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481552312] SERVICE ALERT: prayharroldpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481552322] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481552372] SERVICE ALERT: prayharroldpk;Device Status;OK;SOFT;3;Please open tray 3
[1481552887] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9808 found since last run. oneLevelSearchOps OK. searchOps is too high! 10009 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481552922] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5266 found since last run. oneLevelSearchOps OK. searchOps is too high! 5658 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481552950] HOST ALERT: prayharroldpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1481552976] HOST ALERT: prayharroldpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.57 ms
[1481553003] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (33.4% change >= 20.0% threshold)
[1481553003] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1481553003] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1481553003] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1481553003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1481553003] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1481553003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1481553003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1481553003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1481553003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1481553003] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1481553003] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1481553003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1481553003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1481553003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1481553003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1481553003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1481553007] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481553042] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481553324] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1481553444] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1481553564] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1481553580] SERVICE ALERT: sillpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1481553642] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6726 found since last run. oneLevelSearchOps OK. searchOps is too high! 6985 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481553762] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481554164] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (15 min average)
[1481554207] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 335 found since last run. wholeSubtreeSearchOps is too high! 6320 found since last run. oneLevelSearchOps OK. searchOps is too high! 6502 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481554327] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481554362] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 487 found since last run. wholeSubtreeSearchOps is too high! 6577 found since last run. oneLevelSearchOps OK. searchOps is too high! 7165 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 1145 found since last run.
[1481554482] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481554927] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 329 found since last run. wholeSubtreeSearchOps is too high! 5133 found since last run. oneLevelSearchOps OK. searchOps is too high! 5364 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481555047] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481555082] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 285 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481555202] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481556247] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6617 found since last run. oneLevelSearchOps OK. searchOps is too high! 6817 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481556367] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481556967] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 11100 found since last run. oneLevelSearchOps OK. searchOps is too high! 11298 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481557002] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5510 found since last run. oneLevelSearchOps OK. searchOps is too high! 5976 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481557087] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481557122] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481557131] SERVICE ALERT: halleground1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481557191] SERVICE ALERT: halleground1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481557251] SERVICE ALERT: halleground1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1481557687] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8838 found since last run. oneLevelSearchOps OK. searchOps is too high! 9027 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481557764] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1481557807] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481557884] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 88% (15 min average)
[1481558322] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10431 found since last run. oneLevelSearchOps OK. searchOps is too high! 10795 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481558407] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9163 found since last run. oneLevelSearchOps OK. searchOps is too high! 9264 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481558442] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481558484] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1481558527] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481558604] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1481558724] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1481559042] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8353 found since last run. oneLevelSearchOps OK. searchOps is too high! 8952 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481559127] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8868 found since last run. oneLevelSearchOps OK. searchOps is too high! 8954 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481559162] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481559247] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481559324] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1481559762] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5171 found since last run. oneLevelSearchOps OK. searchOps is too high! 5552 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481559882] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481559923] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1481560447] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7322 found since last run. oneLevelSearchOps OK. searchOps is too high! 7400 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481560482] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5957 found since last run. oneLevelSearchOps OK. searchOps is too high! 6336 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481560523] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1481560567] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481560602] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481560689] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481560748] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481560808] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1481561123] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 50% (15 min average)
[1481561167] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6809 found since last run. oneLevelSearchOps OK. searchOps is too high! 6891 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481561202] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6371 found since last run. oneLevelSearchOps OK. searchOps is too high! 6795 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481561287] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481561322] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481561723] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1481561843] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1481561922] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 11257 found since last run. oneLevelSearchOps OK. searchOps is too high! 11719 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481561963] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1481562042] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481562173] SERVICE ALERT: CAVA2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1481562293] SERVICE ALERT: CAVA2;CPU Load;OK;SOFT;2;CPU Load 74% (15 min average)
[1481562487] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9089 found since last run. oneLevelSearchOps OK. searchOps is too high! 9166 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481562607] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481562642] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9509 found since last run. oneLevelSearchOps OK. searchOps is too high! 10026 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481562762] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481563207] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 11219 found since last run. oneLevelSearchOps OK. searchOps is too high! 11285 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481563208] SERVICE ALERT: markjeffersonpk;Device Status;OK;HARD;3;Processing job from tray 4
[1481563327] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481563357] SERVICE ALERT: owen101lab2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481563362] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 13704 found since last run. oneLevelSearchOps OK. searchOps is too high! 14176 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481563418] SERVICE ALERT: owen101lab2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481563477] SERVICE ALERT: owen101lab2;Device Status;OK;SOFT;3;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1481563482] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481563802] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (41.6% change >= 20.0% threshold)
[1481563803] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1481563803] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1481563803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1481563803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1481563803] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1481563803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1481563803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1481563803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1481563803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1481563803] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1481563803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1481563803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1481563803] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1481563803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1481563803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1481563803] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1481563927] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 13295 found since last run. oneLevelSearchOps OK. searchOps is too high! 13349 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481564046] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481564063] SERVICE ALERT: cava2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1481564082] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8489 found since last run. oneLevelSearchOps OK. searchOps is too high! 8996 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481564183] SERVICE ALERT: cava2;CPU Load;OK;SOFT;2;CPU Load 2% (15 min average)
[1481564201] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481564534] SERVICE ALERT: alexanderpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481564594] SERVICE ALERT: alexanderpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481564645] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8174 found since last run. oneLevelSearchOps OK. searchOps is too high! 8251 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481564654] SERVICE ALERT: alexanderpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1481564764] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481564801] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5427 found since last run. oneLevelSearchOps OK. searchOps is too high! 6117 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481564921] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481565257] SERVICE ALERT: rooseveltpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1481565318] SERVICE ALERT: rooseveltpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1481565362] SERVICE ALERT: rackhampk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1481565364] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7257 found since last run. oneLevelSearchOps OK. searchOps is too high! 7324 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481565378] SERVICE ALERT: rooseveltpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1481565422] SERVICE ALERT: rackhampk;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! 1 of 2 Trays Still Full
[1481565438] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1481565484] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481565499] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1481565520] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8386 found since last run. oneLevelSearchOps OK. searchOps is too high! 8907 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481565558] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1481565639] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481566084] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7663 found since last run. oneLevelSearchOps OK. searchOps is too high! 7737 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481566204] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481566238] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6032 found since last run. oneLevelSearchOps OK. searchOps is too high! 6669 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481566358] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481566397] SERVICE ALERT: eateriespk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1481566759] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1481566804] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5081 found since last run. oneLevelSearchOps OK. searchOps is too high! 5153 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481566924] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481566958] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10532 found since last run. oneLevelSearchOps OK. searchOps is too high! 11289 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481567078] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481567363] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 58% (15 min average)
[1481567468] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1481567524] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 109 found since last run. wholeSubtreeSearchOps is too high! 6204 found since last run. oneLevelSearchOps OK. searchOps is too high! 6267 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481567528] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1481567588] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1481567644] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481567678] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9778 found since last run. oneLevelSearchOps OK. searchOps is too high! 10337 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481567798] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481567963] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1481568083] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1481568201] SERVICE ALERT: porterpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481568203] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1481568244] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 103 found since last run. wholeSubtreeSearchOps is too high! 11467 found since last run. oneLevelSearchOps OK. searchOps is too high! 11534 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481568260] SERVICE ALERT: porterpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481568320] SERVICE ALERT: porterpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1481568364] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481568398] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9850 found since last run. oneLevelSearchOps OK. searchOps is too high! 10535 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481568518] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481568620] SERVICE ALERT: porterpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1481568680] SERVICE ALERT: porterpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1481568740] SERVICE ALERT: porterpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1481568803] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 85% (15 min average)
[1481568854] SERVICE ALERT: alexanderpk;Device Status;OK;HARD;3;Ready
[1481568920] SERVICE ALERT: porterpk;Device Status;OK;HARD;3;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1481568964] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 103 found since last run. wholeSubtreeSearchOps is too high! 14831 found since last run. oneLevelSearchOps OK. searchOps is too high! 14893 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481569084] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481569118] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 14486 found since last run. oneLevelSearchOps OK. searchOps is too high! 15013 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481569238] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481569253] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481569313] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481569373] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1481569403] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1481569523] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1481569577] SERVICE ALERT: rooseveltpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1481569643] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1481569684] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 11694 found since last run. oneLevelSearchOps OK. searchOps is too high! 11743 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481569804] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481569838] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8507 found since last run. oneLevelSearchOps OK. searchOps is too high! 8971 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481569941] SERVICE ALERT: porterpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 96% - OK!
[1481569958] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481570243] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 65% (15 min average)
[1481570404] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 11376 found since last run. oneLevelSearchOps OK. searchOps is too high! 11428 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481570524] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481570558] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6282 found since last run. oneLevelSearchOps OK. searchOps is too high! 6765 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481570678] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481570843] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1481570963] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 83% (15 min average)
[1481571124] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9331 found since last run. oneLevelSearchOps OK. searchOps is too high! 9379 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481571244] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481571278] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6136 found since last run. oneLevelSearchOps OK. searchOps is too high! 6388 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481571336] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1481571366] HOST ALERT: hallefirst2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1481571396] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1481571398] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481571426] HOST ALERT: hallefirst2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1481571456] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1481571486] HOST ALERT: hallefirst2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1481571490] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1481571520] HOST ALERT: hallefirst2;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1481571550] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1481571563] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1481571580] HOST ALERT: hallefirst2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1481571610] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1481571640] HOST ALERT: hallefirst2;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1481571683] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 84% (15 min average)
[1481571704] HOST ALERT: hallefirst2;UP;SOFT;7;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1481571773] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Ready
[1481571844] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 11004 found since last run. oneLevelSearchOps OK. searchOps is too high! 11055 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481571964] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481571998] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5144 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481572050] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 21% - OK!
[1481572118] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481572204] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1481572283] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1481572403] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1481572523] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1481572564] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6553 found since last run. oneLevelSearchOps OK. searchOps is too high! 6601 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481572684] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481572718] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10553 found since last run. oneLevelSearchOps OK. searchOps is too high! 10896 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481572838] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481573123] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1481573284] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8116 found since last run. oneLevelSearchOps OK. searchOps is too high! 8165 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481573404] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481573438] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8049 found since last run. oneLevelSearchOps OK. searchOps is too high! 8470 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481573558] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481573723] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1481574004] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8019 found since last run. oneLevelSearchOps OK. searchOps is too high! 8070 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481574124] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481574158] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 13705 found since last run. oneLevelSearchOps OK. searchOps is too high! 14749 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481574278] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481574323] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 81% (15 min average)
[1481574724] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9958 found since last run. oneLevelSearchOps OK. searchOps is too high! 10004 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481574788] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1481574844] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481574878] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6858 found since last run. oneLevelSearchOps OK. searchOps is too high! 7672 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481574923] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1481574998] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481575043] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1481575163] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1481575444] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8438 found since last run. oneLevelSearchOps OK. searchOps is too high! 8489 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481575564] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481575598] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 118 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5330 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481575718] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481575763] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1481576318] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7349 found since last run. oneLevelSearchOps OK. searchOps is too high! 8077 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481576438] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481576764] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5089 found since last run. oneLevelSearchOps OK. searchOps is too high! 5141 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481576884] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481576963] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1481577038] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8604 found since last run. oneLevelSearchOps OK. searchOps is too high! 9314 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481577083] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1481577158] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481577203] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 52% (15 min average)
[1481577484] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7187 found since last run. oneLevelSearchOps OK. searchOps is too high! 7242 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481577604] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481577758] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5627 found since last run. oneLevelSearchOps OK. searchOps is too high! 6048 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481577803] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1481577878] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481577923] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1481578043] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1481578204] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8474 found since last run. oneLevelSearchOps OK. searchOps is too high! 8545 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481578324] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481578345] SERVICE ALERT: hallefirst3;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1481578375] HOST ALERT: hallefirst3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1481578400] SERVICE ALERT: hallefirst3;Device Status;OK;SOFT;2;Clearing paper path
[1481578401] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1481578404] HOST ALERT: hallefirst3;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1481578457] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 36% - OK!
[1481578478] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5611 found since last run. oneLevelSearchOps OK. searchOps is too high! 6122 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481578493] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;OK;HARD;3;Yellow Cartridge HP CE262A is at 100% - OK!
[1481578598] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481578924] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9238 found since last run. oneLevelSearchOps OK. searchOps is too high! 9418 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481579044] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481579198] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6400 found since last run. oneLevelSearchOps OK. searchOps is too high! 7060 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481579243] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1481579318] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481579599] SERVICE ALERT: hallefirst3;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481579644] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 12154 found since last run. oneLevelSearchOps OK. searchOps is too high! 12325 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481579659] SERVICE ALERT: hallefirst3;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481579720] SERVICE ALERT: hallefirst3;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1481579764] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481579843] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1481580364] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7911 found since last run. oneLevelSearchOps OK. searchOps is too high! 8055 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481580484] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481580518] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6380 found since last run. oneLevelSearchOps OK. searchOps is too high! 6688 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481580638] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481581043] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 79% (15 min average)
[1481581084] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10218 found since last run. oneLevelSearchOps OK. searchOps is too high! 10384 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481581203] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481581238] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7177 found since last run. oneLevelSearchOps OK. searchOps is too high! 7542 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481581250] SERVICE ALERT: halleground1;Device Status;OK;HARD;3;Clearing paper path
[1481581358] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481581643] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1481581763] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1481581851] SERVICE ALERT: halleground1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481581883] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 55% (15 min average)
[1481581910] SERVICE ALERT: halleground1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481581959] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10121 found since last run. oneLevelSearchOps OK. searchOps is too high! 10652 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481581970] SERVICE ALERT: halleground1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1481582078] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481582677] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7701 found since last run. oneLevelSearchOps OK. searchOps is too high! 8196 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481582797] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481583003] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6281 found since last run. oneLevelSearchOps OK. searchOps is too high! 6418 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481583083] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1481583123] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481583203] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1481583323] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1481583723] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7046 found since last run. oneLevelSearchOps OK. searchOps is too high! 7194 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481583843] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481583923] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1481583947] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1481583977] HOST ALERT: halleground1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1481583997] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6357 found since last run. oneLevelSearchOps OK. searchOps is too high! 6659 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481584007] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1481584037] HOST ALERT: halleground1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1481584067] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1481584097] HOST ALERT: halleground1;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1481584117] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481584157] SERVICE ALERT: halleground1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1481584187] HOST ALERT: halleground1;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1481584217] SERVICE ALERT: halleground1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1481584247] HOST ALERT: halleground1;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1481584277] SERVICE ALERT: halleground1;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1481584307] HOST ALERT: halleground1;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1481584376] SERVICE ALERT: halleground1;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1481584397] HOST ALERT: halleground1;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1481584443] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7779 found since last run. oneLevelSearchOps OK. searchOps is too high! 7922 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481584461] HOST ALERT: halleground1;UP;SOFT;8;PING OK - Packet loss = 0%, RTA = 0.55 ms
[1481584563] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481584662] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 31% - OK!
[1481584717] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8641 found since last run. oneLevelSearchOps OK. searchOps is too high! 8780 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481584837] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481584871] SERVICE ALERT: halleground1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1481584970] SERVICE ALERT: halleground1;Device Status;OK;HARD;3;Clearing paper path
[1481584970] SERVICE FLAPPING ALERT: halleground1;Device Status;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1481585437] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8293 found since last run. oneLevelSearchOps OK. searchOps is too high! 8494 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481585557] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481585703] HOST ALERT: halleground1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1481585793] HOST ALERT: halleground1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1481585867] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1481585883] HOST ALERT: halleground1;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1481585927] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1481585957] HOST ALERT: halleground1;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1481585987] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1481586017] HOST ALERT: halleground1;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1481586077] SERVICE ALERT: halleground1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1481586107] HOST ALERT: halleground1;DOWN;SOFT;6;(Host check timed out after 30.00 seconds)
[1481586137] SERVICE ALERT: halleground1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1481586167] HOST ALERT: halleground1;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1481586176] SERVICE ALERT: halleground1;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1481586197] SERVICE ALERT: halleground1;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1481586206] HOST ALERT: halleground1;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1481586236] SERVICE ALERT: halleground1;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1481586266] HOST ALERT: halleground1;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1481586296] SERVICE ALERT: halleground1;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1481586326] HOST ALERT: halleground1;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1481586363] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5608 found since last run. oneLevelSearchOps OK. searchOps is too high! 5739 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481586483] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481587083] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7490 found since last run. oneLevelSearchOps OK. searchOps is too high! 7626 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481587203] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481587803] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6699 found since last run. oneLevelSearchOps OK. searchOps is too high! 6834 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481587923] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481588123] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 52% (15 min average)
[1481588523] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9482 found since last run. oneLevelSearchOps OK. searchOps is too high! 9628 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481588643] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481589243] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9434 found since last run. oneLevelSearchOps OK. searchOps is too high! 9575 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481589323] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1481589363] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481589443] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 76% (15 min average)
[1481589963] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6660 found since last run. oneLevelSearchOps OK. searchOps is too high! 6799 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481590043] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1481590083] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481590163] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1481590283] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1481590357] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1469 found since last run.
[1481590363] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1481590423] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 28% - OK!
[1481590477] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481591283] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5026 found since last run. oneLevelSearchOps OK. searchOps is too high! 5152 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481591403] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481591677] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6123 found since last run. oneLevelSearchOps OK. searchOps is too high! 6359 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481591797] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481592397] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5090 found since last run. oneLevelSearchOps OK. searchOps is too high! 5506 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481592517] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481592603] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5490 found since last run. oneLevelSearchOps OK. searchOps is too high! 5542 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481592723] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481593323] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5693 found since last run. oneLevelSearchOps OK. searchOps is too high! 5743 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481593443] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481594483] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1481594643] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5081 found since last run. oneLevelSearchOps OK. searchOps is too high! 5135 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481594763] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481594917] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5359 found since last run. oneLevelSearchOps OK. searchOps is too high! 5672 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481594984] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;SOFT;1;e:\ - total: 799.87 Gb - used: 649.64 Gb (81%) - free 150.24 Gb (19%)
[1481595037] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481595083] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1481595104] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;SOFT;2;e:\ - total: 799.87 Gb - used: 649.64 Gb (81%) - free 150.24 Gb (19%)
[1481595224] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;HARD;3;e:\ - total: 799.87 Gb - used: 649.64 Gb (81%) - free 150.24 Gb (19%)
[1481596563] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6076 found since last run. oneLevelSearchOps OK. searchOps is too high! 6126 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481596683] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481597096] SERVICE FLAPPING ALERT: halleground1;Device Status;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1481597283] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5862 found since last run. oneLevelSearchOps OK. searchOps is too high! 5916 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1481597403] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481597483] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 90% (15 min average)
[1481597624] SERVICE ALERT: AD-RMAD;E:\ Drive Space;OK;HARD;3;e:\ - total: 799.87 Gb - used: 624.70 Gb (78%) - free 175.17 Gb (22%)
[1481597657] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1481597716] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1481597719] SERVICE ALERT: hallefirst3;Device Status;OK;HARD;3;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1481597776] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1481598083] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1481598203] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1481598323] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1481598376] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 36% - OK!
[1481598923] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1481599523] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1481599643] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1481599763] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1481605163] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)