[1487395801] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1487396539] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1487396659] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 42% (15 min average)
[1487398024] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1487398144] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1487398745] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1487398866] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;2;Warning:  idm1  are more than 120 seconds behind!
[1487398984] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1487400185] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1487400259] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1487400379] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1487400499] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1487400784] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1487401278] SERVICE ALERT: INTLDB;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1487401384] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1487401384] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1487401398] SERVICE ALERT: INTLDB;CPU Load;OK;SOFT;2;CPU Load 0% (15 min average)
[1487401823] SERVICE ALERT: AXIOM-2012;C:\ Drive Space;OK;HARD;3;c:\ - total: 39.66 Gb - used: 35.66 Gb (90%) - free 3.99 Gb (10%)
[1487401984] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1487402299] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1487402899] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1487403321] SERVICE ALERT: FITNESSGRAM;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1487403441] SERVICE ALERT: FITNESSGRAM;CPU Load;OK;SOFT;2;CPU Load 3% (15 min average)
[1487404099] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1487404099] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1487404699] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1487404819] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1487405419] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1487405539] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 87% (15 min average)
[1487406139] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1487406259] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1487406379] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 93% (15 min average)
[1487407579] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1487407699] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 86% (15 min average)
[1487407984] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1487408299] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1487408419] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1487408539] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1487408584] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1487408704] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1487408824] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;3;OK. All replicas within 120 seconds.
[1487409139] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 17% (15 min average)
[1487410200] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service has exited from a period of scheduled downtime
[1487410203] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (31.4% change >= 20.0% threshold)
[1487410203] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (27.1% change >= 20.0% threshold)
[1487410203] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1487410203] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1487410203] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1487410203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1487410203] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1487410203] HOST DOWNTIME ALERT: INTLAPP-ISSMAP;STARTED; Host has entered a period of scheduled downtime
[1487410203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1487410203] HOST DOWNTIME ALERT: INTLTESTDB;STARTED; Host has entered a period of scheduled downtime
[1487410203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1487410203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1487410203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1487410203] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1487410203] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1487410203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1487410203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1487410203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1487410203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1487410203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1487410624] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1487410744] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1487412002] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (29.1% change >= 20.0% threshold)
[1487412002] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1487412002] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1487412002] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1487412002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1487412002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1487412002] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1487412002] HOST DOWNTIME ALERT: INTLAPP-ISSMAP;STARTED; Host has entered a period of scheduled downtime
[1487412002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1487412002] HOST DOWNTIME ALERT: INTLTESTDB;STARTED; Host has entered a period of scheduled downtime
[1487412002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1487412002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1487412002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1487412002] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1487412002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1487412002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1487412002] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1487412002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1487412002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1487412002] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1487413144] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1487413263] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1487414463] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1487416339] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1487416459] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1487416579] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1487417179] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1487418379] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1487418979] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1487421979] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 78% (15 min average)
[1487422579] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1487422699] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1487422819] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1487423419] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 90% (15 min average)
[1487426419] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1487426539] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 22% (15 min average)
[1487428203] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (32.2% change >= 20.0% threshold)
[1487428203] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1487428203] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1487428203] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1487428203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1487428203] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1487428203] HOST DOWNTIME ALERT: INTLAPP-ISSMAP;STARTED; Host has entered a period of scheduled downtime
[1487428203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1487428203] HOST DOWNTIME ALERT: INTLTESTDB;STARTED; Host has entered a period of scheduled downtime
[1487428203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1487428203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1487428203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1487428203] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1487428203] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1487428203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1487428203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1487428203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1487428203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1487428203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1487428250] SERVICE ALERT: ADRODC-LIV;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1487428263] SERVICE ALERT: ADRODC-LIV;Service: Key Distribution Center;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1487428264] SERVICE ALERT: ADRODC-LIV;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1487428268] SERVICE ALERT: ADRODC-LIV;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1487428268] SERVICE ALERT: ADRODC-LIV;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1487428280] SERVICE ALERT: ADRODC-LIV;Service: DNS Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1487428280] SERVICE ALERT: ADRODC-LIV;Service: Remote Management;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1487428280] SERVICE ALERT: ADRODC-LIV;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1487428280] HOST ALERT: ADRODC-LIV;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1487428280] SERVICE ALERT: ADRODC-LIV;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1487428283] SERVICE ALERT: ADRODC-LIV;Service: Distributed File System Replication;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1487428303] SERVICE ALERT: ADRODC-LIV;AD Replication;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 30 seconds.
[1487428310] HOST ALERT: ADRODC-LIV;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1487428312] SERVICE ALERT: ADRODC-LIV;Service: Intersite Messaging;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1487428342] HOST ALERT: ADRODC-LIV;DOWN;SOFT;3;(Host check timed out after 30.00 seconds)
[1487428370] SERVICE ALERT: ADRODC-LIV;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1487428383] SERVICE ALERT: ADRODC-LIV;Service: Key Distribution Center;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1487428384] SERVICE ALERT: ADRODC-LIV;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1487428388] SERVICE ALERT: ADRODC-LIV;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1487428388] SERVICE ALERT: ADRODC-LIV;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1487428400] SERVICE ALERT: ADRODC-LIV;Service: DNS Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1487428400] SERVICE ALERT: ADRODC-LIV;Service: Remote Management;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1487428400] SERVICE ALERT: ADRODC-LIV;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1487428400] HOST ALERT: ADRODC-LIV;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1487428400] SERVICE ALERT: ADRODC-LIV;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1487428403] SERVICE ALERT: ADRODC-LIV;Service: Distributed File System Replication;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1487428403] SERVICE ALERT: ADRODC-LIV;Service: AD Web Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1487428411] SERVICE ALERT: ADRODC-LIV;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1487428423] SERVICE ALERT: ADRODC-LIV;AD Replication;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 30 seconds.
[1487428430] HOST ALERT: ADRODC-LIV;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1487428430] SERVICE ALERT: ADRODC-LIV;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1487428432] SERVICE ALERT: ADRODC-LIV;Service: Intersite Messaging;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1487428460] HOST ALERT: ADRODC-LIV;DOWN;SOFT;6;(Host check timed out after 30.00 seconds)
[1487428490] SERVICE ALERT: ADRODC-LIV;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1487428503] SERVICE ALERT: ADRODC-LIV;Service: Key Distribution Center;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1487428504] SERVICE ALERT: ADRODC-LIV;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1487428508] SERVICE ALERT: ADRODC-LIV;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1487428508] SERVICE ALERT: ADRODC-LIV;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1487428520] SERVICE ALERT: ADRODC-LIV;Service: DNS Service;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1487428520] SERVICE ALERT: ADRODC-LIV;Service: Remote Management;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1487428520] SERVICE ALERT: ADRODC-LIV;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1487428520] HOST ALERT: ADRODC-LIV;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1487428520] SERVICE ALERT: ADRODC-LIV;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1487428523] SERVICE ALERT: ADRODC-LIV;Service: Distributed File System Replication;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1487428523] SERVICE ALERT: ADRODC-LIV;Service: AD Web Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1487428531] SERVICE ALERT: ADRODC-LIV;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1487428543] SERVICE ALERT: ADRODC-LIV;AD Replication;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 30 seconds.
[1487428550] HOST ALERT: ADRODC-LIV;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1487428550] SERVICE ALERT: ADRODC-LIV;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1487428552] SERVICE ALERT: ADRODC-LIV;Service: Intersite Messaging;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1487428580] HOST ALERT: ADRODC-LIV;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1487428582] SERVICE ALERT: ADRODC-LIV;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1487428590] SERVICE ALERT: ADRODC-LIV;Service: DFS Namespace;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1487428591] SERVICE ALERT: ADRODC-LIV;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1487428612] HOST ALERT: ADRODC-LIV;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1487428643] SERVICE ALERT: ADRODC-LIV;Service: AD Web Service;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1487428651] SERVICE ALERT: ADRODC-LIV;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1487428670] SERVICE ALERT: ADRODC-LIV;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1487428688] SERVICE ALERT: ADRODC-LIV;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1487428688] SERVICE ALERT: ADRODC-LIV;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1487428688] SERVICE ALERT: ADRODC-LIV;Service: PXE Boot;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1487428691] SERVICE ALERT: ADRODC-LIV;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1487428692] SERVICE ALERT: ADRODC-LIV;Service: MDT Monitor;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1487428702] SERVICE ALERT: ADRODC-LIV;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1487428703] SERVICE ALERT: ADRODC-LIV;Service: AD Domain Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1487428710] SERVICE ALERT: ADRODC-LIV;Service: DFS Namespace;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1487428711] SERVICE ALERT: ADRODC-LIV;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1487428808] SERVICE ALERT: ADRODC-LIV;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1487428808] SERVICE ALERT: ADRODC-LIV;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1487428808] SERVICE ALERT: ADRODC-LIV;Service: PXE Boot;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1487428811] SERVICE ALERT: ADRODC-LIV;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1487428812] SERVICE ALERT: ADRODC-LIV;Service: MDT Monitor;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1487428822] SERVICE ALERT: ADRODC-LIV;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1487428823] SERVICE ALERT: ADRODC-LIV;Service: AD Domain Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1487428830] SERVICE ALERT: ADRODC-LIV;Service: DFS Namespace;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1487428831] SERVICE ALERT: ADRODC-LIV;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1487428928] SERVICE ALERT: ADRODC-LIV;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1487428928] SERVICE ALERT: ADRODC-LIV;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1487428928] SERVICE ALERT: ADRODC-LIV;Service: PXE Boot;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1487428931] SERVICE ALERT: ADRODC-LIV;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1487428932] SERVICE ALERT: ADRODC-LIV;Service: MDT Monitor;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1487428943] SERVICE ALERT: ADRODC-LIV;Service: AD Domain Service;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1487434938] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1487439003] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1487439003] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1487439003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1487439003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1487439003] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1487439003] HOST DOWNTIME ALERT: INTLAPP-ISSMAP;STARTED; Host has entered a period of scheduled downtime
[1487439003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1487439003] HOST DOWNTIME ALERT: INTLTESTDB;STARTED; Host has entered a period of scheduled downtime
[1487439003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1487439003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1487439003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1487439003] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1487439003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1487439003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1487439003] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1487439003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1487439003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1487439003] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1487439138] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1487439258] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1487439378] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 79% (15 min average)
[1487439625] SERVICE ALERT: eateriespk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1487439685] SERVICE ALERT: eateriespk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1487439745] SERVICE ALERT: eateriespk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1487440945] SERVICE ALERT: eateriespk;Device Status;OK;HARD;3;Clearing paper path
[1487444421] SERVICE ALERT: AXIOM-2012;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 39.66 Gb - used: 35.69 Gb (90%) - free 3.96 Gb (10%)
[1487444540] SERVICE ALERT: AXIOM-2012;C:\ Drive Space;OK;SOFT;2;c:\ - total: 39.66 Gb - used: 35.67 Gb (90%) - free 3.98 Gb (10%)
[1487450776] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1487450896] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 44% (15 min average)
[1487454423] SERVICE ALERT: ADRODC-LIV;Service: AD Web Service;OK;HARD;3;ADWS: Started
[1487454427] HOST ALERT: ADRODC-LIV;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 6.10 ms
[1487454429] SERVICE ALERT: ADRODC-LIV;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1487454448] SERVICE ALERT: ADRODC-LIV;Service: RPC;OK;HARD;3;RpcSs: Started
[1487454608] SERVICE ALERT: ADRODC-LIV;Service: DFS Namespace;OK;HARD;3;Dfs: Started
[1487454609] SERVICE ALERT: ADRODC-LIV;Memory Usage;OK;HARD;3;Memory usage: total:37586.86 MB - used: 4759.74 MB (13%) - free: 32827.12 MB (87%)
[1487454688] SERVICE ALERT: ADRODC-LIV;Uptime;OK;HARD;3;System Uptime - 60 day(s) 8 hour(s) 49 minute(s)
[1487454706] SERVICE ALERT: ADRODC-LIV;Service: PXE Boot;OK;HARD;3;WDSServer: Started
[1487454706] SERVICE ALERT: ADRODC-LIV;CPU Load;OK;HARD;3;CPU Load 1% (15 min average)
[1487454706] SERVICE ALERT: ADRODC-LIV;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1487454709] SERVICE ALERT: ADRODC-LIV;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1487454710] SERVICE ALERT: ADRODC-LIV;Service: MDT Monitor;OK;HARD;3;MDT_Monitor: Started
[1487454723] SERVICE ALERT: ADRODC-LIV;Service: AD Domain Service;OK;HARD;3;NTDS: Started
[1487454868] SERVICE ALERT: ADRODC-LIV;Service: Windows Time;OK;HARD;3;W32Time: Started
[1487454882] SERVICE ALERT: ADRODC-LIV;Service: Key Distribution Center;OK;HARD;3;KDC: Started
[1487454883] SERVICE ALERT: ADRODC-LIV;Service: WMI;OK;HARD;3;Winmgmt: Started
[1487454886] SERVICE ALERT: ADRODC-LIV;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1487454886] SERVICE ALERT: ADRODC-LIV;C:\ Drive Space;OK;HARD;3;c:\ - total: 1115.16 Gb - used: 559.99 Gb (50%) - free 555.17 Gb (50%)
[1487454897] SERVICE ALERT: ADRODC-LIV;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1487454898] SERVICE ALERT: ADRODC-LIV;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1487454898] SERVICE ALERT: ADRODC-LIV;Service: Remote Management;OK;HARD;3;WinRM: Started
[1487454898] SERVICE ALERT: ADRODC-LIV;Service: DNS Service;OK;HARD;3;DNS: Started
[1487454900] SERVICE ALERT: ADRODC-LIV;Service: Distributed File System Replication;OK;HARD;3;DFSR: Started
[1487454912] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;HARD;3;WARNING: Replication with ADDC5 failed!
[1487454930] SERVICE ALERT: ADRODC-LIV;Service: Intersite Messaging;OK;HARD;3;IsmServ: Started
[1487460496] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1487460616] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1487460736] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1487461936] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 53% (15 min average)
[1487463316] SERVICE ALERT: ADRODC-LIV;AD Replication;OK;HARD;3;OK: replication is up and running
[1487464936] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1487465056] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 84% (15 min average)
[1487470939] SERVICE ALERT: AXIOM-2012;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 39.66 Gb - used: 35.70 Gb (90%) - free 3.96 Gb (10%)
[1487471059] SERVICE ALERT: AXIOM-2012;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 39.66 Gb - used: 35.70 Gb (90%) - free 3.96 Gb (10%)
[1487471065] SERVICE ALERT: hallegroundcolor;Cyan Cartridge HP CE261A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1487471120] SERVICE ALERT: hallegroundcolor;Cyan Cartridge HP CE261A;OK;SOFT;2;Cyan Cartridge HP CE261A is at 21% - OK!
[1487471179] SERVICE ALERT: AXIOM-2012;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 39.66 Gb - used: 35.70 Gb (90%) - free 3.96 Gb (10%)
[1487471656] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1487471776] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 78% (15 min average)
[1487472375] HOST ALERT: IMAGES;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1487472376] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1487472453] SERVICE ALERT: IMAGES;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1487472465] HOST ALERT: IMAGES;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1487472496] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1487472501] HOST ALERT: IMAGES;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.21 ms
[1487472553] SERVICE ALERT: IMAGES;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1487472616] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1487473579] SERVICE ALERT: AXIOM-2012;C:\ Drive Space;OK;HARD;3;c:\ - total: 39.66 Gb - used: 35.68 Gb (90%) - free 3.98 Gb (10%)
[1487474415] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 57% (15 min average)
[1487475615] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1487475735] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 49% (15 min average)
[1487476335] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1487476455] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 49% (15 min average)
[1487477055] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1487477175] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1487477295] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1487479095] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1487479095] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)