[1461470647] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 76% (15 min average)
[1461471247] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461471367] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461471487] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461473078] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;1;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1461473198] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;2;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1461473286] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461473318] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;HARD;3;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1461473886] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461475086] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461475686] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1461478087] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1461478207] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1461478327] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 94% (15 min average)
[1461479286] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461479886] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461480127] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1461480246] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 87% (15 min average)
[1461480518] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;OK;HARD;3;OK: No entries found
[1461480846] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1461480965] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1461481085] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 90% (15 min average)
[1461481086] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461481686] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461482285] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461482404] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1461482524] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461482844] SERVICE ALERT: NETENG-PA;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1461482886] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461482964] SERVICE ALERT: NETENG-PA;CPU Load;OK;SOFT;2;CPU Load 3% (15 min average)
[1461483123] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 86% (15 min average)
[1461483486] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461483723] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1461483842] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 71% (15 min average)
[1461486280] SERVICE ALERT: NETIDADPROXY;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1461486400] SERVICE ALERT: NETIDADPROXY;Service: Windows Time;CRITICAL;SOFT;2;W32Time: Stopped
[1461486520] SERVICE ALERT: NETIDADPROXY;Service: Windows Time;CRITICAL;HARD;3;W32Time: Stopped
[1461487083] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1461487683] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1461487803] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1461487923] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461488523] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 72% (15 min average)
[1461489123] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461489243] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1461489363] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461491763] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 90% (15 min average)
[1461492363] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461492483] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1461492603] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461493203] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 84% (15 min average)
[1461493803] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1461493923] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461494043] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461494642] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1461495243] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 61% (15 min average)
[1461495518] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;1;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1461495638] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;2;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1461495758] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;HARD;3;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1461495843] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461495963] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461496083] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 72% (15 min average)
[1461496683] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461496803] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1461496923] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1461497523] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461498123] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461498723] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 87% (15 min average)
[1461499323] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461499443] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461499563] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461500163] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 91% (15 min average)
[1461502358] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;OK;HARD;3;OK: No entries found
[1461502563] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461502683] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461502803] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 91% (15 min average)
[1461503403] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461503523] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1461503643] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 74% (15 min average)
[1461504243] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461504363] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1461504963] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1461505083] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1461505203] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1461507603] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 81% (15 min average)
[1461508203] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461508323] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461508443] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461509182] SERVICE ALERT: eateriespk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461509242] SERVICE ALERT: eateriespk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461509302] SERVICE ALERT: eateriespk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461511350] SERVICE ALERT: eateriespk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1461511410] SERVICE ALERT: eateriespk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1461511470] SERVICE ALERT: eateriespk;Black Print Cartridge HP C8543X;OK;SOFT;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1461511702] SERVICE ALERT: eateriespk;Device Status;OK;HARD;3;Clearing paper path
[1461512643] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461513243] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1461513843] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461515174] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461515233] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461515293] SERVICE ALERT: hallefirst1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461516243] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461516321] SERVICE ALERT: studentcenterpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1461516843] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 73% (15 min average)
[1461516882] SERVICE ALERT: hallesecondsouth;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1461516912] HOST ALERT: hallesecondsouth;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1461516928] SERVICE ALERT: hallesecondsouth;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1461516942] SERVICE ALERT: hallesecondsouth;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1461516958] HOST ALERT: hallesecondsouth;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1461516988] SERVICE ALERT: hallesecondsouth;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1461517002] SERVICE ALERT: hallesecondsouth;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1461517018] HOST ALERT: hallesecondsouth;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1461517048] SERVICE ALERT: hallesecondsouth;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1461517078] HOST ALERT: hallesecondsouth;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1461517168] HOST ALERT: hallesecondsouth;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1461517258] HOST ALERT: hallesecondsouth;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1461517281] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1461517311] HOST ALERT: hallesecondsouth;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1461517401] HOST ALERT: hallesecondsouth;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1461517443] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461517491] HOST ALERT: hallesecondsouth;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1461517563] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461517581] HOST ALERT: hallesecondsouth;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1461517683] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1461518293] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;ORDER CARTRIDGE LESS THAN 1700 PAGES 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1461520632] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461520692] SERVICE ALERT: hallefirst2;Device Status;OK;SOFT;2;50 72 65 73 73 20 17 20 74 6F 20 73 65 65 20 73  74 65 70  by step information
[1461521283] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 68% (15 min average)
[1461521883] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461522003] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461522123] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461523278] SERVICE ALERT: PROCARE;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1461523398] SERVICE ALERT: PROCARE;CPU Load;OK;SOFT;2;CPU Load 0% (15 min average)
[1461523923] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 84% (15 min average)
[1461524523] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1461524598] SERVICE ALERT: PROCARE;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461524643] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461524718] SERVICE ALERT: PROCARE;CPU Load;OK;SOFT;2;CPU Load 0% (15 min average)
[1461524763] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 76% (15 min average)
[1461525321] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1461525363] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461525382] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1461525442] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1461525482] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1461525492] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461525552] SERVICE ALERT: hallefirst2;Device Status;OK;SOFT;2;RETRIEVE JOB  NO STORED JOBS
[1461525602] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1461526152] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461526212] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461526272] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461526802] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1461526872] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Processing job from tray 4
[1461527375] SERVICE ALERT: hallethirdsouth;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461527402] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461527436] SERVICE ALERT: hallethirdsouth;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461527495] SERVICE ALERT: hallethirdsouth;Device Status;OK;SOFT;3;Open tray indicated.
[1461527522] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461527642] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461528242] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 88% (15 min average)
[1461528695] SERVICE ALERT: hallethirdsouth;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461528755] SERVICE ALERT: hallethirdsouth;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461528815] SERVICE ALERT: hallethirdsouth;Device Status;OK;SOFT;3;Open tray indicated.
[1461528842] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1461528962] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461529048] SERVICE ALERT: hallegroundcolor;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 1 Trays Empty
[1461529082] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461529108] SERVICE ALERT: hallegroundcolor;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 1 Trays Empty
[1461529168] SERVICE ALERT: hallegroundcolor;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 1 Trays Empty
[1461529682] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461530015] SERVICE ALERT: hallethirdsouth;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461530075] SERVICE ALERT: hallethirdsouth;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461530136] SERVICE ALERT: hallethirdsouth;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461530282] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1461530735] SERVICE ALERT: hallethirdsouth;Device Status;OK;HARD;3;Ready
[1461530841] SERVICE ALERT: studentcenterpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1461530882] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461531002] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1461531602] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461531722] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1461531842] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1461532623] SERVICE ALERT: eateriespk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1461532722] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;WARNING;SOFT;1;Image Transfer Kit HP CE249A is at 2% - WARNING!
[1461532783] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;WARNING;SOFT;2;Image Transfer Kit HP CE249A is at 2% - WARNING!
[1461532842] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;WARNING;HARD;3;Image Transfer Kit HP CE249A is at 2% - WARNING!
[1461533642] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 71% (15 min average)
[1461534242] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461534362] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461534482] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1461537482] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461538082] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461538092] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461538152] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461538213] SERVICE ALERT: hallefirst1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461539412] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;ORDER CARTRIDGE LESS THAN 1000 PAGES 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Processing job from tray 4
[1461539882] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461540482] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461541682] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 74% (15 min average)
[1461542282] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461542402] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461542522] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461545715] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1461545841] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1461545901] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1461545961] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1461546560] SERVICE ALERT: studentcenterpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1461546722] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 91% (15 min average)
[1461547161] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1461547221] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1461547281] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1461547322] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461547442] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1461547562] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 90% (15 min average)
[1461548162] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461548282] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461548402] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461548711] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1461548771] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1461548830] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1461548966] SERVICE ALERT: hallegroundcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1461551910] SERVICE ALERT: INTLTESTAPP;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461552010] SERVICE ALERT: INTLTESTAPP;VMWare Tools;OK;SOFT;2;VMTools: Started
[1461552022] SERVICE ALERT: LANSWEEPER;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461552064] SERVICE ALERT: CBORD;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461552081] SERVICE ALERT: studentcenterpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1461552081] SERVICE FLAPPING ALERT: studentcenterpk;Paper Tray Status;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1461552122] SERVICE ALERT: LANSWEEPER;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1461552164] SERVICE ALERT: CBORD;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1461554402] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 79% (15 min average)
[1461555002] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461555122] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1461555242] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 94% (15 min average)
[1461555842] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461555962] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461556082] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461556682] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 91% (15 min average)