[1459225461] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 71% (5 min average)
[1459225461] HOST ALERT: hallethirdsouth;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459225551] HOST ALERT: hallethirdsouth;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1459225641] HOST ALERT: hallethirdsouth;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1459225682] HOST ALERT: hallethirdsouth;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1459225757] HOST ALERT: hallethirdsouth;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1459225847] HOST ALERT: hallethirdsouth;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1459225937] HOST ALERT: hallethirdsouth;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1459226027] HOST ALERT: hallethirdsouth;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1459226117] HOST ALERT: hallethirdsouth;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1459226207] HOST ALERT: hallethirdsouth;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1459226661] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459226781] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1459226900] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1459229897] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 23% (5 min average)
[1459230497] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1459230617] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 79% (5 min average)
[1459231217] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459231337] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1459231457] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 94% (5 min average)
[1459232057] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1459232250] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 84.90 Gb - used: 79.01 Gb (93%) - free 5.89 Gb (7%)
[1459232370] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 84.90 Gb - used: 79.01 Gb (93%) - free 5.89 Gb (7%)
[1459232490] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 84.90 Gb - used: 79.01 Gb (93%) - free 5.89 Gb (7%)
[1459232745] SERVICE ALERT: EVENTS-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1459232865] SERVICE ALERT: EVENTS-2012;CPU Load;OK;SOFT;2;CPU Load 21% (5 min average)
[1459233573] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1459233693] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 95% (5 min average)
[1459233813] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 93% (5 min average)
[1459234413] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1459234457] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (5 min average)
[1459235013] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 78% (5 min average)
[1459235057] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1459235177] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1459235297] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1459235613] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1459235733] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 82% (5 min average)
[1459236497] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 93% (5 min average)
[1459237097] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1459237697] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1459238297] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1459240097] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 92% (5 min average)
[1459240697] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 79% (5 min average)
[1459241897] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1459242017] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 59% (5 min average)
[1459242617] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1459242737] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1459242857] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1459244057] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 70% (5 min average)
[1459244657] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1459244777] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 84% (5 min average)
[1459245977] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1459246097] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1459246217] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1459250067] SERVICE ALERT: NETIDADPROXY;Service: DNS Client;CRITICAL;SOFT;1;Dnscache: Stopped
[1459250187] SERVICE ALERT: NETIDADPROXY;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1459250627] SERVICE ALERT: ADCA1;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1459250747] SERVICE ALERT: ADCA1;CPU Load;OK;SOFT;2;CPU Load 22% (5 min average)
[1459251776] SERVICE ALERT: IMAGES;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459251792] SERVICE ALERT: IMAGES;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459251806] HOST ALERT: IMAGES;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459251892] SERVICE ALERT: IMAGES;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1459251894] HOST ALERT: IMAGES;UP;SOFT;2;PING WARNING - Packet loss = 82%, RTA = 0.39 ms
[1459251896] SERVICE ALERT: IMAGES;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1459251996] SERVICE ALERT: IMAGES;Service: Firewall;OK;SOFT;3;MpsSvc: Started
[1459252469] HOST ALERT: owen101lab3;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.49 ms
[1459252512] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 50% - OK!
[1459252583] SERVICE ALERT: owen101lab3;Device Status;OK;HARD;3;Ready
[1459252730] SERVICE ALERT: owen101lab3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1459252817] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 90% (5 min average)
[1459253417] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 52% (5 min average)
[1459254017] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 91% (5 min average)
[1459254137] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 96% (5 min average)
[1459254257] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1459254332] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1459254452] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 44% (5 min average)
[1459254909] SERVICE ALERT: prayharroldpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1459254969] SERVICE ALERT: prayharroldpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1459255029] SERVICE ALERT: prayharroldpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1459255287] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 84.90 Gb - used: 80.80 Gb (95%) - free 4.10 Gb (5%)
[1459255457] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 58% (5 min average)
[1459256274] HOST DOWNTIME ALERT: ADWSUSSERVERS;CANCELLED; Scheduled downtime for host has been cancelled.
[1459256828] SERVICE ALERT: prayharroldpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1459258652] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1459258772] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1459258887] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 84.90 Gb - used: 80.44 Gb (95%) - free 4.46 Gb (5%)
[1459258892] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 96% (5 min average)
[1459259492] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 47% (5 min average)
[1459262487] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;OK;HARD;3;c:\ - total: 84.90 Gb - used: 76.13 Gb (90%) - free 8.77 Gb (10%)
[1459262657] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1459262777] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 85% (5 min average)
[1459263377] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1459263497] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 65% (5 min average)
[1459263722] SERVICE ALERT: rooseveltpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1459263782] SERVICE ALERT: rooseveltpk;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! All Trays Still Full
[1459263858] HOST ALERT: prayharroldpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459263948] HOST ALERT: prayharroldpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1459263983] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459264013] HOST ALERT: prayharroldpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1459264033] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459264043] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459264063] HOST ALERT: prayharroldpk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1459264088] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459264093] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459264103] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459264118] HOST ALERT: prayharroldpk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1459264148] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459264153] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459264178] HOST ALERT: prayharroldpk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1459264208] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459264238] HOST ALERT: prayharroldpk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1459264328] HOST ALERT: prayharroldpk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1459264418] HOST ALERT: prayharroldpk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1459264508] HOST ALERT: prayharroldpk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1459265142] HOST ALERT: prayharroldpk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.60 ms
[1459265297] SERVICE ALERT: prayharroldpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459265347] SERVICE ALERT: prayharroldpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1459265347] SERVICE FLAPPING ALERT: prayharroldpk;Paper Tray Status;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1459265402] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 24% - OK!
[1459265692] SERVICE ALERT: halleground1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459265752] SERVICE ALERT: halleground1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459265812] SERVICE ALERT: halleground1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459265897] SERVICE ALERT: prayharroldpk;Device Status;OK;HARD;3;Processing job from tray 4
[1459266497] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1459266617] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1459266737] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1459267337] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1459267358] SERVICE ALERT: halleground2;Device Status;OK;HARD;3;Warming up
[1459267722] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1459267782] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 46% - OK!
[1459267937] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1459267958] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459268018] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459268078] SERVICE ALERT: halleground2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459268212] SERVICE ALERT: halleground1;Device Status;OK;HARD;3;Ready
[1459268228] SERVICE ALERT: C4SAS-SE;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1459268348] SERVICE ALERT: C4SAS-SE;CPU Load;OK;SOFT;2;CPU Load 2% (5 min average)
[1459268812] SERVICE ALERT: halleground1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459268872] SERVICE ALERT: halleground1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459268932] SERVICE ALERT: halleground1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459270104] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1459270132] SERVICE ALERT: halleground1;Device Status;OK;HARD;3;Ready
[1459270132] SERVICE FLAPPING ALERT: halleground1;Device Status;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1459270164] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 4% - OK!
[1459270688] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459270749] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459270808] SERVICE ALERT: hallesecondsouth;Device Status;OK;SOFT;3;Remove any jammed paper.
[1459270944] HOST ALERT: marshallpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459271034] HOST ALERT: marshallpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1459271070] SERVICE ALERT: marshallpk;Black Cartridge HP CC364X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.137.101. Is the printer off?  Is the network down?
[1459271100] HOST ALERT: marshallpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1459271129] SERVICE ALERT: marshallpk;Black Cartridge HP CC364X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.137.101. Is the printer off?  Is the network down?
[1459271159] HOST ALERT: marshallpk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1459271188] SERVICE ALERT: marshallpk;Black Cartridge HP CC364X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.137.101. Is the printer off?  Is the network down?
[1459271218] HOST ALERT: marshallpk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1459271304] SERVICE ALERT: marshallpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.137.101. Is the printer off?  Is the network down?
[1459271308] HOST ALERT: marshallpk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1459271338] SERVICE ALERT: marshallpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.137.101. Is the printer off?  Is the network down?
[1459271364] SERVICE ALERT: marshallpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.137.101. Is the printer off?  Is the network down?
[1459271368] HOST ALERT: marshallpk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1459271398] SERVICE ALERT: marshallpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.137.101. Is the printer off?  Is the network down?
[1459271424] SERVICE ALERT: marshallpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.137.101. Is the printer off?  Is the network down?
[1459271428] HOST ALERT: marshallpk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1459271458] SERVICE ALERT: marshallpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.137.101. Is the printer off?  Is the network down?
[1459271488] HOST ALERT: marshallpk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1459271578] HOST ALERT: marshallpk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1459272043] SERVICE ALERT: strongpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.130.31. Is the printer off?  Is the network down?
[1459272073] HOST ALERT: strongpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459272103] SERVICE ALERT: strongpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.130.31. Is the printer off?  Is the network down?
[1459272133] HOST ALERT: strongpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1459272163] SERVICE ALERT: strongpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.130.31. Is the printer off?  Is the network down?
[1459272193] HOST ALERT: strongpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1459272283] HOST ALERT: strongpk;DOWN;SOFT;4;(Host check timed out after 30.00 seconds)
[1459272373] HOST ALERT: strongpk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1459272434] SERVICE ALERT: strongpk;Black Cartridge HP CC364X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.130.31. Is the printer off?  Is the network down?
[1459272463] HOST ALERT: strongpk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1459272494] SERVICE ALERT: strongpk;Black Cartridge HP CC364X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.130.31. Is the printer off?  Is the network down?
[1459272511] SERVICE ALERT: strongpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.130.31. Is the printer off?  Is the network down?
[1459272524] HOST ALERT: strongpk;DOWN;SOFT;7;(Host check timed out after 30.00 seconds)
[1459272554] SERVICE ALERT: strongpk;Black Cartridge HP CC364X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.130.31. Is the printer off?  Is the network down?
[1459272571] SERVICE ALERT: strongpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.130.31. Is the printer off?  Is the network down?
[1459272584] HOST ALERT: strongpk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1459272631] SERVICE ALERT: strongpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.130.31. Is the printer off?  Is the network down?
[1459272661] HOST ALERT: strongpk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1459272751] HOST ALERT: strongpk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1459272877] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 1 Trays Empty
[1459272938] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 1 Trays Empty
[1459272997] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 1 Trays Empty
[1459273336] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 72% (5 min average)
[1459273796] SERVICE ALERT: EMUSCRIPT;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459273808] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459273819] HOST ALERT: EMUSCRIPT;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.161)
[1459273865] SERVICE ALERT: EMUSCRIPT;VMWare Tools;CRITICAL;SOFT;1;connect to address 164.76.102.161 and port 12489: No route to host
[1459273868] HOST ALERT: EMUSCRIPT;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.161)
[1459273869] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459273889] SERVICE ALERT: EMUSCRIPT;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.102.161 and port 12489: No route to host
[1459273892] HOST ALERT: EMUSCRIPT;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.102.161)
[1459273899] SERVICE ALERT: EMUSCRIPT;Service: Windows Time;CRITICAL;SOFT;2;connect to address 164.76.102.161 and port 12489: No route to host
[1459273909] SERVICE ALERT: EMUSCRIPT;CPU Load;CRITICAL;SOFT;1;connect to address 164.76.102.161 and port 12489: No route to host
[1459273912] HOST ALERT: EMUSCRIPT;DOWN;SOFT;4;CRITICAL - Host Unreachable (164.76.102.161)
[1459273928] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459273936] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1459273975] HOST ALERT: EMUSCRIPT;DOWN;SOFT;5;CRITICAL - Host Unreachable (164.76.102.161)
[1459273985] SERVICE ALERT: EMUSCRIPT;VMWare Tools;CRITICAL;SOFT;2;connect to address 164.76.102.161 and port 12489: No route to host
[1459274004] SERVICE ALERT: EMUSCRIPT;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.102.161 and port 12489: No route to host
[1459274007] SERVICE ALERT: EMUSCRIPT;C:\ Drive Space;CRITICAL;SOFT;2;connect to address 164.76.102.161 and port 12489: No route to host
[1459274007] HOST ALERT: EMUSCRIPT;DOWN;SOFT;6;CRITICAL - Host Unreachable (164.76.102.161)
[1459274019] SERVICE ALERT: EMUSCRIPT;Service: Windows Time;CRITICAL;HARD;3;connect to address 164.76.102.161 and port 12489: No route to host
[1459274022] HOST ALERT: EMUSCRIPT;DOWN;SOFT;7;CRITICAL - Host Unreachable (164.76.102.161)
[1459274029] SERVICE ALERT: EMUSCRIPT;CPU Load;CRITICAL;SOFT;2;connect to address 164.76.102.161 and port 12489: No route to host
[1459274029] SERVICE ALERT: EMUSCRIPT;Uptime;CRITICAL;SOFT;1;connect to address 164.76.102.161 and port 12489: No route to host
[1459274056] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1459274077] SERVICE ALERT: EMUSCRIPT;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.102.161 and port 12489: No route to host
[1459274080] HOST ALERT: EMUSCRIPT;DOWN;SOFT;8;CRITICAL - Host Unreachable (164.76.102.161)
[1459274096] SERVICE ALERT: EMUSCRIPT;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.102.161 and port 12489: No route to host
[1459274099] HOST ALERT: EMUSCRIPT;DOWN;SOFT;9;CRITICAL - Host Unreachable (164.76.102.161)
[1459274105] SERVICE ALERT: EMUSCRIPT;VMWare Tools;CRITICAL;HARD;3;connect to address 164.76.102.161 and port 12489: No route to host
[1459274124] SERVICE ALERT: EMUSCRIPT;NSClient++ Version;CRITICAL;SOFT;2;connect to address 164.76.102.161 and port 12489: No route to host
[1459274127] HOST ALERT: EMUSCRIPT;DOWN;HARD;10;CRITICAL - Host Unreachable (164.76.102.161)
[1459274130] SERVICE ALERT: EMUSCRIPT;C:\ Drive Space;CRITICAL;HARD;3;connect to address 164.76.102.161 and port 12489: No route to host
[1459274149] SERVICE ALERT: EMUSCRIPT;CPU Load;CRITICAL;HARD;3;connect to address 164.76.102.161 and port 12489: No route to host
[1459274149] SERVICE ALERT: EMUSCRIPT;Uptime;CRITICAL;SOFT;2;connect to address 164.76.102.161 and port 12489: No route to host
[1459274176] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1459274189] SERVICE ALERT: EMUSCRIPT;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.102.161 and port 12489: No route to host
[1459274189] SERVICE ALERT: EMUSCRIPT;Service: NetLogin;CRITICAL;SOFT;1;connect to address 164.76.102.161 and port 12489: No route to host
[1459274197] SERVICE ALERT: EMUSCRIPT;Memory Usage;CRITICAL;SOFT;2;connect to address 164.76.102.161 and port 12489: No route to host
[1459274206] SERVICE ALERT: EMUSQLEXPRESS;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459274209] HOST ALERT: EMUSQLEXPRESS;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.163)
[1459274216] SERVICE ALERT: EMUSCRIPT;Service: Firewall;CRITICAL;SOFT;2;connect to address 164.76.102.161 and port 12489: No route to host
[1459274237] SERVICE ALERT: EMUSQLEXPRESS;Service: NetLogin;CRITICAL;SOFT;1;connect to address 164.76.102.163 and port 12489: No route to host
[1459274239] SERVICE ALERT: EMUSCRIPT;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.102.161 and port 12489: No route to host
[1459274240] HOST ALERT: EMUSQLEXPRESS;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.163)
[1459274242] SERVICE ALERT: EMUSCRIPT;NSClient++ Version;CRITICAL;HARD;3;connect to address 164.76.102.161 and port 12489: No route to host
[1459274271] SERVICE ALERT: EMUSCRIPT;Uptime;CRITICAL;HARD;3;connect to address 164.76.102.161 and port 12489: No route to host
[1459274282] SERVICE ALERT: EMUSCRIPT;Service: User Profile;CRITICAL;SOFT;1;connect to address 164.76.102.161 and port 12489: No route to host
[1459274285] SERVICE ALERT: EMUSCRIPT;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.102.161 and port 12489: No route to host
[1459274303] HOST ALERT: EMUSQLEXPRESS;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.102.163)
[1459274307] SERVICE ALERT: EMUSCRIPT;Service: NetLogin;CRITICAL;SOFT;2;connect to address 164.76.102.161 and port 12489: No route to host
[1459274307] SERVICE ALERT: EMUSCRIPT;Service: Group Policy;CRITICAL;SOFT;2;connect to address 164.76.102.161 and port 12489: No route to host
[1459274308] SERVICE ALERT: EMUSQLEXPRESS;Service: Group Policy;CRITICAL;SOFT;2;connect to address 164.76.102.163 and port 12489: No route to host
[1459274317] SERVICE ALERT: EMUSCRIPT;Memory Usage;CRITICAL;HARD;3;connect to address 164.76.102.161 and port 12489: No route to host
[1459274328] SERVICE ALERT: EMUSQLEXPRESS;Service: User Profile;CRITICAL;SOFT;1;connect to address 164.76.102.163 and port 12489: No route to host
[1459274328] SERVICE ALERT: EMUSQLEXPRESS;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.102.163 and port 12489: No route to host
[1459274329] SERVICE ALERT: EMUSCRIPT;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.102.161 and port 12489: No route to host
[1459274331] HOST ALERT: EMUSQLEXPRESS;DOWN;SOFT;4;CRITICAL - Host Unreachable (164.76.102.163)
[1459274336] SERVICE ALERT: EMUSCRIPT;Service: Firewall;CRITICAL;HARD;3;connect to address 164.76.102.161 and port 12489: No route to host
[1459274357] SERVICE ALERT: EMUSQLEXPRESS;Service: NetLogin;CRITICAL;SOFT;2;connect to address 164.76.102.163 and port 12489: No route to host
[1459274359] SERVICE ALERT: EMUSCRIPT;Service: RPC;CRITICAL;SOFT;2;connect to address 164.76.102.161 and port 12489: No route to host
[1459274360] HOST ALERT: EMUSQLEXPRESS;DOWN;SOFT;5;CRITICAL - Host Unreachable (164.76.102.163)
[1459274377] SERVICE ALERT: EMUSQLEXPRESS;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.102.163 and port 12489: No route to host
[1459274380] HOST ALERT: EMUSQLEXPRESS;DOWN;SOFT;6;CRITICAL - Host Unreachable (164.76.102.163)
[1459274402] SERVICE ALERT: EMUSCRIPT;Service: User Profile;CRITICAL;SOFT;2;connect to address 164.76.102.161 and port 12489: No route to host
[1459274405] SERVICE ALERT: EMUSCRIPT;Service: DNS Client;CRITICAL;SOFT;2;connect to address 164.76.102.161 and port 12489: No route to host
[1459274427] SERVICE ALERT: EMUSCRIPT;Service: NetLogin;CRITICAL;HARD;3;connect to address 164.76.102.161 and port 12489: No route to host
[1459274428] SERVICE ALERT: EMUSQLEXPRESS;Service: Group Policy;CRITICAL;HARD;3;connect to address 164.76.102.163 and port 12489: No route to host
[1459274431] HOST ALERT: EMUSQLEXPRESS;DOWN;SOFT;7;CRITICAL - Host Unreachable (164.76.102.163)
[1459274433] SERVICE ALERT: EMUSCRIPT;Service: Group Policy;CRITICAL;HARD;3;connect to address 164.76.102.161 and port 12489: No route to host
[1459274448] SERVICE ALERT: EMUSQLEXPRESS;Service: User Profile;CRITICAL;SOFT;2;connect to address 164.76.102.163 and port 12489: No route to host
[1459274448] SERVICE ALERT: EMUSQLEXPRESS;NSClient++ Version;CRITICAL;SOFT;2;connect to address 164.76.102.163 and port 12489: No route to host
[1459274449] SERVICE ALERT: EMUSCRIPT;Service: WMI;CRITICAL;SOFT;2;connect to address 164.76.102.161 and port 12489: No route to host
[1459274451] HOST ALERT: EMUSQLEXPRESS;DOWN;SOFT;8;CRITICAL - Host Unreachable (164.76.102.163)
[1459274465] SERVICE ALERT: EMUSQLEXPRESS;Uptime;CRITICAL;SOFT;1;connect to address 164.76.102.163 and port 12489: No route to host
[1459274468] HOST ALERT: EMUSQLEXPRESS;DOWN;SOFT;9;CRITICAL - Host Unreachable (164.76.102.163)
[1459274476] SERVICE ALERT: EMUSQLEXPRESS;Service: NetLogin;CRITICAL;HARD;3;connect to address 164.76.102.163 and port 12489: No route to host
[1459274479] SERVICE ALERT: EMUSCRIPT;Service: RPC;CRITICAL;HARD;3;connect to address 164.76.102.161 and port 12489: No route to host
[1459274497] SERVICE ALERT: EMUSQLEXPRESS;Service: WMI;CRITICAL;SOFT;2;connect to address 164.76.102.163 and port 12489: No route to host
[1459274500] HOST ALERT: EMUSQLEXPRESS;DOWN;HARD;10;CRITICAL - Host Unreachable (164.76.102.163)
[1459274522] SERVICE ALERT: EMUSCRIPT;Service: User Profile;CRITICAL;HARD;3;connect to address 164.76.102.161 and port 12489: No route to host
[1459274525] SERVICE ALERT: EMUSCRIPT;Service: DNS Client;CRITICAL;HARD;3;connect to address 164.76.102.161 and port 12489: No route to host
[1459274528] SERVICE ALERT: hallesecondsouth;Device Status;OK;HARD;3;MENUS  IPDS MENU  RETRIEVE JOB  INFORMATION
[1459274551] SERVICE ALERT: EMUSQLEXPRESS;CPU Load;CRITICAL;SOFT;1;connect to address 164.76.102.163 and port 12489: No route to host
[1459274568] SERVICE ALERT: EMUSQLEXPRESS;NSClient++ Version;CRITICAL;HARD;3;connect to address 164.76.102.163 and port 12489: No route to host
[1459274568] SERVICE ALERT: EMUSQLEXPRESS;Service: User Profile;CRITICAL;HARD;3;connect to address 164.76.102.163 and port 12489: No route to host
[1459274569] SERVICE ALERT: EMUSCRIPT;Service: WMI;CRITICAL;HARD;3;connect to address 164.76.102.161 and port 12489: No route to host
[1459274585] SERVICE ALERT: EMUSQLEXPRESS;Uptime;CRITICAL;SOFT;2;connect to address 164.76.102.163 and port 12489: No route to host
[1459274599] SERVICE ALERT: EMUSQLEXPRESS;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.102.163 and port 12489: No route to host
[1459274617] SERVICE ALERT: EMUSQLEXPRESS;Service: WMI;CRITICAL;HARD;3;connect to address 164.76.102.163 and port 12489: No route to host
[1459274636] SERVICE ALERT: EMUSQLEXPRESS;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.102.163 and port 12489: No route to host
[1459274653] SERVICE ALERT: hallesecondsouth;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1459274671] SERVICE ALERT: EMUSQLEXPRESS;CPU Load;CRITICAL;SOFT;2;connect to address 164.76.102.163 and port 12489: No route to host
[1459274691] SERVICE ALERT: EMUSQLEXPRESS;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.102.163 and port 12489: No route to host
[1459274705] SERVICE ALERT: EMUSQLEXPRESS;Uptime;CRITICAL;HARD;3;connect to address 164.76.102.163 and port 12489: No route to host
[1459274713] SERVICE ALERT: hallesecondsouth;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1459274719] SERVICE ALERT: EMUSQLEXPRESS;Memory Usage;CRITICAL;SOFT;2;connect to address 164.76.102.163 and port 12489: No route to host
[1459274745] SERVICE ALERT: EMUSQLEXPRESS;VMWare Tools;CRITICAL;SOFT;1;connect to address 164.76.102.163 and port 12489: No route to host
[1459274745] SERVICE ALERT: EMUSQLEXPRESS;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.102.163 and port 12489: No route to host
[1459274745] SERVICE ALERT: EMUSQLEXPRESS;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.102.163 and port 12489: No route to host
[1459274748] SERVICE ALERT: EMUSQLEXPRESS;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.102.163 and port 12489: No route to host
[1459274756] SERVICE ALERT: EMUSQLEXPRESS;Service: RPC;CRITICAL;SOFT;2;connect to address 164.76.102.163 and port 12489: No route to host
[1459274773] SERVICE ALERT: hallesecondsouth;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1459274791] SERVICE ALERT: EMUSQLEXPRESS;CPU Load;CRITICAL;HARD;3;connect to address 164.76.102.163 and port 12489: No route to host
[1459274811] SERVICE ALERT: EMUSQLEXPRESS;Service: DNS Client;CRITICAL;SOFT;2;connect to address 164.76.102.163 and port 12489: No route to host
[1459274839] SERVICE ALERT: EMUSQLEXPRESS;Memory Usage;CRITICAL;HARD;3;connect to address 164.76.102.163 and port 12489: No route to host
[1459274865] SERVICE ALERT: EMUSQLEXPRESS;VMWare Tools;CRITICAL;SOFT;2;connect to address 164.76.102.163 and port 12489: No route to host
[1459274865] SERVICE ALERT: EMUSQLEXPRESS;C:\ Drive Space;CRITICAL;SOFT;2;connect to address 164.76.102.163 and port 12489: No route to host
[1459274865] SERVICE ALERT: EMUSQLEXPRESS;Service: Windows Time;CRITICAL;SOFT;2;connect to address 164.76.102.163 and port 12489: No route to host
[1459274874] SERVICE ALERT: EMUSQLEXPRESS;Service: Firewall;CRITICAL;SOFT;2;connect to address 164.76.102.163 and port 12489: No route to host
[1459274874] SERVICE ALERT: EMUSQLEXPRESS;Service: RPC;CRITICAL;HARD;3;connect to address 164.76.102.163 and port 12489: No route to host
[1459274930] SERVICE ALERT: EMUSQLEXPRESS;Service: DNS Client;CRITICAL;HARD;3;connect to address 164.76.102.163 and port 12489: No route to host
[1459274985] SERVICE ALERT: EMUSQLEXPRESS;Service: Windows Time;CRITICAL;HARD;3;connect to address 164.76.102.163 and port 12489: No route to host
[1459274985] SERVICE ALERT: EMUSQLEXPRESS;C:\ Drive Space;CRITICAL;HARD;3;connect to address 164.76.102.163 and port 12489: No route to host
[1459274985] SERVICE ALERT: EMUSQLEXPRESS;VMWare Tools;CRITICAL;HARD;3;connect to address 164.76.102.163 and port 12489: No route to host
[1459274994] SERVICE ALERT: EMUSQLEXPRESS;Service: Firewall;CRITICAL;HARD;3;connect to address 164.76.102.163 and port 12489: No route to host
[1459275134] SERVICE ALERT: hallesecondsouth;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1459275164] HOST ALERT: hallesecondsouth;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459275188] SERVICE ALERT: hallesecondsouth;Device Status;OK;SOFT;2;Genuine HP cartridge installed
[1459275192] HOST ALERT: hallesecondsouth;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1459275376] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 93% (5 min average)
[1459275378] SERVICE ALERT: hallesecondsouth;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1459275386] SERVICE ALERT: NETIDADPROXY;Service: DNS Client;CRITICAL;SOFT;1;Dnscache: Stopped
[1459275408] HOST ALERT: hallesecondsouth;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459275498] HOST ALERT: hallesecondsouth;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1459275506] SERVICE ALERT: NETIDADPROXY;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1459275570] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1459275588] HOST ALERT: hallesecondsouth;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1459275630] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1459275660] HOST ALERT: hallesecondsouth;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1459275690] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1459275720] HOST ALERT: hallesecondsouth;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1459275794] SERVICE ALERT: hallesecondsouth;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1459275810] HOST ALERT: hallesecondsouth;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1459275854] SERVICE ALERT: hallesecondsouth;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1459275884] HOST ALERT: hallesecondsouth;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1459275914] SERVICE ALERT: hallesecondsouth;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1459275944] HOST ALERT: hallesecondsouth;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1459275976] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1459276008] HOST ALERT: hallesecondsouth;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1459276098] HOST ALERT: hallesecondsouth;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1459276148] SERVICE FLAPPING ALERT: prayharroldpk;Paper Tray Status;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1459276477] SERVICE ALERT: halleground2;Device Status;OK;HARD;3;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1459276596] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1459277077] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459277109] SERVICE ALERT: hallesecondsouth;Device Status;OK;HARD;3;Clearing paper path
[1459277109] SERVICE FLAPPING ALERT: hallesecondsouth;Device Status;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1459277113] HOST ALERT: hallesecondsouth;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.45 ms
[1459277137] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459277173] SERVICE ALERT: hallesecondsouth;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1459277197] SERVICE ALERT: halleground2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459277197] SERVICE FLAPPING ALERT: halleground2;Device Status;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1459277484] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 3% - OK!
[1459277554] SERVICE ALERT: hallegroundcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1459277708] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459277767] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459277827] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459278002] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (34.4% change >= 20.0% threshold)
[1459278002] SERVICE FLAPPING ALERT: halleground2;Device Status;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1459278003] SERVICE FLAPPING ALERT: hallesecondsouth;Device Status;STARTED; Service appears to have started flapping (29.3% change >= 20.0% threshold)
[1459278003] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1459278003] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1459278003] SERVICE DOWNTIME ALERT: BOEAPPTEST;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1459278023] SERVICE ALERT: EMUSCRIPT;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1459278026] SERVICE ALERT: EMUSCRIPT;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1459278027] HOST ALERT: EMUSCRIPT;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1459278075] SERVICE ALERT: EMUSCRIPT;Service: RPC;OK;HARD;3;RpcSs: Started
[1459278118] SERVICE ALERT: EMUSCRIPT;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1459278123] SERVICE ALERT: EMUSCRIPT;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1459278165] SERVICE ALERT: EMUSCRIPT;Service: WMI;OK;HARD;3;Winmgmt: Started
[1459278215] SERVICE ALERT: EMUSCRIPT;Service: Windows Time;OK;HARD;3;W32Time: Started
[1459278301] SERVICE ALERT: EMUSCRIPT;VMWare Tools;OK;HARD;3;VMTools: Started
[1459278323] SERVICE ALERT: EMUSCRIPT;C:\ Drive Space;OK;HARD;3;c:\ - total: 79.66 Gb - used: 16.43 Gb (21%) - free 63.23 Gb (79%)
[1459278345] SERVICE ALERT: EMUSCRIPT;CPU Load;OK;HARD;3;CPU Load 2% (5 min average)
[1459278440] SERVICE ALERT: EMUSCRIPT;NSClient++ Version;OK;HARD;3;NSClient++ 0,4,1,73 2012-12-17
[1459278467] SERVICE ALERT: EMUSCRIPT;Uptime;OK;HARD;3;System Uptime - 12 day(s) 7 hour(s) 14 minute(s)
[1459278513] SERVICE ALERT: EMUSCRIPT;Memory Usage;OK;HARD;3;Memory usage: total:10943.64 MB - used: 993.72 MB (9%) - free: 9949.93 MB (91%)
[1459278533] SERVICE ALERT: EMUSCRIPT;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1459279093] SERVICE ALERT: IMAGES;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459279213] SERVICE ALERT: IMAGES;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1459281374] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1459281602] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (25.5% change >= 20.0% threshold)
[1459281603] SERVICE FLAPPING ALERT: hallesecondsouth;Device Status;STARTED; Service appears to have started flapping (25.3% change >= 20.0% threshold)
[1459281603] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1459281603] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1459281603] SERVICE DOWNTIME ALERT: BOEAPPTEST;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1459281623] SERVICE ALERT: EMUSQLEXPRESS;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1459281627] HOST ALERT: EMUSQLEXPRESS;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1459281673] SERVICE ALERT: EMUSQLEXPRESS;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1459281696] SERVICE ALERT: EMUSQLEXPRESS;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1459281764] SERVICE ALERT: EMUSQLEXPRESS;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1459281764] SERVICE ALERT: EMUSQLEXPRESS;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1459281813] SERVICE ALERT: EMUSQLEXPRESS;Service: WMI;OK;HARD;3;Winmgmt: Started
[1459281900] SERVICE ALERT: EMUSQLEXPRESS;Uptime;OK;HARD;3;System Uptime - 12 day(s) 16 hour(s) 51 minute(s)
[1459281960] HOST DOWNTIME ALERT: ADRODC-WEMU;STARTED; Host has entered a period of scheduled downtime
[1459281973] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1459281976] SERVICE ALERT: ADRODC-WEMU;Service: Group Policy;CRITICAL;SOFT;1;gpsvc: Stopped
[1459281986] SERVICE ALERT: EMUSQLEXPRESS;CPU Load;OK;HARD;3;CPU Load 0% (5 min average)
[1459282034] SERVICE ALERT: EMUSQLEXPRESS;Memory Usage;OK;HARD;3;Memory usage: total:10943.64 MB - used: 2258.57 MB (21%) - free: 8685.08 MB (79%)
[1459282070] SERVICE ALERT: EMUSQLEXPRESS;Service: RPC;OK;HARD;3;RpcSs: Started
[1459282115] SERVICE ALERT: ADRODC-WEMU;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1459282124] SERVICE ALERT: EMUSQLEXPRESS;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1459282137] SERVICE ALERT: ADRODC-WEMU;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459282145] HOST ALERT: ADRODC-WEMU;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459282183] SERVICE ALERT: EMUSQLEXPRESS;Service: Windows Time;OK;HARD;3;W32Time: Started
[1459282183] SERVICE ALERT: EMUSQLEXPRESS;C:\ Drive Space;OK;HARD;3;c:\ - total: 79.66 Gb - used: 25.14 Gb (32%) - free 54.52 Gb (68%)
[1459282183] SERVICE ALERT: EMUSQLEXPRESS;VMWare Tools;OK;HARD;3;VMTools: Started
[1459282184] SERVICE ALERT: ADRODC-WEMU;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459282214] HOST ALERT: ADRODC-WEMU;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1459282235] SERVICE ALERT: ADRODC-WEMU;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1459282256] SERVICE ALERT: ADRODC-WEMU;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1459282265] HOST ALERT: ADRODC-WEMU;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1459282268] SERVICE ALERT: ADRODC-WEMU;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459282298] HOST ALERT: ADRODC-WEMU;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1459282304] SERVICE ALERT: ADRODC-WEMU;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1459282314] SERVICE ALERT: ADRODC-WEMU;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459282334] HOST ALERT: ADRODC-WEMU;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1459282357] SERVICE ALERT: ADRODC-WEMU;Service: User Profile;OK;SOFT;3;ProfSvc: Started
[1459282363] SERVICE ALERT: ADRODC-WEMU;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459282368] SERVICE ALERT: ADRODC-WEMU;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 2 minute(s)
[1459282387] HOST ALERT: ADRODC-WEMU;DOWN;SOFT;6;(Host check timed out after 30.00 seconds)
[1459282404] SERVICE ALERT: ADRODC-WEMU;Service: WMI;OK;SOFT;3;Winmgmt: Started
[1459282414] SERVICE ALERT: ADRODC-WEMU;C:\ Drive Space;OK;SOFT;2;c:\ - total: 97.31 Gb - used: 13.71 Gb (14%) - free 83.61 Gb (86%)
[1459282418] HOST ALERT: ADRODC-WEMU;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1459282463] SERVICE ALERT: ADRODC-WEMU;CPU Load;OK;SOFT;2;CPU Load 22% (5 min average)
[1459282465] HOST ALERT: ADRODC-WEMU;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1459282513] HOST ALERT: ADRODC-WEMU;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1459282556] HOST ALERT: ADRODC-WEMU;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1459282573] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1459282815] SERVICE ALERT: ADRODC-WEMU;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1459283173] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1459283883] SERVICE ALERT: markjeffersonpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1459283889] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1459283913] HOST ALERT: markjeffersonpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459283943] SERVICE ALERT: markjeffersonpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1459283949] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1459283973] HOST ALERT: markjeffersonpk;UP;SOFT;2;PING WARNING - Packet loss = 90%, RTA = 0.82 ms
[1459283997] SERVICE ALERT: markjeffersonpk;Device Status;OK;SOFT;3;Warming up
[1459284003] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;OK;SOFT;3;Black Print Cartridge HP C8543X is at 63% - OK!
[1459284592] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Warming up
[1459284596] HOST ALERT: hallefirst2;UP;HARD;10;PING WARNING - Packet loss = 81%, RTA = 0.54 ms
[1459284854] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 5% - OK!
[1459284863] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1459285059] SERVICE ALERT: EVENTS-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1459285179] SERVICE ALERT: EVENTS-2012;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1459285756] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1459285796] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459285816] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1459285826] HOST ALERT: hallefirst2;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1459285846] HOST ALERT: hallefirst1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459285856] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459285876] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1459285883] HOST ALERT: hallefirst1;UP;SOFT;2;PING OK - Packet loss = 37%, RTA = 0.68 ms
[1459285886] HOST ALERT: hallefirst2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1459285916] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459285946] HOST ALERT: hallefirst2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1459286036] HOST ALERT: hallefirst2;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1459286060] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459286069] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459286090] HOST ALERT: hallefirst2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1459286120] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459286129] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459286150] HOST ALERT: hallefirst2;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1459286180] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459286189] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459286210] HOST ALERT: hallefirst2;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1459286225] SERVICE ALERT: hallesecondsouth;Device Status;OK;HARD;3;Clearing paper path
[1459286300] HOST ALERT: hallefirst2;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1459286390] HOST ALERT: hallefirst2;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1459286480] HOST ALERT: hallefirst2;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1459286485] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1459286506] HOST ALERT: hallefirst1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459286547] SERVICE ALERT: hallefirst1;Device Status;OK;SOFT;2;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1459286551] HOST ALERT: hallefirst1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.71 ms
[1459286551] HOST FLAPPING ALERT: hallefirst1;STARTED; Host appears to have started flapping (22.4% change > 20.0% threshold)
[1459286778] HOST ALERT: hallefirst2;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.54 ms
[1459286784] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1459286824] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459286884] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459286944] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459287071] SERVICE ALERT: hallefirst1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1459287110] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459287371] SERVICE ALERT: hallesecondsouth;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1459287431] SERVICE ALERT: hallesecondsouth;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1459287491] SERVICE ALERT: hallesecondsouth;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1459287708] HOST ALERT: hallefirst2;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1459287716] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459287716] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1459287746] HOST ALERT: hallefirst2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1459287836] HOST ALERT: hallefirst2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1459287926] HOST ALERT: hallefirst2;DOWN;SOFT;4;(Host check timed out after 30.00 seconds)
[1459287973] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 88% (5 min average)
[1459287989] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459288010] HOST ALERT: hallefirst2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1459288049] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459288079] HOST ALERT: hallefirst2;DOWN;SOFT;6;(Host check timed out after 30.00 seconds)
[1459288103] SERVICE ALERT: NETIDADPROXY;Service: DNS Client;CRITICAL;SOFT;1;Dnscache: Stopped
[1459288109] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459288109] SERVICE FLAPPING ALERT: hallefirst2;Paper Tray Status;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1459288139] HOST ALERT: hallefirst2;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1459288223] SERVICE ALERT: NETIDADPROXY;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1459288229] HOST ALERT: hallefirst2;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1459288319] HOST ALERT: hallefirst2;DOWN;SOFT;9;(Host check timed out after 30.00 seconds)
[1459288409] HOST ALERT: hallefirst2;DOWN;HARD;10;(Host check timed out after 30.00 seconds)
[1459288573] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1459288693] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1459288813] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1459288870] HOST FLAPPING ALERT: hallefirst1;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1459290519] SERVICE ALERT: owen101labcolor;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1459290520] SERVICE ALERT: owen101labcolor;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1459290549] HOST ALERT: owen101labcolor;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1459290560] SERVICE ALERT: owen101labcolor;Toner Collection Unit HP CE265A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1459290574] SERVICE ALERT: owen101labcolor;Device Status;OK;SOFT;2;YELLOW CARTRIDGE LOW 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Processing...
[1459290574] HOST ALERT: owen101labcolor;UP;SOFT;2;PING OK - Packet loss = 66%, RTA = 0.51 ms
[1459290575] SERVICE ALERT: owen101labcolor;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! All Trays Still Full
[1459290615] SERVICE ALERT: owen101labcolor;Toner Collection Unit HP CE265A;OK;SOFT;2;
[1459293115] SERVICE ALERT: hallesouthcommons;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1459293175] SERVICE ALERT: hallesouthcommons;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! 1 of 2 Trays Still Full
[1459296491] SERVICE ALERT: hallesecondsouth;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1459296543] SERVICE ALERT: hallesecondsouth;Device Status;OK;HARD;3;Clearing paper path
[1459297144] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459297203] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459297223] SERVICE ALERT: NETIDADPROXY;Service: DNS Client;CRITICAL;SOFT;1;Dnscache: Stopped
[1459297263] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459297343] SERVICE ALERT: NETIDADPROXY;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1459298909] SERVICE FLAPPING ALERT: hallefirst2;Paper Tray Status;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1459299116] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1459300213] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1459300735] SERVICE ALERT: IT-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1459300855] SERVICE ALERT: IT-2012;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1459301793] SERVICE ALERT: halleground2;Device Status;OK;HARD;3;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1459301879] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1459301939] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1459301999] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1459302599] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 80% - OK!
[1459303288] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 94% (5 min average)
[1459303408] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 48% (5 min average)
[1459303943] SERVICE ALERT: NETIDADPROXY;Service: DNS Client;CRITICAL;SOFT;1;Dnscache: Stopped
[1459304063] SERVICE ALERT: NETIDADPROXY;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1459304854] SERVICE ALERT: VC;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459304954] SERVICE ALERT: VC;CPU Load;OK;SOFT;2;CPU Load 37% (5 min average)
[1459306470] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1459306530] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1459306590] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1459307190] SERVICE ALERT: eateriespk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1459307794] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 1 Trays Empty
[1459307853] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 1 Trays Empty
[1459307913] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 1 Trays Empty
[1459308664] SERVICE FLAPPING ALERT: hallesecondsouth;Device Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1459309113] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full