[1456224727] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;1;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1456224847] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;2;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1456224967] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;HARD;3;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1456232167] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;OK;HARD;3;OK: No entries found
[1456239175] SERVICE ALERT: ADRODC-DET;Service: MDT Monitor;UNKNOWN;SOFT;1;Failed to open service MDT_Monitor: 1060: The specified service does not exist as an installed service.
[1456239294] SERVICE ALERT: ADRODC-DET;Service: MDT Monitor;UNKNOWN;SOFT;2;Failed to open service MDT_Monitor: 1060: The specified service does not exist as an installed service.
[1456239415] SERVICE ALERT: ADRODC-DET;Service: MDT Monitor;UNKNOWN;HARD;3;Failed to open service MDT_Monitor: 1060: The specified service does not exist as an installed service.
[1456241221] SERVICE ALERT: alexanderpk;Black Cartridge HP CC364X;WARNING;SOFT;1;Black Cartridge HP CC364X is at 2% - WARNING!
[1456241280] SERVICE ALERT: alexanderpk;Black Cartridge HP CC364X;WARNING;SOFT;2;Black Cartridge HP CC364X is at 2% - WARNING!
[1456241340] SERVICE ALERT: alexanderpk;Black Cartridge HP CC364X;WARNING;HARD;3;Black Cartridge HP CC364X is at 2% - WARNING!
[1456244520] SERVICE ALERT: ADDC5;GPO Replication;UNKNOWN;SOFT;1;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1456244537] SERVICE ALERT: ADDC5;GPO Replication;OK;SOFT;2;Export-CSV : The process cannot access the file 'C:\Program
[1456252020] SERVICE ALERT: ADDC5;GPO Replication;UNKNOWN;SOFT;1;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1456252043] SERVICE ALERT: ADDC5;GPO Replication;OK;SOFT;2;Export-CSV : The process cannot access the file 'C:\Program
[1456254394] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1456254460] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1456254490] HOST ALERT: hallefirst1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1456254520] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1456254550] HOST ALERT: hallefirst1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1456254614] HOST ALERT: hallefirst1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.56 ms
[1456255114] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 60% - OK!
[1456256108] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1456256168] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1456256228] SERVICE ALERT: hallefirst1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1456258628] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;Ready
[1456259520] SERVICE ALERT: ADDC5;GPO Replication;UNKNOWN;SOFT;1;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1456259540] SERVICE ALERT: ADDC5;GPO Replication;OK;SOFT;2;Export-CSV : The process cannot access the file 'C:\Program
[1456263461] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1456263521] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1456263581] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1456280508] SERVICE ALERT: DB01-2012;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1456280628] SERVICE ALERT: DB01-2012;CPU Load;WARNING;SOFT;2;CPU Load 90% (5 min average)
[1456280748] SERVICE ALERT: DB01-2012;CPU Load;OK;SOFT;3;CPU Load 87% (5 min average)
[1456285135] SERVICE ALERT: ADWSUSSERVERS;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.102.133 and port 12489: No route to host
[1456285135] SERVICE ALERT: ADWSUSSERVERS;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.102.133 and port 12489: No route to host
[1456285139] HOST ALERT: ADWSUSSERVERS;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.133)
[1456285156] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 12072;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1456285159] HOST ALERT: ADWSUSSERVERS;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.133)
[1456285194] HOST ALERT: ADWSUSSERVERS;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1456285243] SERVICE ALERT: ADWSUSSERVERS;C:\ Drive Space;OK;SOFT;2;c:\ - total: 79.66 Gb - used: 46.41 Gb (58%) - free 33.24 Gb (42%)
[1456285254] SERVICE ALERT: ADWSUSSERVERS;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1456285273] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 12072;OK;SOFT;2;Eventlog check ok