[1454404205] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.88.118! Make sure host is up and SNMP is configured properly.
[1454404259] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 9% - OK!
[1454420844] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1454420878] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1454420882] HOST ALERT: hallefirst1;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1454420904] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1454420964] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1454421156] SERVICE ALERT: hallefirst1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1454426606] SERVICE ALERT: halleground1;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1454426666] SERVICE ALERT: halleground1;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1454426726] SERVICE ALERT: halleground1;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1454428761] SERVICE ALERT: studentcenterpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1454431128] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1454431188] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1454431248] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;OK;SOFT;3;Black Print Cartridge HP C8543X is at 46% - OK!
[1454431524] SERVICE ALERT: halleground1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1454433485] HOST ALERT: mckennypk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1454433556] SERVICE ALERT: mckennypk;Paper Tray Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.152.26! Make sure host is up and SNMP is configured properly.
[1454433575] HOST ALERT: mckennypk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1454433616] SERVICE ALERT: mckennypk;Paper Tray Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.152.26! Make sure host is up and SNMP is configured properly.
[1454433646] HOST ALERT: mckennypk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1454433661] SERVICE ALERT: mckennypk;Black Cartridge HP CC364X;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.152.26! Make sure host is up and SNMP is configured properly.
[1454433676] SERVICE ALERT: mckennypk;Paper Tray Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.152.26! Make sure host is up and SNMP is configured properly.
[1454433691] HOST ALERT: mckennypk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1454433721] SERVICE ALERT: mckennypk;Black Cartridge HP CC364X;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.152.26! Make sure host is up and SNMP is configured properly.
[1454433751] HOST ALERT: mckennypk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1454433781] SERVICE ALERT: mckennypk;Black Cartridge HP CC364X;WARNING;HARD;3;WARNING: No SNMP response from 164.76.152.26! Make sure host is up and SNMP is configured properly.
[1454433811] HOST ALERT: mckennypk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1454433901] HOST ALERT: mckennypk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1454433991] HOST ALERT: mckennypk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1454434081] HOST ALERT: mckennypk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1454434171] HOST ALERT: mckennypk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1454437714] SERVICE ALERT: hallethirdsouth;Device Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.107.213! Make sure host is up and SNMP is configured properly.
[1454437774] SERVICE ALERT: hallethirdsouth;Device Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.107.213! Make sure host is up and SNMP is configured properly.
[1454437834] SERVICE ALERT: hallethirdsouth;Device Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.107.213! Make sure host is up and SNMP is configured properly.
[1454437885] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.152.26! Make sure host is up and SNMP is configured properly.
[1454437945] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.152.26! Make sure host is up and SNMP is configured properly.
[1454438005] SERVICE ALERT: mckennypk;Device Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.152.26! Make sure host is up and SNMP is configured properly.
[1454438575] SERVICE ALERT: mckennypk;Black Cartridge HP CC364X;OK;HARD;3;Black Cartridge HP CC364X is at 19% - OK!
[1454438579] HOST ALERT: mckennypk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1454438599] SERVICE ALERT: mckennypk;Device Status;OK;HARD;3;Ready 164.76.152.26
[1454439070] SERVICE ALERT: mckennypk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1454440270] SERVICE ALERT: mckennypk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1454440331] SERVICE ALERT: mckennypk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1454440391] SERVICE ALERT: mckennypk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1454441236] SERVICE ALERT: ADAUTHDC1;Service: Dell Backup Agent;UNKNOWN;SOFT;1;Failed to open service BackupAgent64: 1060: The specified service does not exist as an installed service.
[1454441281] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.104.232! Make sure host is up and SNMP is configured properly.
[1454441298] SERVICE ALERT: ADAUTHDC3;Service: Dell Backup Agent;UNKNOWN;SOFT;1;Failed to open service BackupAgent64: 1060: The specified service does not exist as an installed service.
[1454441302] SERVICE ALERT: hallefirst2;Paper Tray Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1454441313] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.104.232! Make sure host is up and SNMP is configured properly.
[1454441341] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.104.232! Make sure host is up and SNMP is configured properly.
[1454441356] SERVICE ALERT: ADAUTHDC1;Service: Dell Backup Agent;UNKNOWN;SOFT;2;Failed to open service BackupAgent64: 1060: The specified service does not exist as an installed service.
[1454441361] SERVICE ALERT: ADDC1;Service: Dell Backup Agent;UNKNOWN;SOFT;1;Failed to open service BackupAgent64: 1060: The specified service does not exist as an installed service.
[1454441362] SERVICE ALERT: hallefirst2;Paper Tray Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1454441373] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.104.232! Make sure host is up and SNMP is configured properly.
[1454441394] SERVICE ALERT: ADDC1;Service: Dell Backup Agent;OK;SOFT;2;ErdAgent: Started
[1454441401] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;HARD;3;WARNING: No SNMP response from 164.76.104.232! Make sure host is up and SNMP is configured properly.
[1454441403] HOST ALERT: hallefirst1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1454441418] SERVICE ALERT: ADAUTHDC3;Service: Dell Backup Agent;OK;SOFT;2;ErdAgent: Started
[1454441422] SERVICE ALERT: hallefirst2;Paper Tray Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1454441427] SERVICE ALERT: hallefirst3;Device Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.104.180! Make sure host is up and SNMP is configured properly.
[1454441427] SERVICE ALERT: hallefirst1;Device Status;OK;SOFT;3;Processing job
[1454441431] HOST ALERT: hallefirst1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.01 ms
[1454441476] SERVICE ALERT: ADAUTHDC1;Service: Dell Backup Agent;OK;SOFT;3;ErdAgent: Started
[1454441487] SERVICE ALERT: hallefirst3;Device Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.104.180! Make sure host is up and SNMP is configured properly.
[1454441509] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.104.180! Make sure host is up and SNMP is configured properly.
[1454441547] SERVICE ALERT: hallefirst3;Device Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.104.180! Make sure host is up and SNMP is configured properly.
[1454441557] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1454441569] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.104.180! Make sure host is up and SNMP is configured properly.
[1454441581] SERVICE ALERT: mckennypk;Black Cartridge HP CC364X;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.152.26! Make sure host is up and SNMP is configured properly.
[1454441605] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.152.26! Make sure host is up and SNMP is configured properly.
[1454441611] HOST ALERT: mckennypk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1454441617] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1454441625] SERVICE ALERT: hallefirst3;Paper Tray Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.104.180! Make sure host is up and SNMP is configured properly.
[1454441629] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;WARNING;HARD;3;WARNING: No SNMP response from 164.76.104.180! Make sure host is up and SNMP is configured properly.
[1454441641] SERVICE ALERT: mckennypk;Black Cartridge HP CC364X;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.152.26! Make sure host is up and SNMP is configured properly.
[1454441665] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.152.26! Make sure host is up and SNMP is configured properly.
[1454441671] HOST ALERT: mckennypk;DOWN;SOFT;2;(Host check timed out after 30.00 seconds)
[1454441677] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1454441685] SERVICE ALERT: hallefirst3;Paper Tray Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.104.180! Make sure host is up and SNMP is configured properly.
[1454441701] SERVICE ALERT: mckennypk;Black Cartridge HP CC364X;WARNING;HARD;3;WARNING: No SNMP response from 164.76.152.26! Make sure host is up and SNMP is configured properly.
[1454441707] HOST ALERT: hallefirst2;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1454441715] HOST ALERT: hallefirst3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1454441725] SERVICE ALERT: mckennypk;Device Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.152.26! Make sure host is up and SNMP is configured properly.
[1454441731] HOST ALERT: mckennypk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1454441741] SERVICE ALERT: hallefirst3;Paper Tray Status;OK;SOFT;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1454441745] HOST ALERT: hallefirst3;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.50 ms
[1454441771] HOST ALERT: hallefirst2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.53 ms
[1454441821] HOST ALERT: mckennypk;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1454441911] HOST ALERT: mckennypk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1454441995] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 44% - OK!
[1454442001] HOST ALERT: mckennypk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1454442017] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1454442089] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1454442091] HOST ALERT: mckennypk;DOWN;SOFT;7;(Host check timed out after 30.00 seconds)
[1454442141] SERVICE ALERT: hallefirst3;Device Status;OK;HARD;3;TRAY 2 EMPTY PLAIN LETTER   Ready
[1454442181] HOST ALERT: mckennypk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1454442224] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 12% - OK!
[1454442226] HOST ALERT: mckennypk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1454442271] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;10.00.00 SUPPLY MEMORY ERROR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1454442316] HOST ALERT: mckennypk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1454442690] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at WARNING level!
[1454443951] SERVICE ALERT: ADWSUSSERVERS;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1454443952] SERVICE ALERT: ADWSUSSERVERS;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1454444051] SERVICE ALERT: ADWSUSSERVERS;C:\ Drive Space;OK;SOFT;2;c:\ - total: 79.66 Gb - used: 38.55 Gb (48%) - free 41.11 Gb (52%)
[1454444052] SERVICE ALERT: ADWSUSSERVERS;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1454445895] SERVICE ALERT: mckennypk;Black Cartridge HP CC364X;OK;HARD;3;Black Cartridge HP CC364X is at 18% - OK!
[1454445899] HOST ALERT: mckennypk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.53 ms
[1454446235] SERVICE ALERT: mckennypk;Device Status;OK;HARD;3;Ready 164.76.152.26
[1454446235] SERVICE FLAPPING ALERT: mckennypk;Device Status;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1454446390] SERVICE ALERT: mckennypk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1454446802] SERVICE FLAPPING ALERT: mckennypk;Device Status;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1454447641] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1454447701] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1454447761] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1454447869] SERVICE ALERT: ADWSUSSERVERS;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1454447969] SERVICE ALERT: ADWSUSSERVERS;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1454450161] SERVICE ALERT: eateriespk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1454459905] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 1 Trays Empty
[1454459966] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 1 Trays Empty
[1454460025] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 1 Trays Empty
[1454460557] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1454460617] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! 1 of 3 Trays Still Full
[1454460625] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1454462311] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.111.220! Make sure host is up and SNMP is configured properly.
[1454462365] SERVICE ALERT: hallesouthcommons;Device Status;OK;SOFT;2;Checking printer
[1454466875] SERVICE ALERT: hallesecondsouth;Paper Tray Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.107.53! Make sure host is up and SNMP is configured properly.
[1454466905] HOST ALERT: hallesecondsouth;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1454466935] SERVICE ALERT: hallesecondsouth;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! 1 of 3 Trays Still Full
[1454466939] HOST ALERT: hallesecondsouth;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.57 ms
[1454466945] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;WARNING;SOFT;1;Black Cartridge HP CE260X is at 2% - WARNING!
[1454467006] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;WARNING;SOFT;2;Black Cartridge HP CE260X is at 2% - WARNING!
[1454467065] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;WARNING;HARD;3;Black Cartridge HP CE260X is at 2% - WARNING!