[1458879993] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 85% (5 min average)
[1458880593] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458880713] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1458880833] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1458884872] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1458884992] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 74% (5 min average)
[1458885592] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1458885712] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 83% (5 min average)
[1458886314] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 84.90 Gb - used: 78.93 Gb (93%) - free 5.97 Gb (7%)
[1458886435] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 84.90 Gb - used: 78.93 Gb (93%) - free 5.97 Gb (7%)
[1458886554] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 84.90 Gb - used: 78.93 Gb (93%) - free 5.97 Gb (7%)
[1458887512] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1458887632] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 90% (5 min average)
[1458887752] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 87% (5 min average)
[1458888352] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1458888471] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1458888591] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 85% (5 min average)
[1458889191] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1458889312] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 96% (5 min average)
[1458889431] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 78% (5 min average)
[1458890031] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1458890149] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 82% (5 min average)
[1458890749] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 94% (5 min average)
[1458890869] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1458890989] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 80% (5 min average)
[1458891589] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1458891709] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 95% (5 min average)
[1458891829] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 95% (5 min average)
[1458892429] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 76% (5 min average)
[1458893433] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 65% (5 min average)
[1458894281] SERVICE ALERT: NETIDADPROXY;Service: DNS Client;CRITICAL;SOFT;1;Dnscache: Stopped
[1458894401] SERVICE ALERT: NETIDADPROXY;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1458894633] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1458894753] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 63% (5 min average)
[1458895353] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1458895473] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 95% (5 min average)
[1458895593] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 80% (5 min average)
[1458897993] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458898113] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 70% (5 min average)
[1458900513] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1458900633] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1458900753] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1458901353] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 71% (5 min average)
[1458903753] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458903873] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 77% (5 min average)
[1458904473] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458904593] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1458904713] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 95% (5 min average)
[1458904713] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1458905313] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 44% (5 min average)
[1458905913] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458906033] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 94% (5 min average)
[1458906153] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 92% (5 min average)
[1458906753] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1458908553] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 88% (5 min average)
[1458910353] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1458910473] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 62% (5 min average)
[1458910803] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (35.6% change >= 20.0% threshold)
[1458910803] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1458910803] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1458910803] SERVICE DOWNTIME ALERT: BOEAPPTEST;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1458910813] HOST ALERT: AD-RMAD;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1458911088] SERVICE ALERT: AD-RMAD;VMWare Tools;OK;HARD;3;VMTools: Started
[1458911089] SERVICE ALERT: AD-RMAD;Service: RPC;OK;HARD;3;RpcSs: Started
[1458911089] SERVICE ALERT: AD-RMAD;Service: Windows Time;OK;HARD;3;W32Time: Started
[1458911089] SERVICE ALERT: AD-RMAD;Uptime;OK;HARD;3;System Uptime - 0 day(s) 17 hour(s) 14 minute(s)
[1458911089] SERVICE ALERT: AD-RMAD;E:\ Drive Space;OK;HARD;3;e:\ - total: 499.87 Gb - used: 338.73 Gb (68%) - free 161.15 Gb (32%)
[1458911090] SERVICE ALERT: AD-RMAD;Memory Usage;OK;HARD;3;Memory usage: total:10239.64 MB - used: 1848.49 MB (18%) - free: 8391.15 MB (82%)
[1458911090] SERVICE ALERT: AD-RMAD;CPU Load;OK;HARD;3;CPU Load 0% (5 min average)
[1458911090] SERVICE ALERT: AD-RMAD;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1458911090] SERVICE ALERT: AD-RMAD;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1458911090] SERVICE ALERT: AD-RMAD;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1458911090] SERVICE ALERT: AD-RMAD;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1458911090] SERVICE ALERT: AD-RMAD;Service: WMI;OK;HARD;3;Winmgmt: Started
[1458911091] SERVICE ALERT: AD-RMAD;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1458911092] SERVICE ALERT: AD-RMAD;C:\ Drive Space;OK;HARD;3;c:\ - total: 79.66 Gb - used: 25.50 Gb (32%) - free 54.16 Gb (68%)
[1458911100] SERVICE ALERT: AD-RMAD;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1458912603] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (32.8% change >= 20.0% threshold)
[1458912603] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1458912603] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1458912603] SERVICE DOWNTIME ALERT: BOEAPPTEST;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1458912610] SERVICE ALERT: ADWSUSSERVERS;Service: WMI;OK;HARD;3;Winmgmt: Started
[1458912614] HOST ALERT: ADWSUSSERVERS;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1458912640] SERVICE ALERT: ADWSUSCLIENTS;E:\ Drive Space;OK;HARD;3;e:\ - total: 299.87 Gb - used: 203.86 Gb (68%) - free 96.01 Gb (32%)
[1458912644] HOST ALERT: ADWSUSCLIENTS;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1458912655] SERVICE ALERT: ADWSUSSERVERS;Service: Windows Time;OK;HARD;3;W32Time: Started
[1458912656] SERVICE ALERT: ADWSUSSERVERS;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1458912658] SERVICE ALERT: ADWSUSCLIENTS;Uptime;OK;HARD;3;System Uptime - 0 day(s) 17 hour(s) 26 minute(s)
[1458912660] SERVICE ALERT: ADWSUSCLIENTS;C:\ Drive Space;OK;HARD;3;c:\ - total: 79.66 Gb - used: 39.88 Gb (50%) - free 39.77 Gb (50%)
[1458912680] SERVICE ALERT: ADWSUSCLIENTS;Memory Usage;OK;HARD;3;Memory usage: total:14335.61 MB - used: 5415.16 MB (38%) - free: 8920.45 MB (62%)
[1458912706] SERVICE ALERT: ADWSUSSERVERS;Uptime;OK;HARD;3;System Uptime - 0 day(s) 17 hour(s) 33 minute(s)
[1458912706] SERVICE ALERT: ADWSUSSERVERS;C:\ Drive Space;OK;HARD;3;c:\ - total: 79.66 Gb - used: 48.50 Gb (61%) - free 31.15 Gb (39%)
[1458912721] SERVICE ALERT: ADWSUSSERVERS;Memory Usage;OK;HARD;3;Memory usage: total:32767.77 MB - used: 6270.20 MB (19%) - free: 26497.57 MB (81%)
[1458912728] SERVICE ALERT: ADWSUSCLIENTS;NSClient++ Version;OK;HARD;3;NSClient++ 0,4,1,73 2012-12-17
[1458912728] SERVICE ALERT: ADWSUSCLIENTS;Service: Windows Time;OK;HARD;3;W32Time: Started
[1458912770] SERVICE ALERT: ADWSUSCLIENTS;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1458912829] SERVICE ALERT: ADWSUSCLIENTS;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1458912867] SERVICE ALERT: ADWSUSSERVERS;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1458912875] SERVICE ALERT: ADWSUSCLIENTS;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1458912879] SERVICE ALERT: ADWSUSSERVERS;CPU Load;OK;HARD;3;CPU Load 17% (5 min average)
[1458912880] SERVICE ALERT: ADWSUSSERVERS;VMWare Tools;OK;HARD;3;VMTools: Started
[1458912894] SERVICE ALERT: ADWSUSSERVERS;NSClient++ Version;OK;HARD;3;NSClient++ 0,4,1,73 2012-12-17
[1458912916] SERVICE ALERT: ADWSUSCLIENTS;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1458912940] SERVICE ALERT: ADWSUSSERVERS;E:\ Drive Space;OK;HARD;3;e:\ - total: 99.87 Gb - used: 62.26 Gb (62%) - free 37.61 Gb (38%)
[1458912946] SERVICE ALERT: ADWSUSSERVERS;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1458912953] SERVICE ALERT: ADWSUSSERVERS;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1458912960] SERVICE ALERT: ADWSUSCLIENTS;Service: RPC;OK;HARD;3;RpcSs: Started
[1458913010] SERVICE ALERT: ADWSUSCLIENTS;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1458913059] SERVICE ALERT: ADWSUSCLIENTS;Service: WMI;OK;HARD;3;Winmgmt: Started
[1458913121] SERVICE ALERT: ADWSUSSERVERS;Service: RPC;OK;HARD;3;RpcSs: Started
[1458913162] SERVICE ALERT: ADWSUSSERVERS;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1458913189] SERVICE ALERT: ADWSUSCLIENTS;CPU Load;OK;HARD;3;CPU Load 5% (5 min average)
[1458913190] SERVICE ALERT: ADWSUSCLIENTS;VMWare Tools;OK;HARD;3;VMTools: Started
[1458914073] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1458914193] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1458914313] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1458914913] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1458916113] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 63% (5 min average)
[1458917152] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;OK;HARD;3;c:\ - total: 84.90 Gb - used: 75.77 Gb (89%) - free 9.13 Gb (11%)
[1458918513] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458918633] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1458918753] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1458919953] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 43% (5 min average)
[1458922353] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458922473] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 96% (5 min average)
[1458922593] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1458923193] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 63% (5 min average)
[1458926792] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 91% (5 min average)
[1458926912] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 64% (5 min average)
[1458928712] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1458928832] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1458928952] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 90% (5 min average)
[1458930152] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 56% (5 min average)
[1458930752] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458930872] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 95% (5 min average)
[1458930992] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1458932192] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 69% (5 min average)
[1458932792] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1458932912] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 96% (5 min average)
[1458933032] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1458934832] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 94% (5 min average)
[1458935432] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1458936002] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (42.9% change >= 20.0% threshold)
[1458936002] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1458936002] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1458936002] SERVICE DOWNTIME ALERT: BOEAPPTEST;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1458936018] SERVICE ALERT: ADAUTH-RMAD;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1458936022] HOST ALERT: ADAUTH-RMAD;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1458936070] SERVICE ALERT: ADAUTH-RMAD;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1458936112] SERVICE ALERT: ADAUTH-RMAD;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1458936160] SERVICE ALERT: ADAUTH-RMAD;Service: RPC;OK;HARD;3;RpcSs: Started
[1458936208] SERVICE ALERT: ADAUTH-RMAD;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1458936308] SERVICE ALERT: ADAUTH-RMAD;Service: WMI;OK;HARD;3;Winmgmt: Started
[1458936343] SERVICE ALERT: ADAUTH-RMAD;Uptime;OK;HARD;3;System Uptime - 9 day(s) 6 hour(s) 53 minute(s)
[1458936386] SERVICE ALERT: ADAUTH-RMAD;VMWare Tools;OK;HARD;3;VMTools: Started
[1458936389] SERVICE ALERT: ADAUTH-RMAD;C:\ Drive Space;OK;HARD;3;c:\ - total: 79.66 Gb - used: 24.47 Gb (31%) - free 55.19 Gb (69%)
[1458936418] SERVICE ALERT: ADAUTH-RMAD;Service: Windows Time;OK;HARD;3;W32Time: Started
[1458936431] SERVICE ALERT: ADAUTH-RMAD;CPU Load;OK;HARD;3;CPU Load 0% (5 min average)
[1458936478] SERVICE ALERT: ADAUTH-RMAD;Memory Usage;OK;HARD;3;Memory usage: total:10239.64 MB - used: 1819.39 MB (18%) - free: 8420.25 MB (82%)
[1458936528] SERVICE ALERT: ADAUTH-RMAD;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1458936548] SERVICE ALERT: ADAUTH-RMAD;E:\ Drive Space;OK;HARD;3;e:\ - total: 199.87 Gb - used: 99.62 Gb (50%) - free 100.25 Gb (50%)
[1458936578] SERVICE ALERT: ADAUTH-RMAD;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1458936632] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 88% (5 min average)
[1458937232] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1458937352] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 96% (5 min average)
[1458937472] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 93% (5 min average)
[1458938072] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 49% (5 min average)
[1458939272] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458939392] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 84% (5 min average)
[1458939463] SERVICE ALERT: NETENG-PA;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1458939583] SERVICE ALERT: NETENG-PA;CPU Load;OK;SOFT;2;CPU Load 4% (5 min average)
[1458940592] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1458940712] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 95% (5 min average)
[1458940832] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1458941432] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 71% (5 min average)
[1458942032] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458942152] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1458942272] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 93% (5 min average)
[1458942872] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1458945272] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 92% (5 min average)
[1458945872] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 96% (5 min average)
[1458948400] SERVICE ALERT: NETIDADPROXY;Service: DNS Client;CRITICAL;SOFT;1;Dnscache: Stopped
[1458948520] SERVICE ALERT: NETIDADPROXY;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1458950672] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 92% (5 min average)
[1458951272] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1458954216] SERVICE ALERT: TARGETX-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1458954336] SERVICE ALERT: TARGETX-2012;CPU Load;OK;SOFT;2;CPU Load 11% (5 min average)
[1458955305] SERVICE ALERT: SANDBOX-DB;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1458955425] SERVICE ALERT: SANDBOX-DB;CPU Load;OK;SOFT;2;CPU Load 11% (5 min average)
[1458958687] SERVICE ALERT: CBORD;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458958787] SERVICE ALERT: CBORD;Memory Usage;OK;SOFT;2;Memory usage: total:10943.58 MB - used: 3414.35 MB (31%) - free: 7529.23 MB (69%)
[1458960628] SERVICE ALERT: ADWSUSSERVERS;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458960728] SERVICE ALERT: ADWSUSSERVERS;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1458961427] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 94% (5 min average)
[1458961547] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 92% (5 min average)
[1458961667] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 92% (5 min average)
[1458962267] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 95% (5 min average)
[1458962672] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1458962867] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 83% (5 min average)
[1458963467] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1458963587] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 61% (5 min average)