[1457928008] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1457928008] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1457928245] SERVICE ALERT: SP2010DB;C:\ Drive Space;OK;HARD;3;c:\ - total: 69.90 Gb - used: 55.71 Gb (80%) - free 14.19 Gb (20%)
[1457935895] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 84.90 Gb - used: 78.74 Gb (93%) - free 6.16 Gb (7%)
[1457939406] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1457942200] SERVICE ALERT: SP2010;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457942320] SERVICE ALERT: SP2010;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1457943605] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 85% (5 min average)
[1457944205] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457944325] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1457944445] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 97% (5 min average)
[1457944983] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457945103] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 91% (5 min average)
[1457945223] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 92% (5 min average)
[1457946245] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 90% (5 min average)
[1457946423] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1457946845] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 85% (5 min average)
[1457946845] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1457948223] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 86% (5 min average)
[1457948645] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457948764] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 97% (5 min average)
[1457948823] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1457948882] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 39% (5 min average)
[1457948943] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 85% (5 min average)
[1457949481] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1457949543] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1457949601] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 90% (5 min average)
[1457949663] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1457949721] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 78% (5 min average)
[1457949783] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 83% (5 min average)
[1457950321] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457950383] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1457950441] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1457950503] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1457950562] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1457950623] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 96% (5 min average)
[1457950623] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1457951160] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 93% (5 min average)
[1457951223] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1457951761] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1457952423] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 78% (5 min average)
[1457953023] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1457953143] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 65% (5 min average)
[1457954160] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 87% (5 min average)
[1457954760] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1457954880] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 80% (5 min average)
[1457954943] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457955063] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1457955184] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1457955782] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 68% (5 min average)
[1457955995] SERVICE ALERT: porterpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1457956055] SERVICE ALERT: porterpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1457956080] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457956115] SERVICE ALERT: porterpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1457956200] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1457956320] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1457956716] SERVICE ALERT: porterpk;Device Status;OK;HARD;3;Open top cover.
[1457956716] SERVICE ALERT: studentcenterpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1457958720] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 42% (5 min average)
[1457959920] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 91% (5 min average)
[1457959945] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1457960040] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 56% (5 min average)
[1457960065] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1457960184] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1457960579] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1457960583] HOST ALERT: hallefirst2;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.46 ms
[1457960885] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 29% - OK!
[1457960999] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1457961384] SERVICE ALERT: BOEAPPTEST;CPU Load;OK;HARD;3;CPU Load 25% (5 min average)
[1457961840] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457961960] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1457962079] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1457963184] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1457963279] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 71% (5 min average)
[1457963304] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;SOFT;2;CPU Load 97% (5 min average)
[1457963424] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1457963879] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457963991] SERVICE ALERT: INTLDB;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1457963999] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1457964000] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Ready
[1457964024] SERVICE ALERT: BOEAPPTEST;CPU Load;OK;HARD;3;CPU Load 82% (5 min average)
[1457964024] SERVICE FLAPPING ALERT: BOEAPPTEST;CPU Load;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1457964111] SERVICE ALERT: INTLDB;CPU Load;OK;SOFT;2;CPU Load 8% (5 min average)
[1457964119] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1457964624] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1457964744] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1457964864] SERVICE ALERT: BOEAPPTEST;CPU Load;OK;SOFT;3;CPU Load 25% (5 min average)
[1457965295] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 84.90 Gb - used: 75.57 Gb (89%) - free 9.33 Gb (11%)
[1457965319] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 84% (5 min average)
[1457966519] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457966639] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1457966759] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 90% (5 min average)
[1457967135] SERVICE ALERT: rooseveltpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1457967182] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1457967195] SERVICE ALERT: rooseveltpk;Device Status;OK;SOFT;2;MENUS  SHOW ME HOW  RETRIEVE JOB  INFORMATION
[1457967359] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1457969158] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 82% (5 min average)
[1457969758] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457969878] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1457969998] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1457970598] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1457970755] SERVICE ALERT: marshallpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1457970815] SERVICE ALERT: marshallpk;Device Status;OK;SOFT;2;Open tray indicated.
[1457971198] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1457971363] SERVICE ALERT: ADWSUSCLIENTS;Service: Group Policy;CRITICAL;SOFT;1;gpsvc: Stopped
[1457971396] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 12072;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1457971415] SERVICE ALERT: ADWSUSSERVERS;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.102.133 and port 12489: Connection refused
[1457971415] SERVICE ALERT: ADWSUSSERVERS;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.102.133 and port 12489: Connection refused
[1457971427] HOST ALERT: ADWSUSSERVERS;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1457971461] HOST ALERT: ADWSUSSERVERS;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.39 ms
[1457971483] SERVICE ALERT: ADWSUSCLIENTS;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1457971495] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 12072;OK;SOFT;2;Eventlog check ok
[1457971535] SERVICE ALERT: ADWSUSSERVERS;C:\ Drive Space;OK;SOFT;2;c:\ - total: 79.66 Gb - used: 48.52 Gb (61%) - free 31.13 Gb (39%)
[1457971535] SERVICE ALERT: ADWSUSSERVERS;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1457971798] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (5 min average)
[1457972397] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457972518] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1457972636] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1457973941] SERVICE ALERT: hallegroundcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1457973941] SERVICE ALERT: hallegroundcolor;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1457973971] HOST ALERT: hallegroundcolor;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1457974001] SERVICE ALERT: hallegroundcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1457974001] SERVICE ALERT: hallegroundcolor;Cyan Cartridge HP CE261A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1457974001] SERVICE ALERT: hallegroundcolor;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1457974020] HOST ALERT: hallegroundcolor;UP;SOFT;2;PING OK - Packet loss = 75%, RTA = 0.45 ms
[1457974055] SERVICE ALERT: hallegroundcolor;Device Status;OK;SOFT;3;LOAD TRAY 1 PLAIN LETTER 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1457974056] SERVICE ALERT: hallegroundcolor;Image Fuser Kit HP 110V-CE246A;OK;SOFT;3;Image Fuser Kit HP 110V-CE246A, 220V-CE247A is at 72% - OK!
[1457974056] SERVICE ALERT: hallegroundcolor;Cyan Cartridge HP CE261A;OK;SOFT;2;Cyan Cartridge HP CE261A is at 95% - OK!
[1457974162] SERVICE ALERT: ADDC1;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457974175] SERVICE ALERT: ADDC1;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457974176] SERVICE ALERT: ADDC1;Service: Key Distribution Center;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457974262] SERVICE ALERT: ADDC1;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1457974275] SERVICE ALERT: ADDC1;CPU Load;OK;SOFT;2;CPU Load 1% (5 min average)
[1457974276] SERVICE ALERT: ADDC1;Service: Key Distribution Center;OK;SOFT;2;KDC: Started
[1457974295] SERVICE ALERT: ADDC1;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1457974415] SERVICE ALERT: ADDC1;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1457974826] SERVICE ALERT: ADDC1;Service: Remote Management;CRITICAL;SOFT;1;connect to address 164.76.99.235 and port 12489: Connection refused
[1457974943] SERVICE ALERT: ADDC1;Service: Remote Management;OK;SOFT;2;WinRM: Started
[1457975663] SERVICE FLAPPING ALERT: BOEAPPTEST;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1457977393] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1457977453] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1457977513] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1457978712] SERVICE ALERT: eateriespk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1457978997] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1457979057] SERVICE ALERT: hallefirst2;Device Status;OK;SOFT;2;Clearing paper path
[1457982064] SERVICE ALERT: NETENG-PA;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457982184] SERVICE ALERT: NETENG-PA;CPU Load;OK;SOFT;2;CPU Load 3% (5 min average)
[1457983936] SERVICE ALERT: strongpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1457983995] SERVICE ALERT: strongpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1457984036] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1457984056] SERVICE ALERT: strongpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1457985421] SERVICE ALERT: markjeffersonpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1457985481] SERVICE ALERT: markjeffersonpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1457985541] SERVICE ALERT: markjeffersonpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1457985591] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1457985599] SERVICE ALERT: markjeffersonpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1457985651] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1457985659] SERVICE ALERT: markjeffersonpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1457985711] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1457985719] SERVICE ALERT: markjeffersonpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1457985856] SERVICE ALERT: strongpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1457986135] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1457986306] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 9% - OK!
[1457986313] SERVICE ALERT: markjeffersonpk;Device Status;OK;HARD;3;TRAY 2 EMPTY PLAIN LETTER   Ready
[1457988145] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1457988205] SERVICE ALERT: halleground2;Device Status;OK;SOFT;2;MENUS  RETRIEVE JOB  INFORMATION  PAPER HANDLING
[1457988805] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1457988865] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1457988925] SERVICE ALERT: halleground2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1457990131] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1457990131] SERVICE ALERT: halleground2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1457990147] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1457990161] HOST ALERT: halleground2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1457990185] SERVICE ALERT: halleground2;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! All Trays Still Full
[1457990189] HOST ALERT: halleground2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.62 ms
[1457990201] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 15% - OK!
[1457990725] SERVICE ALERT: halleground2;Device Status;OK;HARD;3;Ready
[1457992736] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1457992796] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1457992856] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1457994251] SERVICE ALERT: owen101lab3;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1457994311] SERVICE ALERT: owen101lab3;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1457994372] SERVICE ALERT: owen101lab3;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1457994971] SERVICE ALERT: owen101lab3;Device Status;OK;HARD;3;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1457995225] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1457995286] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1457995345] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1457997275] SERVICE ALERT: hallesouthcommons;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1457997336] SERVICE ALERT: hallesouthcommons;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1457997396] SERVICE ALERT: hallesouthcommons;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1457998596] SERVICE ALERT: hallesouthcommons;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1457998945] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;CRITICAL;HARD;3;Black Print Cartridge HP C8543X is at 1% - CRITICAL!
[1457999725] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1457999785] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1457999845] SERVICE ALERT: halleground2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1458004440] SERVICE ALERT: SP2010DB;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 69.90 Gb - used: 56.00 Gb (80%) - free 13.90 Gb (20%)
[1458004505] SERVICE ALERT: IT-2012;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1458004560] SERVICE ALERT: SP2010DB;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 69.90 Gb - used: 56.00 Gb (80%) - free 13.90 Gb (20%)
[1458004625] SERVICE ALERT: IT-2012;CPU Load;OK;SOFT;2;CPU Load 84% (5 min average)
[1458004680] SERVICE ALERT: SP2010DB;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 69.90 Gb - used: 56.00 Gb (80%) - free 13.90 Gb (20%)
[1458005280] SERVICE ALERT: SP2010DB;C:\ Drive Space;OK;HARD;3;c:\ - total: 69.90 Gb - used: 55.28 Gb (79%) - free 14.62 Gb (21%)
[1458006835] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 66% (5 min average)
[1458007435] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1458007555] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1458007675] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1458008245] SERVICE ALERT: halleground2;Device Status;OK;HARD;3;Warming up
[1458008275] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 74% (5 min average)
[1458009475] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1458009595] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1458009715] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1458009715] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1458009868] SERVICE ALERT: ADWSUSSERVERS;Service: User Profile;CRITICAL;SOFT;1;connect to address 164.76.102.133 and port 12489: No route to host
[1458009871] HOST ALERT: ADWSUSSERVERS;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.133)
[1458009878] SERVICE ALERT: ADWSUSSERVERS;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.102.133 and port 12489: No route to host
[1458009878] SERVICE ALERT: ADWSUSSERVERS;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.102.133 and port 12489: No route to host
[1458009878] SERVICE ALERT: ADWSUSSERVERS;E:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.102.133 and port 12489: No route to host
[1458009889] HOST ALERT: ADWSUSSERVERS;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1458009976] SERVICE ALERT: ADWSUSSERVERS;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1458009995] SERVICE ALERT: ADWSUSSERVERS;Service: RPC;OK;SOFT;2;RpcSs: Started
[1458009995] SERVICE ALERT: ADWSUSSERVERS;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1458009997] SERVICE ALERT: ADWSUSSERVERS;E:\ Drive Space;OK;SOFT;2;e:\ - total: 99.87 Gb - used: 58.56 Gb (59%) - free 41.31 Gb (41%)