[1457672433] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457672553] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1457672673] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 75% (5 min average)
[1457673273] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1457673393] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1457673513] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 97% (5 min average)
[1457676463] SERVICE ALERT: ADRODC-LIV;Service: PXE Boot;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457676563] SERVICE ALERT: ADRODC-LIV;Service: PXE Boot;OK;SOFT;2;WDSServer: Started
[1457676707] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 84.90 Gb - used: 78.75 Gb (93%) - free 6.15 Gb (7%)
[1457677113] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 93% (5 min average)
[1457677713] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1457680517] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1457680637] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 78% (5 min average)
[1457681359] SERVICE ALERT: EMUFILES;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457681479] SERVICE ALERT: EMUFILES;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1457684837] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1457684957] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 96% (5 min average)
[1457685077] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 75% (5 min average)
[1457686277] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1457686397] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 89% (5 min average)
[1457688513] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 37% (5 min average)
[1457691513] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457691633] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 91% (5 min average)
[1457691753] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1457691790] SERVICE ALERT: OFFICEKMS;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457691910] SERVICE ALERT: OFFICEKMS;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1457692030] SERVICE ALERT: OFFICEKMS;CPU Load;OK;SOFT;3;CPU Load 5% (5 min average)
[1457692353] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 93% (5 min average)
[1457692953] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1457694153] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 68% (5 min average)
[1457695353] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457695473] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1457695593] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1457695997] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457696117] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 53% (5 min average)
[1457698483] SERVICE ALERT: mckennypk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1457698498] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1457698558] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1457698593] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 42% (5 min average)
[1457698618] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1457700883] SERVICE ALERT: mckennypk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1457700942] SERVICE ALERT: mckennypk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1457701003] SERVICE ALERT: mckennypk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1457701593] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1457701617] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1457701713] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1457701833] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1457702717] SERVICE ALERT: IMAGES;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457702747] SERVICE ALERT: IMAGES;MDT Monitor;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457702747] HOST ALERT: IMAGES;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1457702796] SERVICE ALERT: IMAGES;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457702826] HOST ALERT: IMAGES;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1457702837] SERVICE ALERT: IMAGES;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457702840] SERVICE ALERT: IMAGES;Service: PXE Boot;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457702841] HOST ALERT: IMAGES;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1457702851] SERVICE ALERT: IMAGES;MDT Monitor;OK;SOFT;2;MDT_Monitor: Started
[1457702855] SERVICE ALERT: IMAGES;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457702896] SERVICE ALERT: IMAGES;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1457702937] SERVICE ALERT: IMAGES;CPU Load;OK;SOFT;3;CPU Load 3% (5 min average)
[1457702939] SERVICE ALERT: IMAGES;Service: PXE Boot;OK;SOFT;2;WDSServer: Started
[1457702955] SERVICE ALERT: IMAGES;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1457703633] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 44% (5 min average)
[1457703741] HOST ALERT: ADAUTH-RMAD;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1457704608] SERVICE ALERT: NESSUS;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1457704728] SERVICE ALERT: NESSUS;CPU Load;OK;SOFT;2;CPU Load 3% (5 min average)
[1457705201] SERVICE ALERT: mckennypk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1457706465] SERVICE ALERT: IMAGES;MDT Monitor;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457706565] SERVICE ALERT: IMAGES;MDT Monitor;OK;SOFT;2;MDT_Monitor: Started
[1457706577] SERVICE ALERT: ADAUTH-RMAD;C:\ Drive Space;OK;HARD;3;c:\ - total: 79.66 Gb - used: 31.96 Gb (40%) - free 47.70 Gb (60%)
[1457706586] SERVICE ALERT: ADAUTH-RMAD;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1457706674] SERVICE ALERT: ADAUTH-RMAD;CPU Load;OK;HARD;3;CPU Load 0% (5 min average)
[1457706674] SERVICE ALERT: ADAUTH-RMAD;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 11 minute(s)
[1457706674] SERVICE ALERT: ADAUTH-RMAD;Memory Usage;OK;HARD;3;Memory usage: total:10239.64 MB - used: 1163.01 MB (11%) - free: 9076.62 MB (89%)
[1457706674] SERVICE ALERT: ADAUTH-RMAD;Service: Windows Time;OK;HARD;3;W32Time: Started
[1457706674] SERVICE ALERT: ADAUTH-RMAD;Service: WMI;OK;HARD;3;Winmgmt: Started
[1457706674] SERVICE ALERT: ADAUTH-RMAD;VMWare Tools;OK;HARD;3;VMTools: Started
[1457706674] SERVICE ALERT: ADAUTH-RMAD;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1457706674] SERVICE ALERT: ADAUTH-RMAD;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1457706674] SERVICE ALERT: ADAUTH-RMAD;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1457706674] SERVICE ALERT: ADAUTH-RMAD;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1457706674] SERVICE ALERT: ADAUTH-RMAD;Service: RPC;OK;HARD;3;RpcSs: Started
[1457706674] SERVICE ALERT: ADAUTH-RMAD;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1457706705] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 84.90 Gb - used: 75.58 Gb (89%) - free 9.32 Gb (11%)
[1457708432] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1457708450] SERVICE ALERT: INTLTESTDB;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 84.22 Gb - used: 70.80 Gb (84%) - free 13.42 Gb (16%)
[1457708552] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 45% (5 min average)
[1457708570] SERVICE ALERT: INTLTESTDB;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 84.22 Gb - used: 70.80 Gb (84%) - free 13.42 Gb (16%)
[1457708690] SERVICE ALERT: INTLTESTDB;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 84.22 Gb - used: 70.80 Gb (84%) - free 13.42 Gb (16%)
[1457709152] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1457709272] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 96% (5 min average)
[1457709392] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 94% (5 min average)
[1457709392] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1457709486] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1457709992] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1457710403] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (25.7% change >= 20.0% threshold)
[1457710592] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 71% (5 min average)
[1457711192] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1457711312] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 93% (5 min average)
[1457711432] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 73% (5 min average)
[1457712032] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1457712152] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 77% (5 min average)
[1457712857] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1457712918] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1457712977] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1457713307] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1457713368] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1457713428] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1457713641] SERVICE ALERT: C4SAS-SE;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457713761] SERVICE ALERT: C4SAS-SE;CPU Load;OK;SOFT;2;CPU Load 26% (5 min average)
[1457713952] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457714072] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1457714192] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1457714757] SERVICE ALERT: hallefirst3;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1457714817] SERVICE ALERT: hallefirst3;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1457714877] SERVICE ALERT: hallefirst3;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1457715391] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 48% (5 min average)
[1457715992] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1457716111] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 87% (5 min average)
[1457716677] SERVICE ALERT: hallefirst3;Device Status;OK;HARD;3;Clearing paper path
[1457716709] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457716830] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1457716949] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1457718747] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 75% (5 min average)
[1457719347] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 91% (5 min average)
[1457719467] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1457719587] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1457720177] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1457720187] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 90% (5 min average)
[1457720787] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 96% (5 min average)
[1457721387] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 68% (5 min average)
[1457722587] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 91% (5 min average)
[1457722707] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 58% (5 min average)
[1457722836] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (50.7% change >= 20.0% threshold)
[1457722934] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (50.7% change >= 20.0% threshold)
[1457723015] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (50.7% change >= 20.0% threshold)
[1457724507] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457724627] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1457724696] SERVICE ALERT: INTLTESTDB;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457724747] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1457724816] SERVICE ALERT: INTLTESTDB;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1457725347] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 90% (5 min average)
[1457725947] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1457726547] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 90% (5 min average)
[1457727147] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 97% (5 min average)
[1457727226] SERVICE ALERT: crossroadspk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1457728347] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 53% (5 min average)
[1457728947] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457728949] SERVICE ALERT: MAGICDBPROD;Service: Group Policy;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1457729015] SERVICE ALERT: MAGICDBPROD;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457729045] HOST ALERT: MAGICDBPROD;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1457729063] SERVICE ALERT: MAGICDBPROD;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457729065] SERVICE ALERT: MAGICAPPTEST;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457729067] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 97% (5 min average)
[1457729089] SERVICE ALERT: MAGICDBPROD;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457729093] HOST ALERT: MAGICDBPROD;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1457729095] HOST ALERT: MAGICAPPTEST;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1457729109] SERVICE ALERT: MAGICDBPROD;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457729135] SERVICE ALERT: MAGICDBPROD;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457729139] HOST ALERT: MAGICDBPROD;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1457729144] SERVICE ALERT: MAGICAPPTEST;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457729174] HOST ALERT: MAGICAPPTEST;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1457729183] SERVICE ALERT: MAGICDBPROD;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457729185] SERVICE ALERT: MAGICAPPTEST;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457729187] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1457729193] SERVICE ALERT: MAGICDBPROD;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457729193] SERVICE ALERT: MAGICAPPTEST;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457729209] SERVICE ALERT: MAGICDBPROD;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457729213] HOST ALERT: MAGICDBPROD;DOWN;SOFT;4;(Host check timed out after 30.00 seconds)
[1457729215] HOST ALERT: MAGICAPPTEST;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1457729229] SERVICE ALERT: MAGICDBPROD;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457729243] SERVICE ALERT: MAGICDBPROD;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457729243] SERVICE ALERT: MAGICAPPTEST;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457729255] SERVICE ALERT: MAGICDBPROD;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457729259] HOST ALERT: MAGICDBPROD;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1457729264] SERVICE ALERT: MAGICAPPTEST;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457729273] HOST ALERT: MAGICAPPTEST;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1457729283] SERVICE ALERT: MAGICDBPROD;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457729285] SERVICE ALERT: MAGICAPPTEST;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457729285] SERVICE ALERT: MAGICDBPROD;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457729303] SERVICE ALERT: MAGICDBPROD;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457729305] SERVICE ALERT: MAGICAPPTEST;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457729309] SERVICE ALERT: MAGICAPPTEST;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457729313] SERVICE ALERT: MAGICDBPROD;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457729313] SERVICE ALERT: MAGICAPPTEST;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457729313] HOST ALERT: MAGICDBPROD;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1457729315] HOST ALERT: MAGICAPPTEST;DOWN;SOFT;5;(Host check timed out after 30.00 seconds)
[1457729333] SERVICE ALERT: MAGICAPPTEST;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457729334] SERVICE ALERT: MAGICDBPROD;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457729349] SERVICE ALERT: MAGICDBPROD;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457729363] SERVICE ALERT: MAGICDBPROD;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457729363] SERVICE ALERT: MAGICAPPTEST;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457729363] HOST ALERT: MAGICAPPTEST;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1457729364] HOST ALERT: MAGICDBPROD;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1457729379] SERVICE ALERT: MAGICDBPROD;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457729383] SERVICE ALERT: MAGICAPPTEST;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457729384] SERVICE ALERT: MAGICAPPTEST;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457729403] SERVICE ALERT: MAGICDBPROD;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457729405] SERVICE ALERT: MAGICAPPTEST;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457729405] SERVICE ALERT: MAGICDBPROD;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457729409] HOST ALERT: MAGICDBPROD;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1457729413] HOST ALERT: MAGICAPPTEST;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1457729424] SERVICE ALERT: MAGICDBPROD;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457729424] SERVICE ALERT: MAGICDBPROD;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457729429] SERVICE ALERT: MAGICAPPTEST;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457729433] SERVICE ALERT: MAGICAPPTEST;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457729433] SERVICE ALERT: MAGICDBPROD;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457729433] SERVICE ALERT: MAGICAPPTEST;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457729453] SERVICE ALERT: MAGICAPPTEST;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457729454] SERVICE ALERT: MAGICDBPROD;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457729454] HOST ALERT: MAGICDBPROD;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1457729459] HOST ALERT: MAGICAPPTEST;DOWN;SOFT;8;(Host check timed out after 30.00 seconds)
[1457729474] SERVICE ALERT: MAGICDBPROD;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457729475] SERVICE ALERT: MAGICAPPTEST;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457729483] SERVICE ALERT: MAGICDBPROD;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457729483] SERVICE ALERT: MAGICAPPTEST;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457729499] SERVICE ALERT: MAGICDBPROD;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457729503] SERVICE ALERT: MAGICAPPTEST;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457729504] HOST ALERT: MAGICDBPROD;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1457729505] HOST ALERT: MAGICAPPTEST;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1457729519] SERVICE ALERT: MAGICDBPROD;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457729520] SERVICE ALERT: MAGICAPPTEST;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457729523] SERVICE ALERT: MAGICDBPROD;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457729525] SERVICE ALERT: MAGICAPPTEST;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457729525] SERVICE ALERT: MAGICDBPROD;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457729544] SERVICE ALERT: MAGICDBPROD;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457729544] SERVICE ALERT: MAGICDBPROD;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457729549] SERVICE ALERT: MAGICAPPTEST;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457729550] HOST ALERT: MAGICAPPTEST;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1457729553] SERVICE ALERT: MAGICAPPTEST;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457729568] SERVICE ALERT: MAGICAPPTEST;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457729573] SERVICE ALERT: MAGICAPPTEST;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457729574] SERVICE ALERT: MAGICDBPROD;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457729594] SERVICE ALERT: MAGICDBPROD;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457729595] SERVICE ALERT: MAGICAPPTEST;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457729615] SERVICE ALERT: MAGICAPPTEST;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457729619] SERVICE ALERT: MAGICDBPROD;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457729619] SERVICE ALERT: MAGICAPPTEST;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457729623] SERVICE ALERT: MAGICAPPTEST;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457729639] SERVICE ALERT: MAGICDBPROD;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457729640] SERVICE ALERT: MAGICAPPTEST;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457729664] SERVICE ALERT: MAGICDBPROD;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457729664] SERVICE ALERT: MAGICDBPROD;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457729673] SERVICE ALERT: MAGICAPPTEST;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457729688] SERVICE ALERT: MAGICAPPTEST;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457729714] SERVICE ALERT: MAGICDBPROD;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457729715] SERVICE ALERT: MAGICAPPTEST;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457729735] SERVICE ALERT: MAGICAPPTEST;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457729739] SERVICE ALERT: MAGICAPPTEST;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457729759] SERVICE ALERT: MAGICDBPROD;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457729760] SERVICE ALERT: MAGICAPPTEST;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457729808] SERVICE ALERT: MAGICAPPTEST;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457729855] SERVICE ALERT: MAGICAPPTEST;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457729859] SERVICE ALERT: MAGICAPPTEST;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457731587] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 83% (5 min average)
[1457732787] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1457732907] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 96% (5 min average)
[1457733027] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1457733627] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 79% (5 min average)
[1457733725] SERVICE ALERT: INTLTESTAPP;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1457733845] SERVICE ALERT: INTLTESTAPP;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1457734227] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457734347] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 96% (5 min average)
[1457734467] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1457737467] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 80% (5 min average)
[1457738067] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457738187] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 96% (5 min average)
[1457738307] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1457741907] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1457742507] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1457748090] SERVICE ALERT: EVENTSTEST-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457748210] SERVICE ALERT: EVENTSTEST-2012;CPU Load;OK;SOFT;2;CPU Load 81% (5 min average)
[1457748312] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1457748432] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 60% (5 min average)
[1457748795] SERVICE ALERT: IT-2012;CPU Load;WARNING;SOFT;1;CPU Load 94% (5 min average)
[1457748915] SERVICE ALERT: IT-2012;CPU Load;OK;SOFT;2;CPU Load 82% (5 min average)
[1457748945] SERVICE ALERT: TARGETX-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1457749065] SERVICE ALERT: TARGETX-2012;CPU Load;CRITICAL;SOFT;2;CPU Load 97% (5 min average)
[1457749185] SERVICE ALERT: TARGETX-2012;CPU Load;OK;SOFT;3;CPU Load 83% (5 min average)
[1457751498] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1457751822] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1457751882] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1457751943] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1457752098] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1457752542] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;TONER COLLECTION UNIT LOW 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1457753605] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10042;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1457753605] HOST ALERT: ADWSUSSERVERS;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.133)
[1457753669] HOST ALERT: ADWSUSSERVERS;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.45 ms
[1457753713] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10042;OK;SOFT;2;Eventlog check ok
[1457755632] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1457755752] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 92% (5 min average)
[1457755872] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 95% (5 min average)
[1457756472] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 85% (5 min average)
[1457756898] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 74% (5 min average)
[1457757072] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 94% (5 min average)
[1457757192] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 82% (5 min average)
[1457757498] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1457757618] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1457757738] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)