[1478318439] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1478318559] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1478318679] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 42% (15 min average)
[1478320479] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478320599] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 44% (15 min average)
[1478322672] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1478322793] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;2;Warning:  idm1  are more than 120 seconds behind!
[1478322913] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1478322999] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478323119] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1478323513] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1478323719] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478323800] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1478323839] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1478323959] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 55% (15 min average)
[1478324113] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1478324559] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1478324677] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1478324712] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1478324712] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1478324924] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 364;CRITICAL;SOFT;1;Windows Server Update Services(error, 364, success)[Content file download failed.
[1478325043] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 364;CRITICAL;SOFT;2;Windows Server Update Services(error, 364, success)[Content file download failed.
[1478325163] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 364;CRITICAL;HARD;3;Windows Server Update Services(error, 364, success)[Content file download failed.
[1478325312] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1478325432] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;2;Warning:  idm1  are more than 120 seconds behind!
[1478325552] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1478325877] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1478325997] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1478326117] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1478326152] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1478326752] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1478326872] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1478326992] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1478327317] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1478327592] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1478327736] SERVICE ALERT: EVENTSTEST-2012;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1478327856] SERVICE ALERT: EVENTSTEST-2012;CPU Load;OK;SOFT;2;CPU Load 88% (15 min average)
[1478328792] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1478328912] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1478329032] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1478329601] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10032;CRITICAL;SOFT;1;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1478329632] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1478329722] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10032;CRITICAL;SOFT;2;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1478329841] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10032;CRITICAL;HARD;3;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1478329974] SERVICE ALERT: epprogsql2;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478330094] SERVICE ALERT: epprogsql2;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1478330214] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1478330233] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1478330832] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1478330952] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1478331072] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1478331517] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1478332272] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1478332362] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 364;OK;HARD;3;Eventlog check ok
[1478332614] SERVICE ALERT: epprogsql2;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1478332717] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478332837] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1478332957] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 94% (15 min average)
[1478333814] SERVICE ALERT: epprogsql2;CPU Load;OK;HARD;3;CPU Load 48% (15 min average)
[1478334157] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478334277] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 85% (15 min average)
[1478334603] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (63.2% change >= 20.0% threshold)
[1478334603] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1478334603] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1478334603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1478334603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1478334603] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1478334603] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1478334603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1478334603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1478334603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1478334603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1478334603] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1478334603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1478334603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1478334603] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1478334603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1478334603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1478334603] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1478336402] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (54.5% change >= 20.0% threshold)
[1478336402] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1478336402] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1478336402] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1478336402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1478336402] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1478336402] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1478336402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1478336402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1478336402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1478336402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1478336402] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1478336402] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1478336402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1478336402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1478336402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1478336402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1478336402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1478336442] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10032;OK;HARD;3;Eventlog check ok
[1478338477] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1478338597] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1478338707] SERVICE ALERT: hallegroundcolor;Image Transfer Kit HP CE249A;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Image Transfer Kit HP CE249A'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1478338760] SERVICE ALERT: hallegroundcolor;Image Transfer Kit HP CE249A;OK;SOFT;2;Image Transfer Kit HP CE249A is at 55% - OK!
[1478341597] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1478341717] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1478341800] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service has exited from a period of scheduled downtime
[1478341837] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1478343037] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1478343072] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1478343637] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1478343637] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1478344032] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 108 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478344152] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478346637] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (15 min average)
[1478347837] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1478347957] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 15% (15 min average)
[1478350840] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10032;CRITICAL;SOFT;1;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1478350960] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10032;CRITICAL;SOFT;2;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1478351081] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10032;CRITICAL;HARD;3;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1478352602] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1478352602] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1478352602] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1478352602] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1478352602] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1478352602] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1478352602] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1478352602] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1478352602] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1478352602] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1478352602] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1478352602] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1478352602] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1478352602] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1478352602] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1478352602] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1478352659] SERVICE ALERT: ADFS;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1478352663] HOST ALERT: ADFS;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1478352672] SERVICE ALERT: ADFS;VMWare Tools;OK;HARD;3;VMTools: Started
[1478352672] SERVICE ALERT: ADFS;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1478352672] SERVICE ALERT: ADFS;Service: Windows Time;OK;HARD;3;W32Time: Started
[1478352756] SERVICE ALERT: ADFS;Uptime;OK;HARD;3;System Uptime - 0 day(s) 22 hour(s) 41 minute(s)
[1478352756] SERVICE ALERT: ADFS;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1478352922] SERVICE ALERT: ADFS;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1478352923] SERVICE ALERT: ADFS;C:\ Drive Space;OK;HARD;3;c:\ - total: 79.51 Gb - used: 18.55 Gb (23%) - free 60.96 Gb (77%)
[1478353011] SERVICE ALERT: ADFS;CPU Load;OK;HARD;3;CPU Load 1% (15 min average)
[1478353011] SERVICE ALERT: ADFS;Service: RPC;OK;HARD;3;RpcSs: Started
[1478353097] SERVICE ALERT: ADFS;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1478353097] SERVICE ALERT: ADFS;Memory Usage;OK;HARD;3;Memory usage: total:10239.64 MB - used: 1655.40 MB (16%) - free: 8584.24 MB (84%)
[1478353183] SERVICE ALERT: ADFS;Service: WMI;OK;HARD;3;Winmgmt: Started
[1478353184] SERVICE ALERT: ADFS;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1478358281] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10032;OK;HARD;3;Eventlog check ok
[1478358757] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1478358877] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1478358997] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1478361397] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 65% (15 min average)
[1478361774] SERVICE ALERT: CAVA2;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1478361894] SERVICE ALERT: CAVA2;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1478361996] HOST ALERT: hallefirstcolor;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1478361997] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1478362014] SERVICE ALERT: CAVA2;CPU Load;OK;SOFT;3;CPU Load 8% (15 min average)
[1478362066] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;OK;HARD;3;Yellow Cartridge HP CE262A is at 85% - OK!
[1478362076] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;OK;HARD;3;Magenta Cartridge HP CE263A is at 70% - OK!
[1478362117] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 47% (15 min average)
[1478362177] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1478362217] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;OK;HARD;3;Image Transfer Kit HP CE249A is at 69% - OK!
[1478362240] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;OK;HARD;3;Black Cartridge HP CE260X is at 52% - OK!
[1478362282] SERVICE ALERT: hallefirstcolor;Toner Collection Unit HP CE265A;OK;HARD;3;
[1478362553] SERVICE ALERT: hallefirstcolor;Device Status;OK;HARD;3;TONER COLLECTION UNIT FULL 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1478362573] SERVICE ALERT: hallefirstcolor;Image Fuser Kit HP 110V-CE246A;OK;HARD;3;Image Fuser Kit HP 110V-CE246A, 220V-CE247A is at 100% - OK!
[1478362592] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;OK;HARD;3;Cyan Cartridge HP CE261A is at 68% - OK!
[1478362717] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478362837] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1478362957] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 88% (15 min average)
[1478363403] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1478363403] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1478363403] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1478363403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1478363403] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1478363403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1478363403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1478363403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1478363403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1478363403] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1478363403] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1478363403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1478363403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1478363403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1478363403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1478363403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1478363462] SERVICE ALERT: ADFS;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478363465] HOST ALERT: ADFS;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.57)
[1478363475] SERVICE ALERT: ADFS;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478363475] SERVICE ALERT: ADFS;VMWare Tools;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478363478] SERVICE ALERT: ADFS;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478363481] HOST ALERT: ADFS;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.4.57)
[1478363544] HOST ALERT: ADFS;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.4.57)
[1478363559] SERVICE ALERT: ADFS;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478363559] SERVICE ALERT: ADFS;Uptime;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478363562] HOST ALERT: ADFS;DOWN;SOFT;4;CRITICAL - Host Unreachable (164.76.4.57)
[1478363582] SERVICE ALERT: ADFS;Service: Group Policy;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478363585] HOST ALERT: ADFS;DOWN;SOFT;5;CRITICAL - Host Unreachable (164.76.4.57)
[1478363595] SERVICE ALERT: ADFS;VMWare Tools;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478363595] SERVICE ALERT: ADFS;Service: Windows Time;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478363595] SERVICE ALERT: ADFS;Service: DNS Client;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478363679] SERVICE ALERT: ADFS;Service: Firewall;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478363682] HOST ALERT: ADFS;DOWN;SOFT;6;CRITICAL - Host Unreachable (164.76.4.57)
[1478363688] SERVICE ALERT: ADFS;Uptime;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478363702] SERVICE ALERT: ADFS;Service: Group Policy;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478363705] HOST ALERT: ADFS;DOWN;SOFT;7;CRITICAL - Host Unreachable (164.76.4.57)
[1478363715] SERVICE ALERT: ADFS;Service: Windows Time;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478363715] SERVICE ALERT: ADFS;Service: DNS Client;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478363715] SERVICE ALERT: ADFS;VMWare Tools;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478363725] SERVICE ALERT: ADFS;Service: NetLogin;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478363725] SERVICE ALERT: ADFS;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478363728] HOST ALERT: ADFS;DOWN;SOFT;8;CRITICAL - Host Unreachable (164.76.4.57)
[1478363791] HOST ALERT: ADFS;DOWN;SOFT;9;CRITICAL - Host Unreachable (164.76.4.57)
[1478363799] SERVICE ALERT: ADFS;Uptime;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478363799] SERVICE ALERT: ADFS;Service: Firewall;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478363814] SERVICE ALERT: ADFS;CPU Load;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478363814] SERVICE ALERT: ADFS;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478363817] HOST ALERT: ADFS;DOWN;HARD;10;CRITICAL - Host Unreachable (164.76.4.57)
[1478363845] SERVICE ALERT: ADFS;Service: NetLogin;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478363845] SERVICE ALERT: ADFS;C:\ Drive Space;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478363900] SERVICE ALERT: ADFS;Service: User Profile;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478363900] SERVICE ALERT: ADFS;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478363934] SERVICE ALERT: ADFS;Service: RPC;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478363934] SERVICE ALERT: ADFS;CPU Load;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478363965] SERVICE ALERT: ADFS;Service: NetLogin;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478363965] SERVICE ALERT: ADFS;C:\ Drive Space;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478363986] SERVICE ALERT: ADFS;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478363986] SERVICE ALERT: ADFS;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478364020] SERVICE ALERT: ADFS;Service: User Profile;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478364020] SERVICE ALERT: ADFS;Memory Usage;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478364054] SERVICE ALERT: ADFS;Service: RPC;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478364054] SERVICE ALERT: ADFS;CPU Load;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478364106] SERVICE ALERT: ADFS;NSClient++ Version;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478364106] SERVICE ALERT: ADFS;Service: WMI;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478364140] SERVICE ALERT: ADFS;Memory Usage;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478364140] SERVICE ALERT: ADFS;Service: User Profile;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478364157] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1478364226] SERVICE ALERT: ADFS;NSClient++ Version;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478364226] SERVICE ALERT: ADFS;Service: WMI;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478364277] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 56% (15 min average)
[1478365192] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Warming up
[1478365196] HOST ALERT: hallefirst2;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.52 ms
[1478365248] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 9% - OK!
[1478365254] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1478365477] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1478365597] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1478365717] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1478366317] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 24% (15 min average)
[1478366317] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1478366917] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478367035] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 74% (15 min average)
[1478367957] SERVICE ALERT: ADRODC-LIV;Service: DFS Namespace;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1478367987] HOST ALERT: ADRODC-LIV;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1478368046] SERVICE ALERT: ADRODC-LIV;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1478368076] HOST ALERT: ADRODC-LIV;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1478368077] SERVICE ALERT: ADRODC-LIV;Service: DFS Namespace;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1478368105] HOST ALERT: ADRODC-LIV;UP;SOFT;3;PING WARNING - Packet loss = 82%, RTA = 6.02 ms
[1478368146] SERVICE ALERT: ADRODC-LIV;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1478368177] SERVICE ALERT: ADRODC-LIV;Service: DFS Namespace;OK;SOFT;3;Dfs: Started
[1478368362] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 364;CRITICAL;SOFT;1;Windows Server Update Services(error, 364, success)[Content file download failed.
[1478368483] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 364;CRITICAL;SOFT;2;Windows Server Update Services(error, 364, success)[Content file download failed.
[1478368602] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 364;CRITICAL;HARD;3;Windows Server Update Services, eventlog: 1 > critical
[1478368835] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478368955] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 83% (15 min average)
[1478372681] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10032;CRITICAL;SOFT;1;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1478372800] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10032;CRITICAL;SOFT;2;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1478372920] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10032;CRITICAL;HARD;3;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1478373513] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1478373572] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1478373632] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1478374232] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;Ready
[1478375203] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 364;OK;HARD;3;Eventlog check ok
[1478376155] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1478376275] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 47% (15 min average)
[1478377050] SERVICE ALERT: eateriespk;Device Status;OK;HARD;3;Clearing paper path
[1478378075] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1478380120] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10032;OK;HARD;3;Eventlog check ok
[1478383376] SERVICE ALERT: cava2;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1478383496] SERVICE ALERT: cava2;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1478383616] SERVICE ALERT: cava2;CPU Load;OK;SOFT;3;CPU Load 8% (15 min average)
[1478384675] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1478384795] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1478384915] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 82% (15 min average)
[1478392115] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478392235] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 78% (15 min average)
[1478392835] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478392955] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 86% (15 min average)
[1478394520] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10032;CRITICAL;SOFT;1;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1478394641] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10032;CRITICAL;SOFT;2;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1478394760] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10032;CRITICAL;HARD;3;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1478395954] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1478396074] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1478396194] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1478397993] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1478399193] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1478399313] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 82% (15 min average)
[1478399548] SERVICE ALERT: INTLTESTDATA;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1478399648] SERVICE ALERT: INTLTESTDATA;VMWare Tools;OK;SOFT;2;VMTools: Started
[1478401360] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10032;OK;HARD;3;Eventlog check ok
[1478401713] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478401832] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1478401951] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1478402550] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1478402550] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)