[1477022517] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1477022637] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1477022757] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477023957] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 88% (15 min average)
[1477024557] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1477024677] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1477024797] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477025397] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 88% (15 min average)
[1477025997] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1477026117] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1477026237] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 80% (15 min average)
[1477026744] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1477026862] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1477027800] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1477028037] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1477028063] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1477028157] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 83% (15 min average)
[1477028183] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;2;Warning:  idm1  are more than 120 seconds behind!
[1477028302] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1477028757] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1477028877] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 69% (15 min average)
[1477028902] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1477030102] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1477030222] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;2;Warning:  idm1  are more than 120 seconds behind!
[1477030342] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1477030677] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1477030797] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1477030917] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1477030943] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1477030943] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1477031542] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1477032142] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1477033342] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1477033838] SERVICE ALERT: epprogsql2;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1477033942] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1477033958] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1477034078] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477034542] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1477035142] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1477035262] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;2;Warning:  idm1  are more than 120 seconds behind!
[1477035382] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1477036478] SERVICE ALERT: epprogsql2;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477037182] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1477037513] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 72% (15 min average)
[1477037678] SERVICE ALERT: epprogsql2;CPU Load;OK;HARD;3;CPU Load 67% (15 min average)
[1477038603] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (54.9% change >= 20.0% threshold)
[1477038603] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1477038603] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1477038603] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1477038603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1477038603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1477038603] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1477038603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1477038603] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1477038603] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1477038603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1477038603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1477038603] HOST DOWNTIME ALERT: VUM;STARTED; Host has entered a period of scheduled downtime
[1477038603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1477038603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1477038603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1477038603] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1477038603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1477038603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1477040513] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1477040633] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1477040753] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 80% (15 min average)
[1477041953] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1477042073] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 90% (15 min average)
[1477045673] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1477045793] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1477045800] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service has exited from a period of scheduled downtime
[1477045913] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477046513] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477047381] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1477048330] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 191 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477048450] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477049918] HOST ALERT: hallefirst1;UP;HARD;10;PING OK - Packet loss = 72%, RTA = 0.60 ms
[1477050026] HOST ALERT: hallefirst3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1477050090] HOST ALERT: hallefirst3;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.46 ms
[1477050113] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1477050454] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;Ready
[1477050494] SERVICE ALERT: hallefirst1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1477050965] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 35% - OK!
[1477051313] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1477051433] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1477051449] HOST ALERT: hallefirst2;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 1.22 ms
[1477051553] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1477051553] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1477051571] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1477051601] HOST ALERT: hallefirst1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1477051621] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1477051631] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1477051659] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1477051661] HOST ALERT: hallefirst1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1477051691] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1477051699] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1477051719] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1477051721] HOST ALERT: hallefirst1;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1477051759] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1477051769] HOST ALERT: hallefirst2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1477051779] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1477051789] HOST ALERT: hallefirst1;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1477051819] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1477051833] HOST ALERT: hallefirst2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1477051849] HOST ALERT: hallefirst1;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1477051939] HOST ALERT: hallefirst1;UP;SOFT;6;PING WARNING - Packet loss = 96%, RTA = 0.55 ms
[1477052153] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 70% (15 min average)
[1477052227] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1477052281] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 7% - OK!
[1477052285] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 35% - OK!
[1477052333] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1477052340] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1477052373] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;Ready
[1477052413] SERVICE ALERT: hallefirst1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1477052753] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1477052873] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1477052993] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 54% (15 min average)
[1477057915] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1477060193] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1477060313] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1477060989] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5346 found since last run. oneLevelSearchOps OK. searchOps is too high! 5785 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1477061109] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477061513] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1477061515] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1477061575] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1477061633] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 70% (15 min average)
[1477061634] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1477061914] SERVICE ALERT: mckennypk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1477061973] SERVICE ALERT: mckennypk;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! All Trays Still Full
[1477062909] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5579 found since last run. oneLevelSearchOps OK. searchOps is too high! 5861 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1477063029] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477063629] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5310 found since last run. oneLevelSearchOps OK. searchOps is too high! 5498 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1477063749] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477064033] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1477064035] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1477064373] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1477064433] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1477064493] SERVICE ALERT: hallefirst1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1477064649] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5872 found since last run. oneLevelSearchOps OK. searchOps is too high! 5965 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1477064769] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477065233] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1477065353] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1477065369] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5661 found since last run. oneLevelSearchOps OK. searchOps is too high! 5760 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1477065473] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477065489] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477065603] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1477065603] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1477065603] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1477065603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1477065603] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1477065603] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1477065603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1477065603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1477065603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1477065603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1477065603] HOST DOWNTIME ALERT: VUM;STARTED; Host has entered a period of scheduled downtime
[1477065603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1477065603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1477065603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1477065603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1477065603] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1477065603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1477065693] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1477066084] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1477066089] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5640 found since last run. oneLevelSearchOps OK. searchOps is too high! 5760 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1477066144] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 35% - OK!
[1477066149] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5302 found since last run. oneLevelSearchOps OK. searchOps is too high! 5575 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1477066209] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477066269] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477066293] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1477066353] SERVICE ALERT: hallefirst1;Device Status;OK;SOFT;2;Warming up
[1477066673] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 59% (15 min average)
[1477066809] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7170 found since last run. oneLevelSearchOps OK. searchOps is too high! 7493 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1477066929] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477067143] SERVICE ALERT: studentcenterpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1477067204] SERVICE ALERT: studentcenterpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1477067263] SERVICE ALERT: studentcenterpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1477067273] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1477067393] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1477067403] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1477067403] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1477067403] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1477067403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1477067403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1477067403] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1477067403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1477067403] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1477067403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1477067403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1477067403] HOST DOWNTIME ALERT: VUM;STARTED; Host has entered a period of scheduled downtime
[1477067403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1477067403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1477067403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1477067403] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1477067403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1477067403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1477067469] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6973 found since last run. oneLevelSearchOps OK. searchOps is too high! 7178 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1477067513] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1477067589] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477067747] SERVICE ALERT: eateriespk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1477067807] SERVICE ALERT: eateriespk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1477067867] SERVICE ALERT: eateriespk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1477068189] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5568 found since last run. oneLevelSearchOps OK. searchOps is too high! 5708 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1477068309] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477068713] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 72% (15 min average)
[1477068713] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1477068729] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5274 found since last run. oneLevelSearchOps OK. searchOps is too high! 5444 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1477068849] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477069063] SERVICE ALERT: studentcenterpk;Device Status;OK;HARD;3;paper. Close tray and  right door. If  this jam reoccurs,
[1477069067] SERVICE ALERT: eateriespk;Device Status;OK;HARD;3;TRAY 4 EMPTY PLAIN LETTER   Ready
[1477069449] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5070 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1477069543] HOST ALERT: mckennypk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1477069569] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477069625] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1477069633] HOST ALERT: mckennypk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1477069646] SERVICE ALERT: mckennypk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1477069663] SERVICE ALERT: studentcenterpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1477069676] HOST ALERT: mckennypk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1477069685] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1477069706] SERVICE ALERT: mckennypk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1477069715] HOST ALERT: mckennypk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1477069723] SERVICE ALERT: studentcenterpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1477069745] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1477069766] SERVICE ALERT: mckennypk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1477069775] HOST ALERT: mckennypk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1477069779] SERVICE ALERT: mckennypk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1477069783] SERVICE ALERT: studentcenterpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1477069809] HOST ALERT: mckennypk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1477069839] SERVICE ALERT: mckennypk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1477069869] HOST ALERT: mckennypk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1477069899] SERVICE ALERT: mckennypk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1477069929] HOST ALERT: mckennypk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1477069964] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1477070019] HOST ALERT: mckennypk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1477070023] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 100% - OK!
[1477070109] HOST ALERT: mckennypk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1477072569] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7414 found since last run. oneLevelSearchOps OK. searchOps is too high! 7599 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1477072686] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1477072689] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477072716] HOST ALERT: hallefirst2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1477072739] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1477072744] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1477072746] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1477072769] HOST ALERT: hallefirst2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1477072783] SERVICE ALERT: studentcenterpk;Device Status;OK;HARD;3;50 72 65 73 73 20 17 20 74 6F 20 73 65 65 20 73  74 65 70  by step information
[1477072783] SERVICE FLAPPING ALERT: studentcenterpk;Device Status;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1477072804] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1477072806] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1477072834] HOST ALERT: hallefirst2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1477072864] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1477072894] HOST ALERT: hallefirst2;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1477072984] HOST ALERT: hallefirst2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1477073074] HOST ALERT: hallefirst2;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1477073107] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5380 found since last run. oneLevelSearchOps OK. searchOps is too high! 5670 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1477073164] HOST ALERT: hallefirst2;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1477073227] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477073254] HOST ALERT: hallefirst2;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1477073289] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5065 found since last run. oneLevelSearchOps OK. searchOps is too high! 5158 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1477073344] HOST ALERT: hallefirst2;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1477073409] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477073434] HOST ALERT: hallefirst2;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1477073514] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1477073633] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1477073753] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477074427] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5088 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1477074547] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477075553] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477076153] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477077125] HOST ALERT: mckennypk;UP;HARD;10;PING WARNING - Packet loss = 81%, RTA = 1.84 ms
[1477077353] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477077538] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 65% - OK!
[1477077560] SERVICE ALERT: mckennypk;Device Status;OK;HARD;3;Ready
[1477077693] SERVICE ALERT: mckennypk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1477077953] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477078209] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 106 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477078329] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477078343] SERVICE ALERT: C4SAS-SE;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1477078463] SERVICE ALERT: C4SAS-SE;CPU Load;OK;SOFT;2;CPU Load 77% (15 min average)
[1477079064] SERVICE ALERT: C4SAS-SE;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1477079183] SERVICE ALERT: C4SAS-SE;CPU Load;OK;SOFT;2;CPU Load 85% (15 min average)
[1477079529] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 105 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477079649] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477079753] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477080849] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 112 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477080953] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477080969] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477081569] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 111 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477081689] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477081783] SERVICE FLAPPING ALERT: studentcenterpk;Device Status;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1477082153] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477082347] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5080 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1477082467] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477082754] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477083353] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477083489] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 106 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477083609] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477083667] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 122 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477083787] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477083953] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477084553] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477084987] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 115 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477085107] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477085153] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477085754] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477086131] SERVICE ALERT: IT-2012;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1477086251] SERVICE ALERT: IT-2012;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1477086353] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477086371] SERVICE ALERT: IT-2012;CPU Load;OK;SOFT;3;CPU Load 54% (15 min average)
[1477086451] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;CRITICAL;SOFT;1;Black Cartridge HP CE260X is at 1% - CRITICAL!
[1477086511] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;CRITICAL;SOFT;2;Black Cartridge HP CE260X is at 1% - CRITICAL!
[1477086572] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;CRITICAL;HARD;3;Black Cartridge HP CE260X is at 1% - CRITICAL!
[1477086971] SERVICE ALERT: IT-2012;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1477087091] SERVICE ALERT: IT-2012;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1477087211] SERVICE ALERT: IT-2012;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1477087380] HOST ALERT: banprod;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1477087434] SERVICE ALERT: BOEAPP;Service: Group Policy;CRITICAL;SOFT;1;gpsvc: Stopped
[1477087470] HOST ALERT: banprod;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1477087534] HOST ALERT: banprod;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1477087553] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477087554] SERVICE ALERT: BOEAPP;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1477087811] SERVICE ALERT: IT-2012;CPU Load;OK;HARD;3;CPU Load 77% (15 min average)
[1477087813] SERVICE ALERT: NETIDAD;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.4.232 and port 12489: Connection refused
[1477087833] SERVICE ALERT: NETIDAD;Active Directory Federation Services;CRITICAL;SOFT;1;adfssrv: Stopped
[1477087924] SERVICE ALERT: NETIDAD;C:\ Drive Space;OK;SOFT;2;c:\ - total: 79.66 Gb - used: 15.29 Gb (19%) - free 64.37 Gb (81%)
[1477087953] SERVICE ALERT: NETIDAD;Active Directory Federation Services;OK;SOFT;2;adfssrv: Started
[1477088050] SERVICE ALERT: ADRODC-LIV;Service: Distributed File System Replication;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1477088050] SERVICE ALERT: ADRODC-LIV;Service: Key Distribution Center;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1477088050] SERVICE ALERT: ADRODC-LIV;Service: AD Domain Service;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1477088058] SERVICE ALERT: ADRODC-LIV;Service: Intersite Messaging;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1477088063] SERVICE ALERT: ADRODC-LIV;Service: WMI;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1477088068] SERVICE ALERT: ADRODC-DET;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477088084] SERVICE ALERT: ADRODC-DET;Service: MDT Monitor;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477088098] HOST ALERT: ADRODC-DET;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1477088110] SERVICE ALERT: ADRODC-LIV;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477088118] SERVICE ALERT: ADRODC-LIV;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477088124] SERVICE ALERT: ADRODC-DET;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477088140] HOST ALERT: ADRODC-LIV;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1477088154] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477088154] HOST ALERT: ADRODC-DET;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1477088163] SERVICE ALERT: ADRODC-DET;AD Replication;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 30 seconds.
[1477088171] SERVICE ALERT: ADRODC-DET;Service: Firewall;CRITICAL;SOFT;2;connect to address 198.109.113.140 and port 12489: Connection refused
[1477088186] HOST ALERT: ADRODC-DET;UP;SOFT;3;PING OK - Packet loss = 79%, RTA = 9.04 ms
[1477088189] SERVICE ALERT: ADRODC-LIV;Service: Key Distribution Center;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1477088189] SERVICE ALERT: ADRODC-LIV;Service: Distributed File System Replication;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1477088190] SERVICE ALERT: ADRODC-LIV;Service: AD Domain Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1477088198] SERVICE ALERT: ADRODC-LIV;Service: Intersite Messaging;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1477088199] SERVICE ALERT: ADRODC-LIV;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477088199] SERVICE ALERT: ADRODC-LIV;Service: PXE Boot;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477088200] SERVICE ALERT: ADRODC-LIV;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477088203] SERVICE ALERT: ADRODC-LIV;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1477088204] SERVICE ALERT: ADRODC-DET;Service: MDT Monitor;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1477088213] SERVICE ALERT: ADRODC-DET;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477088219] HOST ALERT: ADRODC-LIV;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1477088224] SERVICE ALERT: ADRODC-DET;Memory Usage;OK;SOFT;2;Memory usage: total:37586.86 MB - used: 2375.00 MB (6%) - free: 35211.86 MB (94%)
[1477088230] SERVICE ALERT: ADRODC-LIV;Service: DFS Namespace;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477088230] SERVICE ALERT: ADRODC-LIV;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1477088235] SERVICE ALERT: ADRODC-LIV;Service: Remote Management;CRITICAL;SOFT;1;connect to address 207.74.163.140 and port 12489: Connection refused
[1477088238] SERVICE ALERT: ADRODC-LIV;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1477088249] HOST ALERT: ADRODC-LIV;UP;SOFT;3;PING OK - Packet loss = 75%, RTA = 5.96 ms
[1477088265] SERVICE ALERT: ADRODC-DET;AD Replication;OK;SOFT;2;OK: replication is up and running
[1477088288] SERVICE ALERT: ADRODC-DET;Service: Firewall;OK;SOFT;3;MpsSvc: Started
[1477088289] SERVICE ALERT: ADRODC-LIV;Service: Distributed File System Replication;OK;SOFT;3;DFSR: Started
[1477088289] SERVICE ALERT: ADRODC-LIV;Service: Key Distribution Center;OK;SOFT;3;KDC: Started
[1477088290] SERVICE ALERT: ADRODC-LIV;Service: AD Domain Service;OK;SOFT;3;NTDS: Started
[1477088298] SERVICE ALERT: ADRODC-LIV;Service: Intersite Messaging;OK;SOFT;3;IsmServ: Started
[1477088299] SERVICE ALERT: ADRODC-LIV;Service: Windows Time;CRITICAL;SOFT;2;W32Time: Starting
[1477088299] SERVICE ALERT: ADRODC-LIV;Service: PXE Boot;OK;SOFT;2;WDSServer: Started
[1477088300] SERVICE ALERT: ADRODC-LIV;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1477088303] SERVICE ALERT: ADRODC-LIV;Service: WMI;OK;SOFT;3;Winmgmt: Started
[1477088304] SERVICE ALERT: ADRODC-DET;Service: MDT Monitor;OK;SOFT;3;MDT_Monitor: Started
[1477088313] SERVICE ALERT: ADRODC-DET;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1477088330] SERVICE ALERT: ADRODC-LIV;Service: DFS Namespace;OK;SOFT;2;Dfs: Started
[1477088330] SERVICE ALERT: ADRODC-LIV;C:\ Drive Space;OK;SOFT;3;c:\ - total: 1115.16 Gb - used: 559.49 Gb (50%) - free 555.67 Gb (50%)
[1477088338] SERVICE ALERT: ADRODC-LIV;NSClient++ Version;OK;SOFT;3;NSClient++ 0.4.4.15 2015-11-25
[1477088346] SERVICE ALERT: ADRODC-LIV;Service: Remote Management;OK;SOFT;2;WinRM: Started
[1477088408] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 103 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477088419] SERVICE ALERT: ADRODC-LIV;Service: Windows Time;OK;SOFT;3;W32Time: Started
[1477088460] SERVICE ALERT: TARGETX-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1477088528] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477088580] SERVICE ALERT: TARGETX-2012;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1477088700] SERVICE ALERT: TARGETX-2012;CPU Load;OK;SOFT;3;CPU Load 3% (15 min average)
[1477088779] SERVICE ALERT: ADDC1;Service: Group Policy;CRITICAL;SOFT;1;gpsvc: Stopped
[1477088899] SERVICE ALERT: ADDC1;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1477089057] SERVICE ALERT: ADDC6;Remote Loader for IDM;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477089075] HOST ALERT: ADDC6;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1477089128] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 114 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477089131] SERVICE ALERT: ADDC6;Service: Distributed File System Replication;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477089131] SERVICE ALERT: ADDC6;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477089133] SERVICE ALERT: ADDC6;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477089145] SERVICE ALERT: ADDC6;Service: Remote Management;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477089154] SERVICE ALERT: ADDC6;Service: DFS Namespace;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477089155] SERVICE ALERT: ADDC6;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477089155] SERVICE ALERT: ADDC6;Service: AD Domain Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477089155] SERVICE ALERT: ADDC6;Service: Key Distribution Center;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477089156] SERVICE ALERT: ADDC6;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477089159] SERVICE ALERT: ADDC6;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477089159] SERVICE ALERT: ADDC6;Service: Intersite Messaging;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477089159] SERVICE ALERT: ADDC6;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477089159] SERVICE ALERT: ADDC6;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477089159] SERVICE ALERT: ADDC6;Service: AD Web Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477089161] HOST ALERT: ADDC6;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1477089163] SERVICE ALERT: ADDC6;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477089164] SERVICE ALERT: ADDC6;Service: DNS Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477089166] SERVICE ALERT: ADDC6;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477089171] SERVICE ALERT: ADDC6;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477089177] SERVICE ALERT: ADDC6;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477089177] SERVICE ALERT: ADDC6;Remote Loader for IDM;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1477089179] SERVICE ALERT: ADDC6;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477089193] HOST ALERT: ADDC6;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1477089213] SERVICE ALERT: ADDC6;AD Replication;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 30 seconds.
[1477089243] HOST ALERT: ADDC6;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1477089248] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477089251] SERVICE ALERT: ADDC6;Service: Distributed File System Replication;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1477089251] SERVICE ALERT: ADDC6;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1477089253] SERVICE ALERT: ADDC6;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1477089265] SERVICE ALERT: ADDC6;Service: Remote Management;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1477089274] SERVICE ALERT: ADDC6;Service: DFS Namespace;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1477089275] SERVICE ALERT: ADDC6;Service: Key Distribution Center;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1477089275] SERVICE ALERT: ADDC6;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1477089275] SERVICE ALERT: ADDC6;Service: AD Domain Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1477089276] SERVICE ALERT: ADDC6;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1477089279] SERVICE ALERT: ADDC6;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1477089279] SERVICE ALERT: ADDC6;Service: Intersite Messaging;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1477089279] SERVICE ALERT: ADDC6;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1477089279] SERVICE ALERT: ADDC6;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1477089279] SERVICE ALERT: ADDC6;Service: AD Web Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1477089280] SERVICE ALERT: ADDC6;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477089281] HOST ALERT: ADDC6;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1477089283] SERVICE ALERT: ADDC6;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477089283] SERVICE ALERT: ADDC6;Service: Dell Backup Agent;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477089283] SERVICE ALERT: ADDC6;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1477089284] SERVICE ALERT: ADDC6;Service: DNS Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1477089286] SERVICE ALERT: ADDC6;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1477089291] SERVICE ALERT: ADDC6;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1477089296] SERVICE ALERT: ADDC6;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1477089297] SERVICE ALERT: ADDC6;Remote Loader for IDM;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1477089299] SERVICE ALERT: ADDC6;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1477089313] HOST ALERT: ADDC6;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1477089333] SERVICE ALERT: ADDC6;AD Replication;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 30 seconds.
[1477089351] SERVICE ALERT: ADDC6;Service: Distributed File System Replication;CRITICAL;HARD;3;connect to address 164.76.7.206 and port 12489: Connection refused
[1477089351] SERVICE ALERT: ADDC6;Service: User Profile;CRITICAL;HARD;3;connect to address 164.76.7.206 and port 12489: Connection refused
[1477089353] SERVICE ALERT: ADDC6;Service: WMI;CRITICAL;HARD;3;connect to address 164.76.7.206 and port 12489: Connection refused
[1477089361] HOST ALERT: ADDC6;UP;SOFT;7;PING WARNING - Packet loss = 82%, RTA = 0.33 ms
[1477089375] SERVICE ALERT: ADDC6;Service: Key Distribution Center;OK;SOFT;3;KDC: Started
[1477089375] SERVICE ALERT: ADDC6;Service: AD Domain Service;OK;SOFT;3;NTDS: Started
[1477089375] SERVICE ALERT: ADDC6;Uptime;OK;SOFT;3;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1477089376] SERVICE ALERT: ADDC6;Service: Firewall;OK;SOFT;3;MpsSvc: Started
[1477089377] SERVICE ALERT: ADDC6;Service: DFS Namespace;OK;SOFT;3;Dfs: Started
[1477089379] SERVICE ALERT: ADDC6;NSClient++ Version;OK;SOFT;3;NSClient++ 0.4.4.15 2015-11-25
[1477089379] SERVICE ALERT: ADDC6;Service: Intersite Messaging;OK;SOFT;3;IsmServ: Started
[1477089379] SERVICE ALERT: ADDC6;Memory Usage;OK;SOFT;3;Memory usage: total:20479.73 MB - used: 888.39 MB (4%) - free: 19591.33 MB (96%)
[1477089379] SERVICE ALERT: ADDC6;C:\ Drive Space;OK;SOFT;3;c:\ - total: 79.66 Gb - used: 35.31 Gb (44%) - free 44.34 Gb (56%)
[1477089379] SERVICE ALERT: ADDC6;Service: AD Web Service;OK;SOFT;3;ADWS: Started
[1477089380] SERVICE ALERT: ADDC6;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1477089383] SERVICE ALERT: ADDC6;Service: Windows Time;CRITICAL;SOFT;2;W32Time: Stopped
[1477089383] SERVICE ALERT: ADDC6;Service: Dell Backup Agent;CRITICAL;SOFT;2;ErdAgent: Starting
[1477089383] SERVICE ALERT: ADDC6;Service: RPC;OK;SOFT;3;RpcSs: Started
[1477089383] SERVICE ALERT: ADDC6;Service: DNS Service;CRITICAL;HARD;3;DNS: Starting
[1477089384] SERVICE ALERT: ADDC6;Service: Remote Management;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1477089386] SERVICE ALERT: ADDC6;VMWare Tools;OK;SOFT;3;VMTools: Started
[1477089391] SERVICE ALERT: ADDC6;Service: NetLogin;OK;SOFT;3;Netlogon: Started
[1477089396] SERVICE ALERT: ADDC6;Service: DNS Client;OK;SOFT;3;Dnscache: Started
[1477089399] SERVICE ALERT: ADDC6;CPU Load;OK;SOFT;3;CPU Load 6% (15 min average)
[1477089436] SERVICE ALERT: ADDC6;AD Replication;OK;SOFT;3;OK: replication is up and running
[1477089503] SERVICE ALERT: ADDC6;Service: Windows Time;CRITICAL;HARD;3;W32Time: Stopped
[1477089503] SERVICE ALERT: ADDC6;Service: Dell Backup Agent;CRITICAL;HARD;3;ErdAgent: Starting
[1477089620] SERVICE ALERT: ADDC2;Service: Remote Management;CRITICAL;SOFT;1;connect to address 164.76.7.202 and port 12489: Connection refused
[1477089620] SERVICE ALERT: ADDC2;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477089624] SERVICE ALERT: ADDC2;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.7.202 and port 12489: Connection refused
[1477089625] SERVICE ALERT: ADDC2;Service: DNS Service;CRITICAL;SOFT;1;connect to address 164.76.7.202 and port 12489: Connection refused
[1477089625] SERVICE ALERT: ADDC2;Service: DFS Namespace;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477089720] SERVICE ALERT: ADDC2;Service: RPC;OK;SOFT;2;RpcSs: Started
[1477089725] SERVICE ALERT: ADDC2;Service: DFS Namespace;OK;SOFT;2;Dfs: Started
[1477089737] SERVICE ALERT: ADDC2;Service: Remote Management;OK;SOFT;2;WinRM: Started
[1477089744] SERVICE ALERT: ADDC2;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1477089745] SERVICE ALERT: ADDC2;Service: DNS Service;OK;SOFT;2;DNS: Started
[1477089877] SERVICE ALERT: ADDC6;Remote Loader for IDM;OK;HARD;3;dirxml_remote.exe: Running
[1477089900] SERVICE ALERT: TARGETX-2012;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1477089951] SERVICE ALERT: ADDC6;Service: Distributed File System Replication;OK;HARD;3;DFSR: Started
[1477089951] SERVICE ALERT: ADDC6;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1477089953] SERVICE ALERT: ADDC6;Service: WMI;OK;HARD;3;Winmgmt: Started
[1477089964] SERVICE ALERT: ADDC6;Service: Remote Management;OK;HARD;3;WinRM: Started
[1477089983] SERVICE ALERT: ADDC6;Service: DNS Service;OK;HARD;3;DNS: Started
[1477090020] SERVICE ALERT: TARGETX-2012;CPU Load;OK;SOFT;2;CPU Load 87% (15 min average)
[1477090103] SERVICE ALERT: ADDC6;Service: Windows Time;OK;HARD;3;W32Time: Started
[1477090103] SERVICE ALERT: ADDC6;Service: Dell Backup Agent;OK;HARD;3;ErdAgent: Started
[1477090276] SERVICE ALERT: ADDC4;Service: Group Policy;CRITICAL;SOFT;1;gpsvc: Stopped
[1477090323] SERVICE ALERT: ADDC4;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.7.204 and port 12489: Connection refused
[1477090371] SERVICE ALERT: ADDC5;AD Group Changes;WARNING;SOFT;1;WARNING:  0  groups have changed!
[1477090381] SERVICE ALERT: SNAP;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1477090384] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.138)
[1477090396] SERVICE ALERT: ADDC4;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1477090416] SERVICE ALERT: SNAP;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1477090419] HOST ALERT: SNAP;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.4.138)
[1477090425] SERVICE ALERT: ADDC5;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477090434] SERVICE ALERT: ADDC4;Memory Usage;OK;SOFT;2;Memory usage: total:20479.73 MB - used: 1049.21 MB (5%) - free: 19430.52 MB (95%)
[1477090443] HOST ALERT: SNAP;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1477090464] SERVICE ALERT: hallesecondsouth;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1477090480] SERVICE ALERT: ADDC1;AD Replication;WARNING;SOFT;1;WARNING: Replication with ADDC5 failed!
[1477090489] SERVICE ALERT: SNAP;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1477090493] SERVICE ALERT: ADDC5;AD Group Changes;OK;SOFT;2;WARNING:  0  groups have changed!
[1477090494] HOST ALERT: hallesecondsouth;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1477090525] SERVICE ALERT: ADDC5;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1477090533] SERVICE ALERT: SNAP;C:\ Drive Space;OK;SOFT;2;c:\ - total: 1116.65 Gb - used: 37.40 Gb (3%) - free 1079.25 Gb (97%)
[1477090546] SERVICE ALERT: hallesecondsouth;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1477090576] HOST ALERT: hallesecondsouth;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1477090580] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1477090605] SERVICE ALERT: ADDC1;AD Replication;OK;SOFT;2;OK: replication is up and running
[1477090606] SERVICE ALERT: hallesecondsouth;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1477090610] HOST ALERT: hallesecondsouth;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1477090639] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1477090666] SERVICE ALERT: hallesecondsouth;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1477090669] HOST ALERT: hallesecondsouth;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1477090699] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1477090729] HOST ALERT: hallesecondsouth;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1477090819] HOST ALERT: hallesecondsouth;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1477090909] HOST ALERT: hallesecondsouth;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1477090999] HOST ALERT: hallesecondsouth;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1477091088] HOST ALERT: hallesecondsouth;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1477091178] HOST ALERT: hallesecondsouth;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1477091195] SERVICE ALERT: ITRF;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477091204] SERVICE ALERT: INTLAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1477091225] HOST ALERT: ITRF;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1477091252] HOST ALERT: ITRF;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1477091275] SERVICE ALERT: ADAUTHDC3-NEW;Service: NetLogin;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1477091295] SERVICE ALERT: ITRF;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1477091296] SERVICE ALERT: ADAUTHDC3-NEW;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.4.208 and port 12489: Connection refused
[1477091296] SERVICE ALERT: ADAUTHDC3-NEW;Service: DFS Namespace;CRITICAL;SOFT;1;connect to address 164.76.4.208 and port 12489: Connection refused
[1477091298] SERVICE ALERT: ADAUTHDC3-NEW;Service: DNS Service;CRITICAL;SOFT;1;connect to address 164.76.4.208 and port 12489: Connection refused
[1477091323] SERVICE ALERT: ADAUTHDC3-NEW;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477091323] SERVICE ALERT: ADAUTHDC3-NEW;Service: Dell Backup Agent;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477091323] SERVICE ALERT: ADAUTHDC3-NEW;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477091323] SERVICE ALERT: ADAUTHDC3-NEW;Service: Remote Management;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477091324] SERVICE ALERT: INTLAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1477091325] SERVICE ALERT: ADAUTHDC3-NEW;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477091395] SERVICE ALERT: ADAUTHDC3-NEW;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1477091400] SERVICE ALERT: INTLDB;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477091416] SERVICE ALERT: ADAUTHDC3-NEW;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1477091416] SERVICE ALERT: ADAUTHDC3-NEW;Service: DFS Namespace;OK;SOFT;2;Dfs: Started
[1477091418] SERVICE ALERT: ADAUTHDC3-NEW;Service: DNS Service;OK;SOFT;2;DNS: Started
[1477091423] SERVICE ALERT: ADAUTHDC3-NEW;Service: Dell Backup Agent;CRITICAL;SOFT;2;ErdAgent: Starting
[1477091423] SERVICE ALERT: ADAUTHDC3-NEW;VMWare Tools;OK;SOFT;2;VMTools: Started
[1477091423] SERVICE ALERT: ADAUTHDC3-NEW;Service: Remote Management;OK;SOFT;2;WinRM: Started
[1477091423] SERVICE ALERT: ADAUTHDC3-NEW;Service: RPC;OK;SOFT;2;RpcSs: Started
[1477091425] SERVICE ALERT: ADAUTHDC3-NEW;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1477091445] SERVICE ALERT: INTLAPP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477091500] SERVICE ALERT: INTLDB;Service: RPC;OK;SOFT;2;RpcSs: Started
[1477091543] SERVICE ALERT: ADAUTHDC3-NEW;Service: Dell Backup Agent;OK;SOFT;3;ErdAgent: Started
[1477091605] SERVICE ALERT: INTLDB;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1477091725] SERVICE ALERT: INTLDB;CPU Load;OK;SOFT;2;CPU Load 78% (15 min average)
[1477091753] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477092048] SERVICE ALERT: ADAUTHDC3-NEW;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1477092168] SERVICE ALERT: ADAUTHDC3-NEW;CPU Load;OK;SOFT;2;CPU Load 88% (15 min average)
[1477092353] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477092603] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (62.4% change >= 20.0% threshold)
[1477092603] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1477092603] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1477092603] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1477092603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1477092603] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1477092603] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1477092603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1477092603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1477092603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1477092603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1477092603] HOST DOWNTIME ALERT: VUM;STARTED; Host has entered a period of scheduled downtime
[1477092603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1477092603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1477092603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1477092603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1477092603] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1477092603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1477092644] SERVICE ALERT: INTLAPP;CPU Load;OK;HARD;3;CPU Load 8% (15 min average)
[1477092734] SERVICE ALERT: EVENTS-2012;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1477092735] SERVICE ALERT: ADRODC-WEMU;Service: Group Policy;CRITICAL;SOFT;1;gpsvc: Stopped
[1477092771] SERVICE ALERT: ADRODC-WEMU;Service: Windows Time;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1477092797] SERVICE ALERT: ADRODC-WEMU;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477092798] SERVICE ALERT: ADRODC-WEMU;Service: Remote Management;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477092798] SERVICE ALERT: ADRODC-WEMU;Service: Distributed File System Replication;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477092821] SERVICE ALERT: ADRODC-WEMU;Service: DNS Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477092851] HOST ALERT: ADRODC-WEMU;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1477092854] SERVICE ALERT: EVENTS-2012;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1477092875] SERVICE ALERT: ADRODC-WEMU;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1477092905] HOST ALERT: ADRODC-WEMU;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1477092911] SERVICE ALERT: ADRODC-WEMU;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1477092917] SERVICE ALERT: ADRODC-WEMU;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1477092918] SERVICE ALERT: ADRODC-WEMU;Service: Remote Management;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1477092918] SERVICE ALERT: ADRODC-WEMU;Service: Distributed File System Replication;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1477092941] SERVICE ALERT: ADRODC-WEMU;Service: DNS Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1477092941] HOST ALERT: ADRODC-WEMU;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1477092953] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477092974] SERVICE ALERT: EVENTS-2012;CPU Load;OK;SOFT;3;CPU Load 93% (15 min average)
[1477092983] SERVICE ALERT: C4SAS-SE;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1477092995] SERVICE ALERT: ADRODC-WEMU;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1477092999] HOST ALERT: ADRODC-WEMU;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.54 ms
[1477093012] SERVICE ALERT: ADRODC-WEMU;Service: Windows Time;OK;SOFT;3;W32Time: Started
[1477093017] SERVICE ALERT: ADRODC-WEMU;Service: WMI;OK;SOFT;3;Winmgmt: Started
[1477093018] SERVICE ALERT: ADRODC-WEMU;Service: Remote Management;OK;SOFT;3;WinRM: Started
[1477093018] SERVICE ALERT: ADRODC-WEMU;Service: Distributed File System Replication;OK;SOFT;3;DFSR: Started
[1477093041] SERVICE ALERT: ADRODC-WEMU;Service: DNS Service;OK;SOFT;3;DNS: Started
[1477093103] SERVICE ALERT: C4SAS-SE;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1477093121] SERVICE ALERT: epprogsql2;Service: Group Policy;CRITICAL;SOFT;1;gpsvc: Stopped
[1477093223] SERVICE ALERT: C4SAS-SE;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477093224] SERVICE ALERT: AX-PROD;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.7.144 and port 12489: Connection refused
[1477093224] SERVICE ALERT: AX-PROD;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.7.144 and port 12489: Connection refused
[1477093241] SERVICE ALERT: epprogsql2;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1477093344] SERVICE ALERT: AX-PROD;Memory Usage;OK;SOFT;2;Memory usage: total:20199.72 MB - used: 1459.80 MB (7%) - free: 18739.93 MB (93%)
[1477093344] SERVICE ALERT: AX-PROD;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1477093553] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477093565] SERVICE ALERT: EMUFILES;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1477093575] SERVICE ALERT: ADRODC-WEMU;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1477093685] SERVICE ALERT: EMUFILES;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1477093726] SERVICE ALERT: EMUPRINTPS;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477093726] SERVICE ALERT: EMUPRINTPS;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477093805] SERVICE ALERT: EMUFILES;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477093826] SERVICE ALERT: EMUPRINTPS;Service: RPC;OK;SOFT;2;RpcSs: Started
[1477093826] SERVICE ALERT: EMUPRINTPS;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1477093960] SERVICE ALERT: C4SAS-SE;C:\ Drive Space;UNKNOWN;SOFT;1;NSClient - ERROR: Invalid return from command: check_drivesize
[1477094009] SERVICE ALERT: C4SAS-SE;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.158.242 and port 12489: Connection refused
[1477094080] SERVICE ALERT: C4SAS-SE;C:\ Drive Space;OK;SOFT;2;c:\ - total: 241.76 Gb - used: 56.38 Gb (23%) - free 185.38 Gb (77%)
[1477094121] SERVICE ALERT: C4SAS-SE;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1477094403] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (62.4% change >= 20.0% threshold)
[1477094403] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1477094403] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1477094403] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1477094403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1477094403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1477094403] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1477094403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1477094403] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1477094403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1477094403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1477094403] HOST DOWNTIME ALERT: VUM;STARTED; Host has entered a period of scheduled downtime
[1477094403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1477094403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1477094403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1477094403] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1477094403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1477094403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1477094423] SERVICE ALERT: C4SAS-SE;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477095005] SERVICE ALERT: EMUFILES;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477095225] SERVICE ALERT: CSCFILESERV;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.4.42 and port 12489: No route to host
[1477095228] HOST ALERT: CSCFILESERV;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.42)
[1477095262] SERVICE ALERT: CSCFILESERV;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.4.42 and port 12489: No route to host
[1477095265] SERVICE ALERT: CSCFILESERV;Uptime;CRITICAL;SOFT;1;connect to address 164.76.4.42 and port 12489: No route to host
[1477095265] HOST ALERT: CSCFILESERV;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.4.42)
[1477095268] SERVICE ALERT: CSCFILESERV;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.4.42 and port 12489: No route to host
[1477095287] SERVICE ALERT: EMUFILES;C:\ Drive Space;UNKNOWN;SOFT;1;NSClient - ERROR: Invalid return from command: check_drivesize
[1477095306] SERVICE ALERT: CSCFILESERV;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.4.42 and port 12489: No route to host
[1477095309] HOST ALERT: CSCFILESERV;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.4.42)
[1477095333] SERVICE ALERT: CSCFILESERV;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1477095337] HOST ALERT: CSCFILESERV;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1477095353] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477095379] SERVICE ALERT: CSCFILESERV;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1477095385] SERVICE ALERT: CSCFILESERV;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 2 minute(s)
[1477095385] SERVICE ALERT: CSCFILESERV;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1477095407] SERVICE ALERT: EMUFILES;C:\ Drive Space;OK;SOFT;2;c:\ - total: 59.90 Gb - used: 35.61 Gb (59%) - free 24.29 Gb (41%)
[1477095423] SERVICE ALERT: CSCFILESERV;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1477096223] SERVICE ALERT: C4SAS-SE;CPU Load;OK;HARD;3;CPU Load 8% (15 min average)
[1477096553] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477096805] SERVICE ALERT: EMUFILES;CPU Load;OK;HARD;3;CPU Load 1% (15 min average)
[1477097154] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477097753] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477098353] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1477099223] SERVICE ALERT: C4SAS-SE;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1477099343] SERVICE ALERT: C4SAS-SE;CPU Load;OK;SOFT;2;CPU Load 87% (15 min average)
[1477099554] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477100153] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477100753] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477101005] SERVICE ALERT: EMUFILES;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1477101125] SERVICE ALERT: EMUFILES;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1477101643] SERVICE ALERT: INTLAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1477101763] SERVICE ALERT: INTLAPP;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1477101883] SERVICE ALERT: INTLAPP;CPU Load;OK;SOFT;3;CPU Load 5% (15 min average)
[1477102341] SERVICE ALERT: C4SAS-SE;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1477102461] SERVICE ALERT: C4SAS-SE;CPU Load;OK;SOFT;2;CPU Load 80% (15 min average)
[1477102553] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477103153] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477103648] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 118 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477103753] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477103768] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477104354] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477104368] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 127 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477104488] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477104907] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 165 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477104953] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477105027] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477105088] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 136 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477105208] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477105553] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477105627] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 144 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477105747] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477105808] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 123 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477105928] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477106153] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477106753] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477107953] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)