[1477109153] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477109753] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 91% (15 min average)
[1477110353] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1477110473] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1477110593] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 71% (15 min average)
[1477111793] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1477111913] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1477112033] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1477112778] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1477112898] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1477113233] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1477113499] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1477113618] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1477113834] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1477113954] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1477114073] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477114200] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1477114220] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1477114338] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1477114937] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1477115057] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1477115657] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1477115777] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1477115873] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477115897] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;3;OK. All replicas within 120 seconds.
[1477116473] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477116497] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1477116617] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1477116737] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1477117073] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477117337] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1477117937] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1477118273] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477118538] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1477118538] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1477118873] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477119473] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1477119736] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1477120073] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477120288] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1477120408] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1477120528] SERVICE ALERT: epprogsql2;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477120673] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477120936] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1477121128] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477121536] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1477121656] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;2;Warning:  idm1  are more than 120 seconds behind!
[1477121776] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1477122976] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1477123576] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1477123673] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477123696] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1477124273] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477124728] SERVICE ALERT: epprogsql2;CPU Load;OK;HARD;3;CPU Load 17% (15 min average)
[1477125002] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (58.3% change >= 20.0% threshold)
[1477125002] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (41.1% change >= 20.0% threshold)
[1477125002] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1477125002] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1477125002] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1477125002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1477125002] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1477125002] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1477125002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1477125002] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1477125002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1477125002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1477125002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1477125002] HOST DOWNTIME ALERT: VUM;STARTED; Host has entered a period of scheduled downtime
[1477125002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1477125002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1477125002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1477125002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1477125002] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1477125002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1477127273] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477127873] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477128473] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477129074] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477131473] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477131653] SERVICE ALERT: PAPERCUT;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 127.54 Gb - used: 118.96 Gb (93%) - free 8.58 Gb (7%)
[1477131773] SERVICE ALERT: PAPERCUT;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 127.54 Gb - used: 118.96 Gb (93%) - free 8.58 Gb (7%)
[1477131893] SERVICE ALERT: PAPERCUT;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 127.54 Gb - used: 118.96 Gb (93%) - free 8.58 Gb (7%)
[1477132073] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477132200] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service has exited from a period of scheduled downtime
[1477133896] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1477135673] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477136873] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477137473] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1477138073] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1477138193] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1477138313] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477139513] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477140113] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 78% (15 min average)
[1477140713] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1477140833] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1477140953] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1477141553] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477142153] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477142753] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477143353] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477146353] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477147553] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477149353] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477149953] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477150553] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477151153] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477153553] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477153803] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (49.6% change >= 20.0% threshold)
[1477153803] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1477153803] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1477153803] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1477153803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1477153803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1477153803] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1477153803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1477153803] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1477153803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1477153803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1477153803] HOST DOWNTIME ALERT: VUM;STARTED; Host has entered a period of scheduled downtime
[1477153803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1477153803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1477153803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1477153803] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1477153803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1477153803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1477154153] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477154413] SERVICE ALERT: hallefirst1;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1477154473] SERVICE ALERT: hallefirst1;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! All Trays Still Full
[1477154753] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1477155353] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1477155473] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1477155593] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477157393] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477157787] HOST ALERT: hallefirst1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1477157851] HOST ALERT: hallefirst1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.62 ms
[1477157862] SERVICE ALERT: eateriespk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1477157921] SERVICE ALERT: eateriespk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1477157981] SERVICE ALERT: eateriespk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1477157993] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477159194] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477159793] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1477159925] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 104 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477160045] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477160982] SERVICE ALERT: eateriespk;Device Status;OK;HARD;3;MENUS  RETRIEVE JOB  INFORMATION  PAPER HANDLING
[1477164593] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1477166393] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1477166513] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1477166633] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 79% (15 min average)
[1477167234] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1477167353] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1477167473] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477168074] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477168673] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477169274] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477169873] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477170430] SERVICE ALERT: hallegroundcolor;Image Transfer Kit HP CE249A;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Image Transfer Kit HP CE249A'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1477170483] SERVICE ALERT: hallegroundcolor;Image Transfer Kit HP CE249A;OK;SOFT;2;Image Transfer Kit HP CE249A is at 56% - OK!
[1477170858] SERVICE ALERT: C4SAS-SE;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1477170978] SERVICE ALERT: C4SAS-SE;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1477171098] SERVICE ALERT: C4SAS-SE;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477171673] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 82% (15 min average)
[1477171698] SERVICE ALERT: C4SAS-SE;CPU Load;OK;HARD;3;CPU Load 4% (15 min average)
[1477172873] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1477172993] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1477173113] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477173713] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477174313] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1477174913] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1477175033] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1477175153] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477175753] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477176354] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 86% (15 min average)
[1477176953] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1477177073] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1477177193] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477179593] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477180193] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477180794] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477182593] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477183193] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (15 min average)
[1477183794] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1477183913] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1477184033] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477186433] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 20% (15 min average)
[1477187033] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1477187153] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1477187274] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477187644] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 106 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477187764] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477188474] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 75% (15 min average)
[1477189073] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1477189193] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1477189313] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477190513] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477191114] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 91% (15 min average)
[1477191713] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1477191833] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1477191953] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477193153] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477193753] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477194353] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)