[1455004563] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1455004624] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1455004683] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1455004822] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1455004882] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1455004941] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1455021722] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Open right door.
[1455021723] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Open right door.
[1455022321] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1455022323] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1455022381] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1455022383] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1455022441] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1455022443] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1455024141] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1455024340] SERVICE ALERT: ADWSUSSERVERS;Service: WMI;CRITICAL;SOFT;1;Winmgmt: Stopped
[1455024460] SERVICE ALERT: ADWSUSSERVERS;Service: WMI;CRITICAL;SOFT;2;Winmgmt: Stopped
[1455024468] SERVICE ALERT: ADWSUSSERVERS;Service: WMI;CRITICAL;SOFT;1;Winmgmt: Stopped
[1455024484] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1455024580] SERVICE ALERT: ADWSUSSERVERS;Service: WMI;OK;SOFT;3;Winmgmt: Started
[1455024588] SERVICE ALERT: ADWSUSSERVERS;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1455024725] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1455024785] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1455024785] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1455024842] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1455024843] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1455024845] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1455024845] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1455024904] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1455026877] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1455026936] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1455026979] HOST ALERT: hallefirst2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1455026996] SERVICE ALERT: crossroadspk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1455026997] SERVICE ALERT: hallefirst2;Paper Tray Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1455027027] HOST ALERT: hallefirst2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1455027057] SERVICE ALERT: hallefirst2;Paper Tray Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1455027064] HOST ALERT: hallefirst2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1455027087] HOST ALERT: hallefirst2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1455027117] SERVICE ALERT: hallefirst2;Paper Tray Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1455027147] HOST ALERT: hallefirst2;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1455027154] HOST ALERT: hallefirst2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1455027179] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1455027237] HOST ALERT: hallefirst2;DOWN;SOFT;5;(Host check timed out after 30.00 seconds)
[1455027239] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1455027244] HOST ALERT: hallefirst2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1455027247] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1455027249] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1455027251] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;WARNING;HARD;3;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1455027277] HOST ALERT: hallefirst2;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1455027279] HOST ALERT: hallefirst2;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1455027298] SERVICE ALERT: hallefirst2;Paper Tray Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1455027299] SERVICE ALERT: crossroadspk;Device Status;OK;SOFT;3;Open top cover.
[1455027303] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10042;UNKNOWN;SOFT;1;Could not open the 'application' event log: 1717: The interface is unknown.
[1455027307] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1455027308] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1455027310] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;WARNING;HARD;3;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1455027328] HOST ALERT: hallefirst2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1455027338] HOST ALERT: hallefirst2;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1455027359] SERVICE ALERT: hallefirst2;Paper Tray Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1455027367] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1455027367] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1455027368] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1455027368] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1455027389] HOST ALERT: hallefirst2;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1455027398] HOST ALERT: hallefirst2;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1455027418] SERVICE ALERT: hallefirst2;Paper Tray Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1455027423] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10042;OK;SOFT;2;Eventlog check ok
[1455027448] HOST ALERT: hallefirst2;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1455027488] HOST ALERT: hallefirst2;DOWN;SOFT;9;(Host check timed out after 30.00 seconds)
[1455027538] HOST ALERT: hallefirst2;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1455027578] HOST ALERT: hallefirst2;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1455027596] SERVICE ALERT: crossroadspk;Device Status;OK;HARD;3;Open top cover.
[1455027627] HOST ALERT: hallefirst2;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1455027717] HOST ALERT: hallefirst2;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1455028196] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1455028208] SERVICE ALERT: prayharroldpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1455028256] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1455028269] SERVICE ALERT: prayharroldpk;Device Status;OK;SOFT;2;Remove any  jammed  paper.
[1455028302] HOST ALERT: hallefirst2;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1455028312] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1455028316] SERVICE ALERT: crossroadspk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1455028351] HOST ALERT: hallefirst2;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1455028445] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 42% - OK!
[1455028499] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1455028504] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 42% - OK!
[1455028559] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1455028561] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;TRAY 3 OPEN  46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1455028562] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;TRAY 3 OPEN  46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1455028611] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1455028619] SERVICE ALERT: crossroadspk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1455028917] SERVICE ALERT: crossroadspk;Device Status;OK;HARD;3;Open top cover.
[1455028917] SERVICE FLAPPING ALERT: crossroadspk;Device Status;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1455029219] SERVICE ALERT: crossroadspk;Device Status;OK;HARD;3;Open top cover.
[1455029818] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1455029878] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1455029938] SERVICE ALERT: crossroadspk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1455030115] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1455030175] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1455030235] SERVICE ALERT: crossroadspk;Device Status;OK;SOFT;3;Open top cover.
[1455030538] SERVICE ALERT: crossroadspk;Device Status;OK;HARD;3;Open top cover.
[1455030538] SERVICE FLAPPING ALERT: crossroadspk;Device Status;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1455032045] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1455032104] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 100% - OK!
[1455032229] SERVICE ALERT: studentcenterpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1455032569] SERVICE ALERT: studentcenterpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1455032710] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1455032710] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1455032740] HOST ALERT: hallefirst2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1455032740] HOST ALERT: hallefirst2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1455032766] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1455032767] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1455032770] HOST ALERT: hallefirst2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.67 ms
[1455032771] HOST ALERT: hallefirst2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.68 ms
[1455032823] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;SOFT;3;Black Print Cartridge HP C8543X is at 42% - OK!
[1455032824] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;SOFT;3;Black Print Cartridge HP C8543X is at 42% - OK!
[1455034780] HOST DOWNTIME ALERT: IMAGES;STARTED; Host has entered a period of scheduled downtime
[1455034791] SERVICE ALERT: IMAGES;MDT Monitor;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1455034791] SERVICE ALERT: IMAGES;Service: RPC;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1455034816] SERVICE ALERT: IMAGES;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1455034869] SERVICE ALERT: IMAGES;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1455034884] SERVICE ALERT: IMAGES;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1455034899] HOST ALERT: IMAGES;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1455034928] SERVICE ALERT: IMAGES;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1455034929] SERVICE ALERT: IMAGES;Service: PXE Boot;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1455034931] SERVICE ALERT: IMAGES;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1455034931] SERVICE ALERT: IMAGES;MDT Monitor;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1455034936] SERVICE ALERT: IMAGES;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1455034937] SERVICE ALERT: IMAGES;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1455034939] SERVICE ALERT: IMAGES;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1455034948] SERVICE ALERT: IMAGES;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1455034958] HOST ALERT: IMAGES;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1455034961] HOST ALERT: IMAGES;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1455034972] SERVICE ALERT: IMAGES;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1455034974] HOST ALERT: IMAGES;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1455034976] HOST ALERT: IMAGES;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1455034984] SERVICE ALERT: IMAGES;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1455035028] SERVICE ALERT: IMAGES;C:\ Drive Space;OK;SOFT;2;c:\ - total: 1841.16 Gb - used: 979.64 Gb (53%) - free 861.52 Gb (47%)
[1455035029] SERVICE ALERT: IMAGES;Service: PXE Boot;OK;SOFT;2;WDSServer: Started
[1455035031] SERVICE ALERT: IMAGES;MDT Monitor;OK;SOFT;3;MDT_Monitor: Started
[1455035031] SERVICE ALERT: IMAGES;Service: RPC;OK;SOFT;3;RpcSs: Started
[1455035036] SERVICE ALERT: IMAGES;Service: NetLogin;OK;SOFT;3;Netlogon: Started
[1455035037] SERVICE ALERT: IMAGES;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1455035039] SERVICE ALERT: IMAGES;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1455035048] SERVICE ALERT: IMAGES;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1455036761] SERVICE ALERT: ADWSUSSERVERS;Service: VMWare Tools;CRITICAL;SOFT;1;VMTools: Stopped
[1455036767] SERVICE ALERT: ADWSUSSERVERS;Service: Windows Time;WARNING;SOFT;1;W32Time: Error
[1455036790] SERVICE ALERT: ADWSUSSERVERS;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1455036881] SERVICE ALERT: ADWSUSSERVERS;Service: VMWare Tools;OK;SOFT;2;VMTools: Started
[1455036887] SERVICE ALERT: ADWSUSSERVERS;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1455036890] SERVICE ALERT: ADWSUSSERVERS;CPU Load;OK;SOFT;2;CPU Load 9% (5 min average)
[1455038361] HOST DOWNTIME ALERT: IMAGES;STOPPED; Host has exited from a period of scheduled downtime
[1455038375] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1455038399] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1455038435] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1455038459] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1455038495] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1455038519] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1455039360] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1455039362] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1455039694] SERVICE ALERT: eateriespk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1455039718] SERVICE ALERT: eateriespk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1455040433] SERVICE FLAPPING ALERT: crossroadspk;Device Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1455041155] SERVICE ALERT: owen101lab3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1455041159] HOST ALERT: owen101lab3;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.58 ms
[1455041161] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 30% - OK!
[1455041260] SERVICE ALERT: owen101lab3;Device Status;OK;HARD;3;Clearing paper path
[1455041489] HOST ALERT: owen101lab3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1455041579] HOST ALERT: owen101lab3;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1455041669] HOST ALERT: owen101lab3;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1455041759] HOST ALERT: owen101lab3;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1455041760] SERVICE ALERT: owen101lab3;Paper Tray Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1455041766] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1455041790] HOST ALERT: owen101lab3;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1455041820] SERVICE ALERT: owen101lab3;Paper Tray Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1455041826] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1455041850] HOST ALERT: owen101lab3;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1455041866] SERVICE ALERT: owen101lab3;Device Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1455041880] SERVICE ALERT: owen101lab3;Paper Tray Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1455041886] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;WARNING;HARD;3;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1455041896] HOST ALERT: owen101lab3;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1455041926] SERVICE ALERT: owen101lab3;Device Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1455041936] SERVICE FLAPPING ALERT: crossroadspk;Device Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1455041956] HOST ALERT: owen101lab3;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1455041986] SERVICE ALERT: owen101lab3;Device Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1455042016] HOST ALERT: owen101lab3;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1455042106] HOST ALERT: owen101lab3;DOWN;HARD;10;(Host check timed out after 30.00 seconds)
[1455042601] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at 5% - WARNING!
[1455042661] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at 5% - WARNING!
[1455042688] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at 5% - WARNING!
[1455042721] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at 5% - WARNING!
[1455042748] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at 5% - WARNING!
[1455042808] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at 5% - WARNING!
[1455043696] SERVICE ALERT: ADWSUSSERVERS;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1455043796] SERVICE ALERT: ADWSUSSERVERS;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1455047281] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 21% - OK!
[1455047285] HOST ALERT: halleground3;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.59 ms
[1455047324] HOST ALERT: halleground3;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.53 ms
[1455047382] SERVICE ALERT: halleground3;Device Status;OK;HARD;3;Waiting for tray 4 to lift   Ready
[1455047443] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 21% - OK!
[1455047460] SERVICE ALERT: halleground3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1455047461] SERVICE ALERT: halleground3;Device Status;OK;HARD;3;Waiting for tray 4 to lift   Ready
[1455047584] SERVICE ALERT: halleground3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1455048048] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1455048051] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at 3% - WARNING!
[1455048065] SERVICE ALERT: halleground3;Paper Tray Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1455048067] SERVICE ALERT: halleground3;Device Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1455048078] HOST ALERT: halleground3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1455048095] HOST ALERT: halleground3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1455048108] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1455048110] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at 3% - WARNING!
[1455048125] SERVICE ALERT: halleground3;Paper Tray Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1455048127] SERVICE ALERT: halleground3;Device Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1455048138] HOST ALERT: halleground3;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1455048155] HOST ALERT: halleground3;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1455048168] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;WARNING;HARD;3;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1455048170] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at 3% - WARNING!
[1455048185] SERVICE ALERT: halleground3;Paper Tray Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1455048187] SERVICE ALERT: halleground3;Device Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1455048190] SERVICE ALERT: halleground3;Paper Tray Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1455048198] HOST ALERT: halleground3;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1455048215] HOST ALERT: halleground3;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1455048250] SERVICE ALERT: halleground3;Paper Tray Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1455048280] HOST ALERT: halleground3;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1455048305] HOST ALERT: halleground3;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1455048310] SERVICE ALERT: halleground3;Paper Tray Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1455048327] HOST ALERT: halleground3;UP;SOFT;5;PING OK - Packet loss = 72%, RTA = 0.55 ms
[1455048350] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 12072;UNKNOWN;SOFT;1;Could not open the 'application' event log: 1717: The interface is unknown.
[1455048369] HOST ALERT: halleground3;UP;SOFT;5;PING OK - Packet loss = 0%, RTA = 0.54 ms
[1455048381] SERVICE ALERT: ADWSUSSERVERS;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1455048470] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 12072;OK;SOFT;2;Eventlog check ok
[1455048481] SERVICE ALERT: ADWSUSSERVERS;NSClient++ Version;OK;SOFT;2;NSClient++ 0,4,1,73 2012-12-17
[1455048588] SERVICE ALERT: halleground3;Device Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1455048601] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 1 Trays Empty
[1455048618] HOST ALERT: halleground3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1455048618] HOST FLAPPING ALERT: halleground3;STARTED; Host appears to have started flapping (20.8% change > 20.0% threshold)
[1455048648] SERVICE ALERT: halleground3;Device Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1455048661] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 1 Trays Empty
[1455048678] HOST ALERT: halleground3;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1455048699] HOST ALERT: halleground3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1455048699] HOST FLAPPING ALERT: halleground3;STARTED; Host appears to have started flapping (21.1% change > 20.0% threshold)
[1455048708] SERVICE ALERT: halleground3;Device Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1455048722] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;SOFT;3;Printer Paper Is OK! All Trays Still Full
[1455048738] HOST ALERT: halleground3;DOWN;SOFT;3;(Host check timed out after 30.00 seconds)
[1455048789] HOST ALERT: halleground3;DOWN;SOFT;2;(Host check timed out after 30.00 seconds)
[1455048798] HOST ALERT: halleground3;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1455048879] HOST ALERT: halleground3;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1455048888] HOST ALERT: halleground3;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1455048940] HOST ALERT: halleground3;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1455048969] HOST ALERT: halleground3;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1455049030] HOST ALERT: halleground3;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1455049059] HOST ALERT: halleground3;DOWN;SOFT;5;(Host check timed out after 30.00 seconds)
[1455049086] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1455049116] HOST ALERT: halleground3;DOWN;SOFT;6;(Host check timed out after 30.01 seconds)
[1455049120] HOST ALERT: halleground3;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1455049146] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1455049176] HOST ALERT: halleground3;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1455049206] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;WARNING;HARD;3;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1455049210] HOST ALERT: halleground3;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1455049236] HOST ALERT: halleground3;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1455049300] HOST ALERT: halleground3;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1455049326] HOST ALERT: halleground3;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1455049416] HOST ALERT: halleground3;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1455049510] HOST FLAPPING ALERT: halleground3;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1455051596] SERVICE ALERT: ADWSUSSERVERS;Service: DNS Client;WARNING;SOFT;1;Dnscache: Error
[1455051716] SERVICE ALERT: ADWSUSSERVERS;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1455055971] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;CRITICAL;HARD;3;Black Print Cartridge HP C8543X is at 1% - CRITICAL!
[1455057720] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1455057809] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1455058320] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at 3% - WARNING!
[1455058408] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at 3% - WARNING!
[1455058764] SERVICE ALERT: halleground1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1455058823] SERVICE ALERT: halleground1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1455058883] SERVICE ALERT: halleground1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1455059186] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1455059233] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1455059246] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 100% - OK!
[1455059293] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 11% - OK!
[1455059483] SERVICE ALERT: halleground1;Device Status;OK;HARD;3;Processing job from tray 4
[1455061940] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1455062001] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1455062061] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1455062283] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1455062342] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1455062403] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1455062586] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;WARNING;SOFT;1;Magenta Cartridge HP CE263A is at 3% - WARNING!
[1455062645] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;WARNING;SOFT;2;Magenta Cartridge HP CE263A is at 3% - WARNING!
[1455062668] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;WARNING;SOFT;1;Magenta Cartridge HP CE263A is at 3% - WARNING!
[1455062706] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;WARNING;HARD;3;Magenta Cartridge HP CE263A is at 3% - WARNING!
[1455062727] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;WARNING;SOFT;2;Magenta Cartridge HP CE263A is at 3% - WARNING!
[1455062788] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;WARNING;HARD;3;Magenta Cartridge HP CE263A is at 3% - WARNING!
[1455063859] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1455064460] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1455064520] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1455064579] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1455064920] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;CRITICAL;HARD;3;Black Print Cartridge HP C8543X is at 1% - CRITICAL!
[1455065007] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;CRITICAL;HARD;3;Black Print Cartridge HP C8543X is at 1% - CRITICAL!
[1455065105] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;OK;HARD;3;Magenta Cartridge HP CE263A is at 100% - OK!
[1455065186] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;OK;HARD;3;Magenta Cartridge HP CE263A is at 100% - OK!
[1455067320] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1455067320] SERVICE FLAPPING ALERT: hallefirst1;Black Print Cartridge HP C8543X;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1455067407] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1455067407] SERVICE FLAPPING ALERT: hallefirst1;Black Print Cartridge HP C8543X;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1455067477] SERVICE ALERT: DB01-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1455067597] SERVICE ALERT: DB01-2012;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1455069849] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;SOFT;1;e:\ - total: 99.87 Gb - used: 80.60 Gb (81%) - free 19.27 Gb (19%)
[1455069969] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;SOFT;2;e:\ - total: 99.87 Gb - used: 80.60 Gb (81%) - free 19.27 Gb (19%)
[1455069991] SERVICE ALERT: SANDBOX-DB;CPU Load;WARNING;SOFT;1;CPU Load 81% (5 min average)
[1455070089] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;HARD;3;e:\ - total: 99.87 Gb - used: 80.60 Gb (81%) - free 19.27 Gb (19%)
[1455070111] SERVICE ALERT: SANDBOX-DB;CPU Load;WARNING;SOFT;2;CPU Load 88% (5 min average)
[1455070147] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;SOFT;1;e:\ - total: 99.87 Gb - used: 80.60 Gb (81%) - free 19.27 Gb (19%)
[1455070231] SERVICE ALERT: SANDBOX-DB;CPU Load;OK;SOFT;3;CPU Load 78% (5 min average)
[1455070267] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;SOFT;2;e:\ - total: 99.87 Gb - used: 80.60 Gb (81%) - free 19.27 Gb (19%)
[1455070387] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;HARD;3;e:\ - total: 99.87 Gb - used: 80.60 Gb (81%) - free 19.27 Gb (19%)
[1455070597] SERVICE ALERT: DB01-2012;CPU Load;WARNING;SOFT;1;CPU Load 88% (5 min average)
[1455070717] SERVICE ALERT: DB01-2012;CPU Load;WARNING;SOFT;2;CPU Load 84% (5 min average)
[1455070837] SERVICE ALERT: DB01-2012;CPU Load;WARNING;HARD;3;CPU Load 86% (5 min average)
[1455071437] SERVICE ALERT: DB01-2012;CPU Load;OK;HARD;3;CPU Load 73% (5 min average)
[1455071889] SERVICE ALERT: AD-RMAD;E:\ Drive Space;OK;HARD;3;e:\ - total: 99.87 Gb - used: 79.60 Gb (80%) - free 20.27 Gb (20%)
[1455072187] SERVICE ALERT: AD-RMAD;E:\ Drive Space;OK;HARD;3;e:\ - total: 99.87 Gb - used: 79.60 Gb (80%) - free 20.27 Gb (20%)
[1455073608] SERVICE ALERT: alexanderpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1455073668] SERVICE ALERT: alexanderpk;Device Status;OK;SOFT;2;MENUS  SHOW ME HOW  RETRIEVE JOB  INFORMATION
[1455076919] SERVICE FLAPPING ALERT: hallefirst1;Black Print Cartridge HP C8543X;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1455077007] SERVICE FLAPPING ALERT: hallefirst1;Black Print Cartridge HP C8543X;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1455078840] SERVICE ALERT: ADWU1;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 399.66 Gb - used: 321.11 Gb (80%) - free 78.55 Gb (20%)
[1455078960] SERVICE ALERT: ADWU1;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 399.66 Gb - used: 322.42 Gb (81%) - free 77.23 Gb (19%)
[1455079080] SERVICE ALERT: ADWU1;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 399.66 Gb - used: 323.09 Gb (81%) - free 76.57 Gb (19%)