[1459656119] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 68% (5 min average)
[1459656172] SERVICE ALERT: SP2010DB;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459656292] SERVICE ALERT: SP2010DB;CPU Load;OK;SOFT;2;CPU Load 30% (5 min average)
[1459656719] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459656839] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 61% (5 min average)
[1459657439] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1459657559] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 92% (5 min average)
[1459657679] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 97% (5 min average)
[1459659116] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;1;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1459659236] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;2;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1459659356] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;HARD;3;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1459660079] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 92% (5 min average)
[1459660679] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 97% (5 min average)
[1459661279] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 93% (5 min average)
[1459661879] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 95% (5 min average)
[1459663079] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 87% (5 min average)
[1459663516] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1459663636] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 91% (5 min average)
[1459663679] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459663683] SERVICE ALERT: MBAM;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459663756] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 84% (5 min average)
[1459663799] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1459663803] SERVICE ALERT: MBAM;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1459663919] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1459664356] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1459664476] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 89% (5 min average)
[1459666556] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;OK;HARD;3;OK: No entries found
[1459666876] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459666996] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1459667116] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1459668316] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 84% (5 min average)
[1459669516] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 94% (5 min average)
[1459669636] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1459669756] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 97% (5 min average)
[1459670356] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 89% (5 min average)
[1459670356] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1459671119] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 88% (5 min average)
[1459671719] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1459671839] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 61% (5 min average)
[1459673039] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1459673159] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1459673279] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 95% (5 min average)
[1459674479] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 37% (5 min average)
[1459675079] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 94% (5 min average)
[1459675199] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 54% (5 min average)
[1459676999] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1459677119] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 95% (5 min average)
[1459677239] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1459677839] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 93% (5 min average)
[1459678439] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 96% (5 min average)
[1459681755] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1459682639] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 80% (5 min average)
[1459683839] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1459683959] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 97% (5 min average)
[1459684079] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 92% (5 min average)
[1459684679] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1459685879] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 85% (5 min average)
[1459688879] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 94% (5 min average)
[1459688998] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 40% (5 min average)
[1459689597] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1459689717] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 97% (5 min average)
[1459689837] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1459690437] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1459691037] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 80% (5 min average)
[1459691637] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 91% (5 min average)
[1459691757] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 74% (5 min average)
[1459692357] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459692477] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 93% (5 min average)
[1459692597] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 52% (5 min average)
[1459693197] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1459693317] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 77% (5 min average)
[1459693917] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1459694037] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 95% (5 min average)
[1459694157] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1459694757] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 88% (5 min average)
[1459695957] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459696077] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 93% (5 min average)
[1459696197] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1459696820] SERVICE ALERT: halleground2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1459696825] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1459696880] SERVICE ALERT: halleground2;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1459696885] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1459696907] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1459696940] SERVICE ALERT: halleground2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1459696945] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1459696967] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1459697027] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1459697255] SERVICE ALERT: halleground1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1459697284] SERVICE ALERT: NETIDADPROXY;Service: DNS Client;CRITICAL;SOFT;1;Dnscache: Stopped
[1459697315] SERVICE ALERT: halleground1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1459697317] SERVICE ALERT: halleground1;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1459697361] HOST ALERT: halleground2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459697375] SERVICE ALERT: halleground1;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1459697375] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1459697377] SERVICE ALERT: halleground1;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1459697397] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 81% (5 min average)
[1459697404] SERVICE ALERT: NETIDADPROXY;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1459697425] HOST ALERT: halleground2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.61 ms
[1459697435] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1459697437] SERVICE ALERT: halleground1;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1459697465] HOST ALERT: halleground1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459697495] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1459697525] HOST ALERT: halleground1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1459697615] HOST ALERT: halleground1;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1459697705] HOST ALERT: halleground1;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1459697795] HOST ALERT: halleground1;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1459697885] HOST ALERT: halleground1;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1459697975] HOST ALERT: halleground1;DOWN;SOFT;7;(Host check timed out after 30.00 seconds)
[1459697997] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1459698065] HOST ALERT: halleground1;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1459698117] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1459698125] HOST ALERT: halleground1;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1459698215] HOST ALERT: halleground1;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1459698237] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1459700037] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (5 min average)
[1459700637] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1459700757] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1459700877] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1459702677] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 94% (5 min average)
[1459703277] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 85% (5 min average)
[1459703877] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459703997] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1459704117] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1459704332] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Ready
[1459704717] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 63% (5 min average)
[1459704906] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459704966] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459705026] SERVICE ALERT: hallefirst1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459705317] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1459705437] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 85% (5 min average)
[1459706037] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 94% (5 min average)
[1459706157] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 97% (5 min average)
[1459706227] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;Clearing paper path
[1459706277] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1459706369] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1459706430] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1459706489] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;SOFT;3;Black Print Cartridge HP C8543X is at 56% - OK!
[1459706877] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1459707477] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1459711677] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 94% (5 min average)
[1459712277] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1459715277] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 83% (5 min average)
[1459715877] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1459715997] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 91% (5 min average)
[1459716117] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 70% (5 min average)
[1459716717] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1459716837] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1459716957] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1459719957] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 60% (5 min average)
[1459720557] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1459720677] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 74% (5 min average)
[1459721277] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1459721397] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1459721517] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1459724172] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1459724232] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1459724292] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1459724517] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 81% (5 min average)
[1459724891] SERVICE ALERT: eateriespk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1459725117] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1459725237] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 94% (5 min average)
[1459725357] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1459725439] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459725499] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459725559] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459726159] SERVICE ALERT: hallesouthcommons;Device Status;OK;HARD;3;Open top cover.
[1459727359] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459727420] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459727479] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459727959] HOST ALERT: halleground2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459728049] HOST ALERT: halleground2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1459728079] SERVICE ALERT: hallesouthcommons;Device Status;OK;HARD;3;Open top cover.
[1459728079] SERVICE FLAPPING ALERT: hallesouthcommons;Device Status;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1459728139] HOST ALERT: halleground2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1459728175] HOST ALERT: halleground2;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1459728255] HOST ALERT: halleground2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1459728345] HOST ALERT: halleground2;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1459728435] HOST ALERT: halleground2;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1459728525] HOST ALERT: halleground2;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1459728614] HOST ALERT: halleground2;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1459728679] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459728704] HOST ALERT: halleground2;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1459728739] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459728799] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459729399] SERVICE ALERT: hallesouthcommons;Device Status;OK;HARD;3;Open top cover.
[1459730825] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459730886] SERVICE ALERT: hallefirst1;Device Status;OK;SOFT;2;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1459731735] SERVICE ALERT: hallefirst3;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1459731785] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1459731795] SERVICE ALERT: hallefirst3;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1459731795] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459731845] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1459731855] SERVICE ALERT: hallefirst3;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1459731855] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459731855] SERVICE ALERT: hallefirst3;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1459731905] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1459731915] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459731915] SERVICE ALERT: hallefirst3;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1459731937] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459731975] SERVICE ALERT: hallefirst3;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1459731997] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459732025] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1459732057] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459732085] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1459732090] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1459732145] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1459732150] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1459732204] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459732210] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1459732229] SERVICE ALERT: IT-2012;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1459732264] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459732295] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1459732324] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459732325] HOST ALERT: hallefirst1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459732349] SERVICE ALERT: IT-2012;CPU Load;WARNING;SOFT;2;CPU Load 90% (5 min average)
[1459732355] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1459732366] HOST ALERT: hallefirst1;UP;SOFT;2;PING OK - Packet loss = 58%, RTA = 0.66 ms
[1459732409] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;SOFT;3;Black Print Cartridge HP C8543X is at 55% - OK!
[1459732449] SERVICE ALERT: hallefirst3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1459732469] SERVICE ALERT: IT-2012;CPU Load;OK;SOFT;3;CPU Load 86% (5 min average)
[1459732500] SERVICE ALERT: hallefirst1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1459732509] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1459732570] SERVICE ALERT: hallefirst3;Device Status;OK;HARD;3;ORDER CARTRIDGE LESS THAN 1000 PAGES 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1459732651] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Ready
[1459732675] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1459732704] SERVICE ALERT: halleground3;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1459732705] SERVICE ALERT: halleground3;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1459732735] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1459732739] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 3% - OK!
[1459732764] SERVICE ALERT: halleground3;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1459732765] SERVICE ALERT: halleground3;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1459732795] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1459732804] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;Processing job from tray 4
[1459732824] SERVICE ALERT: halleground3;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1459732825] SERVICE ALERT: halleground3;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1459732918] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 37% - OK!
[1459734589] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 9% - OK!
[1459734618] SERVICE ALERT: halleground3;Device Status;OK;HARD;3;Ready
[1459734620] SERVICE ALERT: halleground3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1459735369] SERVICE ALERT: hallesouthcommons;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1459735430] SERVICE ALERT: hallesouthcommons;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1459735490] SERVICE ALERT: hallesouthcommons;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1459736757] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 94% (5 min average)
[1459736926] SERVICE ALERT: VC;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459737026] SERVICE ALERT: VC;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1459737357] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1459737676] SERVICE ALERT: CBORD;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459737776] SERVICE ALERT: CBORD;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1459739353] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1459739473] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1459739593] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 95% (5 min average)
[1459740793] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1459740799] SERVICE FLAPPING ALERT: hallesouthcommons;Device Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1459741393] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 97% (5 min average)
[1459741993] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 78% (5 min average)
[1459741993] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)