[1490760126] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1490760185] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1490760243] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1490761443] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 81% (15 min average)
[1490761443] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1490761801] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1490762643] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1490762703] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1490762763] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1490763138] SERVICE ALERT: ADWU1;Service: NetLogin;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1490763138] SERVICE ALERT: ADWU1;Service: User Profile;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1490763198] SERVICE ALERT: ADWU1;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1490763198] SERVICE ALERT: ADWU1;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1490763363] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 90% (15 min average)
[1490763551] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;1;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1490763671] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;2;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1490763791] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;HARD;3;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1490763963] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1490764023] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 48% (15 min average)
[1490764354] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1490764413] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1490764473] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1490765074] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1490765674] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1490766273] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1490766333] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1490766393] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1490766393] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1490766993] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1490767593] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1490768223] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1490768283] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1490768343] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1490769393] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1490769770] SERVICE ALERT: NETENG-PA;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1490769830] SERVICE ALERT: NETENG-PA;CPU Load;OK;SOFT;2;CPU Load 1% (15 min average)
[1490769993] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1490770143] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 86% (15 min average)
[1490770389] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;OK;HARD;3;OK: No entries found
[1490771241] SERVICE ALERT: ADCA1;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1490771301] SERVICE ALERT: ADCA1;CPU Load;OK;SOFT;2;CPU Load 9% (15 min average)
[1490771343] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1490771403] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1490771463] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1490772063] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1490772993] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1490773594] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1490775063] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1490775663] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1490775723] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 89% (15 min average)
[1490776200] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service has exited from a period of scheduled downtime
[1490776203] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (26.6% change >= 20.0% threshold)
[1490776203] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (45.5% change >= 20.0% threshold)
[1490776203] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1490776203] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1490776203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1490776203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1490776203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1490776203] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1490776203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1490776203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1490776203] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1490776203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1490776203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1490776203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1490776203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1490776203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1490776203] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1490776923] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1490776983] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1490778002] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (24.7% change >= 20.0% threshold)
[1490778002] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (31.1% change >= 20.0% threshold)
[1490778002] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1490778002] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1490778002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1490778002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1490778002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1490778002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1490778002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1490778002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1490778002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1490778002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1490778002] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1490778002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1490778002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1490778002] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1490778002] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1490782383] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1490782443] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1490782503] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1490783103] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1490784303] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 80% (15 min average)
[1490784903] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1490784963] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1490784993] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1490785023] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1490785623] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (15 min average)
[1490786223] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1490786283] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1490786343] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1490786943] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1490787543] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1490787603] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1490788837] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1490788898] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1490788958] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1490789243] SERVICE ALERT: strongpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1490789573] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1490789990] HOST ALERT: itservices;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490790080] HOST ALERT: itservices;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1490790170] HOST ALERT: itservices;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1490790260] HOST ALERT: itservices;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1490790350] HOST ALERT: itservices;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1490790440] HOST ALERT: itservices;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1490790530] HOST ALERT: itservices;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1490790620] HOST ALERT: itservices;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1490790710] HOST ALERT: itservices;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1490790747] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6392 found since last run. oneLevelSearchOps OK. searchOps is too high! 6399 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490790800] HOST ALERT: itservices;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1490790807] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490790849] HOST ALERT: biapp;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.229)
[1490790912] HOST ALERT: biapp;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.4.229)
[1490790976] HOST ALERT: biapp;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1490791203] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1490791263] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1490791323] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1490791923] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 40% (15 min average)
[1490792052] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6741 found since last run. oneLevelSearchOps OK. searchOps is too high! 7040 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490792112] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490792712] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8062 found since last run. oneLevelSearchOps OK. searchOps is too high! 8322 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490792772] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490793010] SERVICE ALERT: hallefirst3;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1490793027] HOST ALERT: hallefirst3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490793053] HOST ALERT: hallefirst3;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1490793064] SERVICE ALERT: hallefirst3;Device Status;OK;SOFT;2;Processing job from tray 4
[1490793372] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7479 found since last run. oneLevelSearchOps OK. searchOps is too high! 7821 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490793432] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490793495] SERVICE ALERT: strongpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1490793554] SERVICE ALERT: strongpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1490793614] SERVICE ALERT: strongpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1490793807] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10833 found since last run. oneLevelSearchOps OK. searchOps is too high! 10841 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490793867] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490794032] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5175 found since last run. oneLevelSearchOps OK. searchOps is too high! 5368 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490794092] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490794203] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (42.2% change >= 20.0% threshold)
[1490794203] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1490794203] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1490794203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1490794203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1490794203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1490794203] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1490794203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1490794203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1490794203] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1490794203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1490794203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1490794203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1490794203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1490794203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1490794203] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1490794467] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9527 found since last run. oneLevelSearchOps OK. searchOps is too high! 9537 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490794527] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490794692] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7246 found since last run. oneLevelSearchOps OK. searchOps is too high! 7493 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490794752] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490795127] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6299 found since last run. oneLevelSearchOps OK. searchOps is too high! 6313 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490795187] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490795414] SERVICE ALERT: strongpk;Device Status;OK;HARD;3;Ready
[1490795464] SERVICE ALERT: prayharroldpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1490795524] SERVICE ALERT: prayharroldpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1490795584] SERVICE ALERT: prayharroldpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1490795952] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6233 found since last run. oneLevelSearchOps OK. searchOps is too high! 6564 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490796012] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490796387] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6878 found since last run. oneLevelSearchOps OK. searchOps is too high! 6892 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490796447] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490796612] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5350 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490796633] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1490796672] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490796693] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1490796752] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1490796859] SERVICE ALERT: hallegroundcolor;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 0 Trays Empty
[1490796914] SERVICE ALERT: hallegroundcolor;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! All Trays Still Full
[1490797272] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7074 found since last run. oneLevelSearchOps OK. searchOps is too high! 7452 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490797332] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490797593] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1490797647] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7411 found since last run. oneLevelSearchOps OK. searchOps is too high! 7421 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490797653] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1490797707] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490797983] SERVICE ALERT: prayharroldpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1490798253] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1490798307] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7097 found since last run. oneLevelSearchOps OK. searchOps is too high! 7109 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490798313] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;2;Warning:  idm1  are more than 120 seconds behind!
[1490798367] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490798373] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1490798532] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7960 found since last run. oneLevelSearchOps OK. searchOps is too high! 8207 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490798592] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490798967] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6566 found since last run. oneLevelSearchOps OK. searchOps is too high! 6577 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490798973] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1490799027] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490799192] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 17488 found since last run. oneLevelSearchOps OK. searchOps is too high! 17826 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490799252] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490799627] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9840 found since last run. oneLevelSearchOps OK. searchOps is too high! 9848 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490799687] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490799852] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9756 found since last run. oneLevelSearchOps OK. searchOps is too high! 10159 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490799912] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490800287] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 11614 found since last run. oneLevelSearchOps OK. searchOps is too high! 11624 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490800322] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1490800347] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490800382] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1490800512] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6138 found since last run. oneLevelSearchOps OK. searchOps is too high! 6389 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490800572] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490800947] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8050 found since last run. oneLevelSearchOps OK. searchOps is too high! 8058 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490800982] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1490801007] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490801042] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 79% (15 min average)
[1490801172] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5201 found since last run. oneLevelSearchOps OK. searchOps is too high! 5450 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490801232] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490801642] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1490801702] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1490801762] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1490801832] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5237 found since last run. oneLevelSearchOps OK. searchOps is too high! 5530 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490801892] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490802207] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6507 found since last run. oneLevelSearchOps OK. searchOps is too high! 6517 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490802267] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490802492] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5057 found since last run. oneLevelSearchOps OK. searchOps is too high! 5445 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490802552] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490802757] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1490802867] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10292 found since last run. oneLevelSearchOps OK. searchOps is too high! 10302 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490802927] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490802962] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 74% (15 min average)
[1490803152] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5280 found since last run. oneLevelSearchOps OK. searchOps is too high! 5475 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490803212] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490803527] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10900 found since last run. oneLevelSearchOps OK. searchOps is too high! 10920 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490803562] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1490803587] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490803622] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1490803682] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1490803812] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5053 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490803872] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490804187] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 12970 found since last run. oneLevelSearchOps OK. searchOps is too high! 12974 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490804247] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490804282] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1490804472] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9590 found since last run. oneLevelSearchOps OK. searchOps is too high! 10041 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490804532] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490804847] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 16993 found since last run. oneLevelSearchOps OK. searchOps is too high! 17005 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490804882] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1490804907] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490805002] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (29.8% change >= 20.0% threshold)
[1490805002] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1490805002] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1490805002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1490805002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1490805002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1490805002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1490805002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1490805002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1490805002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1490805002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1490805002] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1490805002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1490805002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1490805002] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1490805002] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1490805132] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9388 found since last run. oneLevelSearchOps OK. searchOps is too high! 9903 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490805192] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490805502] SERVICE ALERT: studentcenterpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1490805507] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 11985 found since last run. oneLevelSearchOps OK. searchOps is too high! 11989 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490805561] SERVICE ALERT: studentcenterpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1490805567] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490805620] SERVICE ALERT: studentcenterpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1490805792] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5600 found since last run. oneLevelSearchOps OK. searchOps is too high! 5839 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490805852] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490806167] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10469 found since last run. oneLevelSearchOps OK. searchOps is too high! 10491 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490806220] SERVICE ALERT: studentcenterpk;Device Status;OK;HARD;3;TRAY 3 EMPTY PLAIN LETTER   Ready
[1490806227] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490806452] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6330 found since last run. oneLevelSearchOps OK. searchOps is too high! 6666 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490806512] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490806682] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 79% (15 min average)
[1490806764] SERVICE ALERT: owen101lab1;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1490806794] HOST ALERT: owen101lab1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490806824] SERVICE ALERT: owen101lab1;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1490806827] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10772 found since last run. oneLevelSearchOps OK. searchOps is too high! 10792 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490806834] HOST ALERT: owen101lab1;UP;SOFT;2;PING OK - Packet loss = 54%, RTA = 0.52 ms
[1490806878] SERVICE ALERT: owen101lab1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1490806887] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490807112] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8618 found since last run. oneLevelSearchOps OK. searchOps is too high! 9061 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490807172] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490807282] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1490807342] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1490807402] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1490807478] SERVICE ALERT: owen101lab1;Device Status;OK;HARD;3;Ready
[1490807487] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6682 found since last run. oneLevelSearchOps OK. searchOps is too high! 6705 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490807547] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490807772] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7906 found since last run. oneLevelSearchOps OK. searchOps is too high! 8405 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490807832] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490808002] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 91% (15 min average)
[1490808432] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9528 found since last run. oneLevelSearchOps OK. searchOps is too high! 10094 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490808492] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490808576] HOST ALERT: hallefirstcolor;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490808602] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1490808640] HOST ALERT: hallefirstcolor;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1490808662] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1490808722] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1490808753] SERVICE ALERT: studentcenterpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1490809092] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9629 found since last run. oneLevelSearchOps OK. searchOps is too high! 10145 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490809152] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490809322] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 61% (15 min average)
[1490809348] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6025 found since last run. oneLevelSearchOps OK. searchOps is too high! 6052 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490809407] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490809752] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8757 found since last run. oneLevelSearchOps OK. searchOps is too high! 9146 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490809812] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490809922] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1490809982] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 55% (15 min average)
[1490810006] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8186 found since last run. oneLevelSearchOps OK. searchOps is too high! 8213 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490810065] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490810113] SERVICE ALERT: hallegroundcolor;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 1 Trays Empty
[1490810173] SERVICE ALERT: hallegroundcolor;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 1 Trays Empty
[1490810233] SERVICE ALERT: hallegroundcolor;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 1 Trays Empty
[1490810314] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1490810374] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1490810412] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6752 found since last run. oneLevelSearchOps OK. searchOps is too high! 7434 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490810434] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1490810472] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490810582] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1490810642] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1490810665] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6666 found since last run. oneLevelSearchOps OK. searchOps is too high! 6697 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490810702] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 56% (15 min average)
[1490810725] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490811302] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1490811325] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5000 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490811362] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 82% (15 min average)
[1490811385] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490811443] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;CRITICAL;SOFT;1;Black Cartridge HP CE260X is at 1% - CRITICAL!
[1490811502] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;CRITICAL;SOFT;2;Black Cartridge HP CE260X is at 1% - CRITICAL!
[1490811562] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;OK;SOFT;3;
[1490811672] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5060 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490811732] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490811921] SERVICE ALERT: hallefirstcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1490811962] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1490811981] SERVICE ALERT: hallefirstcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1490811985] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6832 found since last run. oneLevelSearchOps OK. searchOps is too high! 6858 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490812022] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1490812041] SERVICE ALERT: hallefirstcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1490812045] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490812082] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1490812332] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5237 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490812392] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490812632] SERVICE ALERT: hallegroundcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1490812641] SERVICE ALERT: hallefirstcolor;Device Status;OK;HARD;3;Open right door.
[1490812682] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 27% (15 min average)
[1490812992] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5371 found since last run. oneLevelSearchOps OK. searchOps is too high! 5905 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490813052] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490813245] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8489 found since last run. oneLevelSearchOps OK. searchOps is too high! 8510 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490813282] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1490813305] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490813342] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1490813402] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1490813652] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5347 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490813712] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490813905] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 11542 found since last run. oneLevelSearchOps OK. searchOps is too high! 11572 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490813963] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;CRITICAL;SOFT;1;Black Cartridge HP CE260X is at 1% - CRITICAL!
[1490813965] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490814002] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 85% (15 min average)
[1490814022] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;CRITICAL;SOFT;2;Black Cartridge HP CE260X is at 1% - CRITICAL!
[1490814083] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;CRITICAL;HARD;3;Black Cartridge HP CE260X is at 1% - CRITICAL!
[1490814312] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6719 found since last run. oneLevelSearchOps OK. searchOps is too high! 7080 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490814372] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490814565] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10531 found since last run. oneLevelSearchOps OK. searchOps is too high! 10573 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490814602] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1490814625] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490814662] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1490814722] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1490814722] SERVICE ALERT: sillpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1490814783] SERVICE ALERT: sillpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1490814843] SERVICE ALERT: sillpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1490814859] HOST ALERT: ADWSUSSERVERS;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.79)
[1490814872] SERVICE ALERT: ADWSUSSERVERS;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.4.79 and port 12489: No route to host
[1490814872] SERVICE ALERT: ADWSUSSERVERS;Service: NetLogin;CRITICAL;SOFT;1;connect to address 164.76.4.79 and port 12489: No route to host
[1490814872] SERVICE ALERT: ADWSUSSERVERS;Service: User Profile;CRITICAL;SOFT;1;connect to address 164.76.4.79 and port 12489: No route to host
[1490814877] HOST ALERT: ADWSUSSERVERS;UP;SOFT;2;PING OK - Packet loss = 16%, RTA = 111.07 ms
[1490814929] SERVICE ALERT: ADWSUSSERVERS;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1490814929] SERVICE ALERT: ADWSUSSERVERS;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1490814929] SERVICE ALERT: ADWSUSSERVERS;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1490814972] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8207 found since last run. oneLevelSearchOps OK. searchOps is too high! 8824 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490815032] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490815225] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9050 found since last run. oneLevelSearchOps OK. searchOps is too high! 9099 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490815233] SERVICE ALERT: eateriespk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1490815285] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490815322] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1490815632] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 11351 found since last run. oneLevelSearchOps OK. searchOps is too high! 11922 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490815692] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490815833] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1490815885] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6570 found since last run. oneLevelSearchOps OK. searchOps is too high! 6614 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490815893] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1490815922] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 62% (15 min average)
[1490815945] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490815953] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1490816292] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 11062 found since last run. oneLevelSearchOps OK. searchOps is too high! 11439 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490816352] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490816522] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1490816582] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1490816766] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1490816826] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1490816886] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1490816952] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8680 found since last run. oneLevelSearchOps OK. searchOps is too high! 9033 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490817012] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490817020] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1490817050] HOST ALERT: hallefirst2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490817072] HOST ALERT: hallefirst2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.58 ms
[1490817074] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! All Trays Still Full
[1490817145] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7699 found since last run. oneLevelSearchOps OK. searchOps is too high! 7744 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490817152] SERVICE ALERT: eateriespk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1490817152] SERVICE FLAPPING ALERT: eateriespk;Paper Tray Status;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1490817182] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1490817205] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490817242] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 70% (15 min average)
[1490817612] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6268 found since last run. oneLevelSearchOps OK. searchOps is too high! 6637 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490817672] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490817674] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1490817680] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1490817735] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1490817740] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1490817765] HOST ALERT: hallefirst2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490817794] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1490817800] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1490817824] HOST ALERT: hallefirst2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1490817914] HOST ALERT: hallefirst2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1490818004] HOST ALERT: hallefirst2;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1490818093] HOST ALERT: hallefirst2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1490818183] HOST ALERT: hallefirst2;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1490818272] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5428 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490818273] HOST ALERT: hallefirst2;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1490818332] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490818363] HOST ALERT: hallefirst2;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1490818423] HOST ALERT: hallefirst2;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1490818442] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1490818502] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1490818513] HOST ALERT: hallefirst2;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1490818562] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 82% (15 min average)
[1490818932] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5021 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490818992] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490819005] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7934 found since last run. oneLevelSearchOps OK. searchOps is too high! 7957 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490819065] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490819162] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1490819222] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1490819282] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1490819592] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8876 found since last run. oneLevelSearchOps OK. searchOps is too high! 9326 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490819652] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490819665] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5839 found since last run. oneLevelSearchOps OK. searchOps is too high! 5864 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490819699] SERVICE ALERT: marshallpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.137.49. Is the printer off?  Is the network down?
[1490819725] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490819759] SERVICE ALERT: marshallpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.137.49. Is the printer off?  Is the network down?
[1490819819] SERVICE ALERT: marshallpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.137.49. Is the printer off?  Is the network down?
[1490819888] SERVICE ALERT: marshallpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.137.49. Is the printer off?  Is the network down?
[1490819894] SERVICE ALERT: marshallpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.137.49. Is the printer off?  Is the network down?
[1490819948] SERVICE ALERT: marshallpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.137.49. Is the printer off?  Is the network down?
[1490819954] SERVICE ALERT: marshallpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.137.49. Is the printer off?  Is the network down?
[1490820008] SERVICE ALERT: marshallpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.137.49. Is the printer off?  Is the network down?
[1490820014] SERVICE ALERT: marshallpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.137.49. Is the printer off?  Is the network down?
[1490820252] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10827 found since last run. oneLevelSearchOps OK. searchOps is too high! 11254 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490820312] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490820413] SERVICE ALERT: marshallpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1490820482] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 81% (15 min average)
[1490820602] SERVICE ALERT: marshallpk;Device Status;OK;HARD;3;TRAY 3 EMPTY PLAIN LETTER   Ready
[1490820609] SERVICE ALERT: marshallpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 49% - OK!
[1490820912] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7397 found since last run. oneLevelSearchOps OK. searchOps is too high! 7727 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490820925] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9233 found since last run. oneLevelSearchOps OK. searchOps is too high! 9267 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490820972] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490820985] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490821572] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9809 found since last run. oneLevelSearchOps OK. searchOps is too high! 10165 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490821585] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6423 found since last run. oneLevelSearchOps OK. searchOps is too high! 6465 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490821632] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490821645] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490821682] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1490821742] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1490821802] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1490822232] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5663 found since last run. oneLevelSearchOps OK. searchOps is too high! 6100 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490822245] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9004 found since last run. oneLevelSearchOps OK. searchOps is too high! 9030 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490822292] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490822305] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490822334] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1490822394] SERVICE ALERT: halleground2;Device Status;OK;SOFT;2;50 72 65 73 73 20 17 20 74 6F 20 73 65 65 20 73  74 65 70  by step information
[1490822553] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1490822613] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1490822673] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1490822892] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7072 found since last run. oneLevelSearchOps OK. searchOps is too high! 7398 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490822905] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 11367 found since last run. oneLevelSearchOps OK. searchOps is too high! 11390 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490822952] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490822965] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490823364] SERVICE ALERT: prayharroldpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1490823423] SERVICE ALERT: prayharroldpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1490823483] SERVICE ALERT: prayharroldpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1490823565] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7247 found since last run. oneLevelSearchOps OK. searchOps is too high! 7279 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490823625] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490824083] SERVICE ALERT: prayharroldpk;Device Status;OK;HARD;3;Ready
[1490824152] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5172 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490824212] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490824802] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1490824812] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7663 found since last run. oneLevelSearchOps OK. searchOps is too high! 7968 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490824872] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490825472] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6706 found since last run. oneLevelSearchOps OK. searchOps is too high! 6900 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490825532] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490826132] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5323 found since last run. oneLevelSearchOps OK. searchOps is too high! 5521 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490826192] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490826602] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1490826625] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7135 found since last run. oneLevelSearchOps OK. searchOps is too high! 7158 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490826662] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 90% (15 min average)
[1490826685] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490827235] SERVICE ALERT: owen101lab2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1490827262] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1490827285] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6466 found since last run. oneLevelSearchOps OK. searchOps is too high! 6488 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490827294] SERVICE ALERT: owen101lab2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1490827322] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1490827345] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490827354] SERVICE ALERT: owen101lab2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1490827382] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 83% (15 min average)
[1490827945] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8484 found since last run. oneLevelSearchOps OK. searchOps is too high! 8511 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490827952] SERVICE FLAPPING ALERT: eateriespk;Paper Tray Status;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1490828005] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490828592] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6971 found since last run. oneLevelSearchOps OK. searchOps is too high! 7275 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490828652] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490829182] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1490829242] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1490829242] SERVICE ALERT: sillpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1490829302] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1490829902] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1490830405] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5936 found since last run. oneLevelSearchOps OK. searchOps is too high! 5957 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490830465] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490830502] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1490830562] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1490830622] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1490831652] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5548 found since last run. oneLevelSearchOps OK. searchOps is too high! 5673 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490831712] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490833872] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1490833932] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 9% - OK!
[1490833954] SERVICE ALERT: owen101lab2;Device Status;OK;HARD;3;Clearing paper path
[1490834065] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5196 found since last run. oneLevelSearchOps OK. searchOps is too high! 5212 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490834125] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490834554] SERVICE ALERT: owen101lab2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1490834615] SERVICE ALERT: owen101lab2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1490834674] SERVICE ALERT: owen101lab2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1490835422] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 49% (15 min average)
[1490836499] SERVICE ALERT: TARGETX-2012;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1490836559] SERVICE ALERT: TARGETX-2012;CPU Load;OK;SOFT;2;CPU Load 90% (15 min average)
[1490836622] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1490836682] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 51% (15 min average)
[1490837282] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1490837342] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 88% (15 min average)
[1490837400] HOST DOWNTIME ALERT: IMAGES;STARTED; Host has entered a period of scheduled downtime
[1490837712] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5376 found since last run. oneLevelSearchOps OK. searchOps is too high! 5596 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490837772] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490837942] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1490838002] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 75% (15 min average)
[1490838359] SERVICE ALERT: TARGETX-2012;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1490838419] SERVICE ALERT: TARGETX-2012;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1490838450] SERVICE ALERT: IMAGES;MDT Monitor;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1490838454] HOST ALERT: IMAGES;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490838479] SERVICE ALERT: TARGETX-2012;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1490838518] HOST ALERT: IMAGES;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1490839030] SERVICE ALERT: IMAGES;MDT Monitor;UNKNOWN;HARD;3;Failed to open service MDT_Monitor: 1060: The specified service does not exist as an installed service.
[1490839079] SERVICE ALERT: TARGETX-2012;CPU Load;OK;HARD;3;CPU Load 9% (15 min average)
[1490839202] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1490839262] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1490839322] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1490839470] SERVICE ALERT: studentcenterpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1490840522] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 78% (15 min average)
[1490840725] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5010 found since last run. oneLevelSearchOps OK. searchOps is too high! 5017 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490840785] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490841000] HOST DOWNTIME ALERT: R25-ALPHA2012;STARTED; Host has entered a period of scheduled downtime
[1490841000] HOST DOWNTIME ALERT: PROCARE;STARTED; Host has entered a period of scheduled downtime
[1490841000] HOST DOWNTIME ALERT: PNCWEB-2012;STARTED; Host has entered a period of scheduled downtime
[1490841000] HOST DOWNTIME ALERT: XDM;STARTED; Host has entered a period of scheduled downtime
[1490841000] HOST DOWNTIME ALERT: PNCDB-2012;STARTED; Host has entered a period of scheduled downtime
[1490841000] HOST DOWNTIME ALERT: PNCAPP-2012;STARTED; Host has entered a period of scheduled downtime
[1490841000] HOST DOWNTIME ALERT: WINKMS;STARTED; Host has entered a period of scheduled downtime
[1490841000] HOST DOWNTIME ALERT: OPIE;STARTED; Host has entered a period of scheduled downtime
[1490841000] HOST DOWNTIME ALERT: OFFICEKMS;STARTED; Host has entered a period of scheduled downtime
[1490841000] HOST DOWNTIME ALERT: UHSPHARM-2012;STARTED; Host has entered a period of scheduled downtime
[1490841000] HOST DOWNTIME ALERT: NETENG-PA;STARTED; Host has entered a period of scheduled downtime
[1490841000] HOST DOWNTIME ALERT: NESSUS;STARTED; Host has entered a period of scheduled downtime
[1490841000] HOST DOWNTIME ALERT: TI22-PSYCLINIC;STARTED; Host has entered a period of scheduled downtime
[1490841000] HOST DOWNTIME ALERT: MBAMDB;STARTED; Host has entered a period of scheduled downtime
[1490841000] HOST DOWNTIME ALERT: MBAM;STARTED; Host has entered a period of scheduled downtime
[1490841000] HOST DOWNTIME ALERT: TI22-COE;STARTED; Host has entered a period of scheduled downtime
[1490841000] HOST DOWNTIME ALERT: HDBACKUPS;STARTED; Host has entered a period of scheduled downtime
[1490841000] HOST DOWNTIME ALERT: FTK;STARTED; Host has entered a period of scheduled downtime
[1490841000] HOST DOWNTIME ALERT: TI22;STARTED; Host has entered a period of scheduled downtime
[1490841000] HOST DOWNTIME ALERT: EMUSQLEXPRESS;STARTED; Host has entered a period of scheduled downtime
[1490841000] HOST DOWNTIME ALERT: EMUSCRIPT;STARTED; Host has entered a period of scheduled downtime
[1490841000] HOST DOWNTIME ALERT: TABLEAU;STARTED; Host has entered a period of scheduled downtime
[1490841000] HOST DOWNTIME ALERT: DPS-SECURE02;STARTED; Host has entered a period of scheduled downtime
[1490841000] HOST DOWNTIME ALERT: DIRECTACCESS;STARTED; Host has entered a period of scheduled downtime
[1490841000] HOST DOWNTIME ALERT: SCCM-PRIM01;STARTED; Host has entered a period of scheduled downtime
[1490841000] HOST DOWNTIME ALERT: COE-WEBDAV;STARTED; Host has entered a period of scheduled downtime
[1490841000] HOST DOWNTIME ALERT: CEDARON;STARTED; Host has entered a period of scheduled downtime
[1490841000] HOST DOWNTIME ALERT: SANDBOX-DB;STARTED; Host has entered a period of scheduled downtime
[1490841000] HOST DOWNTIME ALERT: CBORD;STARTED; Host has entered a period of scheduled downtime
[1490841000] HOST DOWNTIME ALERT: ADFS;STARTED; Host has entered a period of scheduled downtime
[1490841000] HOST DOWNTIME ALERT: REFWIN;STARTED; Host has entered a period of scheduled downtime
[1490841000] HOST DOWNTIME ALERT: 4WINDSMANAGER;STARTED; Host has entered a period of scheduled downtime
[1490841000] HOST DOWNTIME ALERT: R25-ZETA2012;STARTED; Host has entered a period of scheduled downtime
[1490841000] HOST DOWNTIME ALERT: R25-WEBVIEWER;STARTED; Host has entered a period of scheduled downtime
[1490841000] HOST DOWNTIME ALERT: R25-DELTA2012;STARTED; Host has entered a period of scheduled downtime
[1490841000] HOST DOWNTIME ALERT: R25-WEBSERVICES;STARTED; Host has entered a period of scheduled downtime
[1490841000] HOST DOWNTIME ALERT: R25-INTERFACE;STARTED; Host has entered a period of scheduled downtime
[1490841122] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1490841182] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1490841242] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 91% (15 min average)
[1490841385] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5276 found since last run. oneLevelSearchOps OK. searchOps is too high! 5297 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490841445] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490841842] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1490841902] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1490841962] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1490842562] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 73% (15 min average)
[1490843026] SERVICE ALERT: SANDBOX-DB;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 39.66 Gb - used: 35.90 Gb (91%) - free 3.76 Gb (9%)
[1490843043] SERVICE ALERT: PROCARE;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1490843086] SERVICE ALERT: SANDBOX-DB;C:\ Drive Space;OK;SOFT;2;c:\ - total: 39.66 Gb - used: 35.15 Gb (89%) - free 4.51 Gb (11%)
[1490843103] SERVICE ALERT: PROCARE;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1490843122] SERVICE ALERT: MBAM;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1490843162] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1490843163] SERVICE ALERT: PROCARE;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1490843182] SERVICE ALERT: MBAM;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1490843197] SERVICE ALERT: NESSUS;Service: Group Policy;CRITICAL;SOFT;1;gpsvc: Stopped
[1490843222] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 65% (15 min average)
[1490843230] SERVICE ALERT: REFWIN;Service: DNS Client;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1490843242] SERVICE ALERT: MBAM;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1490843248] SERVICE ALERT: NESSUS;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.2.69 and port 12489: Connection refused
[1490843253] SERVICE ALERT: REFWIN;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1490843257] SERVICE ALERT: NESSUS;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1490843290] SERVICE ALERT: REFWIN;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1490843293] SERVICE ALERT: REFWIN;C:\ Drive Space;OK;SOFT;2;c:\ - total: 59.90 Gb - used: 44.22 Gb (74%) - free 15.68 Gb (26%)
[1490843302] SERVICE ALERT: NETENG-PA;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.4.142 and port 12489: Connection refused
[1490843308] SERVICE ALERT: NESSUS;Service: RPC;OK;SOFT;2;RpcSs: Started
[1490843316] HOST ALERT: NETENG-PA;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490843317] HOST ALERT: PROCARE;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490843362] SERVICE ALERT: NETENG-PA;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1490843366] HOST ALERT: NETENG-PA;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.36 ms
[1490843381] HOST ALERT: PROCARE;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.40 ms
[1490843536] HOST ALERT: FTK;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490843539] SERVICE ALERT: FTK;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1490843549] SERVICE ALERT: FTK;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1490843554] SERVICE ALERT: FTK;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1490843569] HOST ALERT: FTK;DOWN;HARD;2;PING CRITICAL - Packet loss = 100%
[1490843569] SERVICE ALERT: FTK;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1490843569] SERVICE ALERT: FTK;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1490843587] SERVICE ALERT: OPIE;Service: Group Policy;CRITICAL;SOFT;1;gpsvc: Stopped
[1490843592] SERVICE ALERT: FTK;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1490843599] SERVICE ALERT: FTK;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1490843602] SERVICE ALERT: FTK;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1490843602] SERVICE ALERT: FTK;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1490843603] SERVICE ALERT: TI22;Service: DNS Client;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1490843607] SERVICE ALERT: FTK;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1490843608] SERVICE ALERT: FTK;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1490843609] SERVICE ALERT: FTK;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1490843614] SERVICE ALERT: FTK;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1490843614] SERVICE ALERT: FTK;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1490843618] SERVICE ALERT: FTK;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1490843618] SERVICE ALERT: OPIE;Service: DNS Client;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1490843622] SERVICE ALERT: HDBACKUPS;Service: Group Policy;CRITICAL;SOFT;1;gpsvc: Stopped
[1490843628] SERVICE ALERT: TI22;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.4.202 and port 12489: Connection refused
[1490843629] SERVICE ALERT: FTK;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1490843629] SERVICE ALERT: FTK;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1490843632] SERVICE ALERT: TABLEAU;C:\ Drive Space;UNKNOWN;SOFT;1;NSClient - ERROR: Invalid return from command: check_drivesize
[1490843643] SERVICE ALERT: OPIE;VMWare Tools;CRITICAL;SOFT;1;connect to address 164.76.158.60 and port 12489: Connection refused
[1490843645] SERVICE ALERT: WINKMS;Service: Group Policy;CRITICAL;SOFT;1;gpsvc: Stopped
[1490843649] SERVICE ALERT: OPIE;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1490843650] SERVICE ALERT: OPIE;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1490843652] SERVICE ALERT: FTK;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1490843652] HOST ALERT: R25-INTERFACE;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490843654] SERVICE ALERT: R25-ZETA2012;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.4.174 and port 12489: No route to host
[1490843659] SERVICE ALERT: FTK;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1490843662] SERVICE ALERT: FTK;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1490843662] SERVICE ALERT: FTK;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1490843663] SERVICE ALERT: FTK;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.7.103 and port 12489: Connection refused
[1490843663] SERVICE ALERT: TI22;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1490843663] SERVICE ALERT: FTK;C:\ Drive Space;CRITICAL;SOFT;2;connect to address 164.76.7.103 and port 12489: Connection refused
[1490843663] SERVICE ALERT: FTK;CPU Load;CRITICAL;HARD;3;connect to address 164.76.7.103 and port 12489: Connection refused
[1490843666] SERVICE ALERT: FTK;Service: DNS Client;CRITICAL;SOFT;2;connect to address 164.76.7.103 and port 12489: Connection refused
[1490843667] SERVICE ALERT: FTK;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1490843668] SERVICE ALERT: FTK;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1490843668] SERVICE ALERT: EMUSQLEXPRESS;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1490843669] SERVICE ALERT: FTK;NSClient++ Version;CRITICAL;HARD;3;connect to address 164.76.7.103 and port 12489: Connection refused
[1490843669] SERVICE ALERT: FTK;Service: RPC;CRITICAL;HARD;3;connect to address 164.76.7.103 and port 12489: Connection refused
[1490843669] SERVICE ALERT: FTK;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1490843669] SERVICE ALERT: R25-WEBVIEWER;CPU Load;CRITICAL;SOFT;1;connect to address 164.76.4.84 and port 12489: Connection refused
[1490843670] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1490843675] SERVICE ALERT: R25-DELTA2012;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1490843675] SERVICE ALERT: PNCWEB-2012;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1490843677] SERVICE ALERT: EMUSQLEXPRESS;VMWare Tools;CRITICAL;SOFT;1;connect to address 164.76.7.169 and port 12489: Connection refused
[1490843677] SERVICE ALERT: PNCWEB-2012;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.2.178 and port 12489: Connection refused
[1490843678] SERVICE ALERT: CBORD;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.4.184 and port 12489: No route to host
[1490843678] SERVICE ALERT: OPIE;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1490843679] HOST ALERT: OPIE;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490843680] SERVICE ALERT: SANDBOX-DB;Service: Windows Time;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1490843680] SERVICE ALERT: TI22-COE;E:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.4.108 and port 12489: Connection refused
[1490843681] HOST ALERT: CBORD;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.184)
[1490843682] SERVICE ALERT: WINKMS;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1490843683] SERVICE ALERT: CEDARON;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1490843684] SERVICE ALERT: HDBACKUPS;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1490843685] SERVICE ALERT: R25-DELTA2012;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1490843686] SERVICE ALERT: R25-WEBSERVICES;Service: RPC;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1490843688] SERVICE ALERT: TI22;Service: RPC;OK;SOFT;2;RpcSs: Started
[1490843689] SERVICE ALERT: OPIE;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1490843692] SERVICE ALERT: TABLEAU;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1490843692] SERVICE ALERT: TABLEAU;C:\ Drive Space;OK;SOFT;2;c:\ - total: 149.66 Gb - used: 87.61 Gb (59%) - free 62.04 Gb (41%)
[1490843693] SERVICE ALERT: HDBACKUPS;Service: Group Policy;CRITICAL;SOFT;2;connect to address 164.76.4.109 and port 12489: No route to host
[1490843693] HOST ALERT: HDBACKUPS;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.109)
[1490843701] SERVICE ALERT: FTK;Service: NetLogin;OK;SOFT;3;Netlogon: Started
[1490843701] SERVICE ALERT: WINKMS;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1490843702] SERVICE ALERT: FTK;Service: Group Policy;OK;SOFT;3;gpsvc: Started
[1490843702] SERVICE ALERT: FTK;Service: User Profile;OK;SOFT;3;ProfSvc: Started
[1490843702] SERVICE ALERT: R25-ZETA2012;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1490843703] SERVICE ALERT: OPIE;VMWare Tools;OK;SOFT;2;VMTools: Started
[1490843705] SERVICE ALERT: WINKMS;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1490843707] SERVICE ALERT: FTK;Uptime;OK;SOFT;3;System Uptime - 0 day(s) 0 hour(s) 1 minute(s)
[1490843707] SERVICE ALERT: FTK;Service: Windows Time;CRITICAL;HARD;3;W32Time: Stopped
[1490843708] SERVICE ALERT: EMUSQLEXPRESS;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1490843714] SERVICE ALERT: FTK;C:\ Drive Space;OK;SOFT;3;c:\ - total: 408.28 Gb - used: 141.24 Gb (35%) - free 267.03 Gb (65%)
[1490843715] SERVICE ALERT: R25-DELTA2012;Service: RPC;OK;SOFT;2;RpcSs: Started
[1490843715] SERVICE ALERT: PNCWEB-2012;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1490843716] HOST ALERT: R25-INTERFACE;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 8.89 ms
[1490843717] SERVICE ALERT: FTK;Service: DNS Client;OK;SOFT;3;Dnscache: Started
[1490843719] HOST ALERT: OPIE;UP;SOFT;2;PING WARNING - Packet loss = 86%, RTA = 0.30 ms
[1490843719] SERVICE ALERT: SANDBOX-DB;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.4.164 and port 12489: Connection refused
[1490843719] HOST ALERT: FTK;UP;HARD;2;PING OK - Packet loss = 73%, RTA = 0.38 ms
[1490843722] SERVICE ALERT: ADWU1;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1490843722] SERVICE ALERT: WINKMS;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1490843723] SERVICE ALERT: FTK;Memory Usage;OK;SOFT;2;Memory usage: total:65524.62 MB - used: 19233.99 MB (29%) - free: 46290.63 MB (71%)
[1490843723] SERVICE ALERT: CEDARON;C:\ Drive Space;OK;SOFT;2;c:\ - total: 249.66 Gb - used: 42.34 Gb (17%) - free 207.31 Gb (83%)
[1490843725] SERVICE ALERT: R25-DELTA2012;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1490843726] SERVICE ALERT: CBORD;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1490843726] SERVICE ALERT: R25-WEBVIEWER;CPU Load;OK;SOFT;2;CPU Load 79% (15 min average)
[1490843727] SERVICE ALERT: LANSWEEPER;CPU Load;CRITICAL;SOFT;1;connect to address 164.76.2.33 and port 12489: Connection refused
[1490843727] SERVICE ALERT: HDBACKUPS;Service: Windows Time;CRITICAL;SOFT;2;connect to address 164.76.4.109 and port 12489: No route to host
[1490843728] SERVICE ALERT: EMUSQLEXPRESS;VMWare Tools;OK;SOFT;2;VMTools: Started
[1490843730] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1490843730] HOST ALERT: CBORD;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1490843730] HOST ALERT: IMAGES;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490843730] HOST ALERT: HDBACKUPS;DOWN;HARD;2;CRITICAL - Host Unreachable (164.76.4.109)
[1490843732] SERVICE ALERT: UHSPHARM-2012;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.4.209 and port 12489: Connection refused
[1490843732] SERVICE ALERT: LANSWEEPER;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1490843732] SERVICE ALERT: TABLEAU;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1490843734] SERVICE ALERT: HDBACKUPS;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.4.109 and port 12489: No route to host
[1490843737] SERVICE ALERT: TI22-COE;E:\ Drive Space;OK;SOFT;2;e:\ - total: 39.87 Gb - used: 0.25 Gb (1%) - free 39.62 Gb (99%)
[1490843737] SERVICE ALERT: PNCWEB-2012;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1490843740] SERVICE ALERT: SANDBOX-DB;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1490843741] SERVICE ALERT: WINKMS;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1490843745] SERVICE ALERT: HDBACKUPS;Service: Group Policy;CRITICAL;HARD;3;connect to address 164.76.4.109 and port 12489: No route to host
[1490843746] SERVICE ALERT: R25-WEBSERVICES;Service: RPC;OK;SOFT;2;RpcSs: Started
[1490843753] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1490843757] SERVICE ALERT: ADWSUSCLIENTS;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1490843762] SERVICE ALERT: ADWU1;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 1 minute(s)
[1490843763] SERVICE ALERT: PROCARE;CPU Load;OK;HARD;3;CPU Load 4% (15 min average)
[1490843766] SERVICE ALERT: HDBACKUPS;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.4.109 and port 12489: No route to host
[1490843772] SERVICE ALERT: LANSWEEPER;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1490843779] SERVICE ALERT: SANDBOX-DB;Service: RPC;OK;SOFT;2;RpcSs: Started
[1490843783] SERVICE ALERT: UHSPHARM-2012;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1490843787] SERVICE ALERT: LANSWEEPER;CPU Load;OK;SOFT;2;CPU Load 67% (15 min average)
[1490843787] SERVICE ALERT: HDBACKUPS;Service: Windows Time;CRITICAL;HARD;3;connect to address 164.76.4.109 and port 12489: No route to host
[1490843790] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1490843793] SERVICE ALERT: HDBACKUPS;Service: WMI;CRITICAL;SOFT;2;connect to address 164.76.4.109 and port 12489: No route to host
[1490843797] SERVICE ALERT: ADWSUSCLIENTS;Memory Usage;OK;SOFT;2;Memory usage: total:14335.67 MB - used: 2968.09 MB (21%) - free: 11367.58 MB (79%)
[1490843820] HOST ALERT: IMAGES;DOWN;HARD;2;PING CRITICAL - Packet loss = 100%
[1490843822] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1490843826] SERVICE ALERT: HDBACKUPS;Service: DNS Client;CRITICAL;SOFT;2;connect to address 164.76.4.109 and port 12489: No route to host
[1490843850] SERVICE ALERT: IMAGES;MDT Monitor;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1490843850] SERVICE ALERT: HDBACKUPS;Service: WMI;OK;SOFT;3;Winmgmt: Started
[1490843850] HOST ALERT: HDBACKUPS;UP;HARD;2;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1490843852] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;OK;SOFT;2;Eventlog check ok
[1490843860] HOST ALERT: IMAGES;UP;HARD;2;PING OK - Packet loss = 54%, RTA = 0.30 ms
[1490843882] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1490843883] SERVICE ALERT: HDBACKUPS;Service: DNS Client;OK;SOFT;3;Dnscache: Started
[1490843942] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1490843945] SERVICE ALERT: SANDBOX-DB;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1490844005] SERVICE ALERT: SANDBOX-DB;CPU Load;OK;SOFT;2;CPU Load 87% (15 min average)
[1490844018] SERVICE ALERT: MBAM;Service: Group Policy;CRITICAL;SOFT;1;gpsvc: Stopped
[1490844039] SERVICE ALERT: MBAMDB;Service: NetLogin;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1490844062] SERVICE ALERT: MBAM;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.7.98 and port 12489: Connection refused
[1490844067] SERVICE ALERT: MBAMDB;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1490844078] SERVICE ALERT: MBAM;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1490844079] SERVICE ALERT: MBAMDB;VMWare Tools;CRITICAL;SOFT;1;connect to address 164.76.7.97 and port 12489: Connection refused
[1490844082] SERVICE ALERT: MBAMDB;Service: NetLogin;CRITICAL;SOFT;2;connect to address 164.76.7.97 and port 12489: Connection refused
[1490844095] SERVICE ALERT: MBAMDB;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1490844097] HOST ALERT: MBAMDB;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490844107] SERVICE ALERT: MBAMDB;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1490844111] HOST ALERT: MBAMDB;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1490844122] SERVICE ALERT: MBAM;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1490844139] SERVICE ALERT: MBAMDB;VMWare Tools;OK;SOFT;2;VMTools: Started
[1490844142] SERVICE ALERT: MBAMDB;Service: NetLogin;OK;SOFT;3;Netlogon: Started
[1490844155] SERVICE ALERT: MBAMDB;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1490844166] SERVICE ALERT: OFFICEKMS;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1490844226] SERVICE ALERT: OFFICEKMS;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1490844239] SERVICE ALERT: FTK;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1490844249] SERVICE ALERT: FTK;Service: WMI;OK;HARD;3;Winmgmt: Started
[1490844254] SERVICE ALERT: FTK;CPU Load;OK;HARD;3;CPU Load 0% (15 min average)
[1490844269] SERVICE ALERT: FTK;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1490844269] SERVICE ALERT: FTK;Service: RPC;OK;HARD;3;RpcSs: Started
[1490844285] SERVICE ALERT: OFFICEKMS;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1490844307] SERVICE ALERT: FTK;Service: Windows Time;OK;HARD;3;W32Time: Started
[1490844342] SERVICE ALERT: HDBACKUPS;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1490844384] SERVICE ALERT: HDBACKUPS;Service: Windows Time;OK;HARD;3;W32Time: Started
[1490844430] SERVICE ALERT: IMAGES;MDT Monitor;UNKNOWN;HARD;3;Failed to open service MDT_Monitor: 1060: The specified service does not exist as an installed service.
[1490844430] SERVICE FLAPPING ALERT: IMAGES;MDT Monitor;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1490844442] SERVICE ALERT: MBAM;CPU Load;OK;HARD;3;CPU Load 3% (15 min average)
[1490844600] HOST DOWNTIME ALERT: IMAGES;STOPPED; Host has exited from a period of scheduled downtime
[1490844651] SERVICE ALERT: CLIENTLANSWEEP;Service: User Profile;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1490844653] SERVICE ALERT: PNCDB-2012;Service: Firewall;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1490844684] SERVICE ALERT: PNCDB-2012;CPU Load;CRITICAL;SOFT;1;connect to address 164.76.2.176 and port 12489: Connection refused
[1490844689] SERVICE ALERT: PNCDB-2012;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1490844711] SERVICE ALERT: CLIENTLANSWEEP;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1490844713] SERVICE ALERT: PNCDB-2012;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1490844729] SERVICE ALERT: PNCDB-2012;Memory Usage;OK;SOFT;2;Memory usage: total:17663.70 MB - used: 1607.19 MB (9%) - free: 16056.51 MB (91%)
[1490844732] SERVICE ALERT: DPS-SECURE02;Service: RPC;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1490844741] SERVICE ALERT: DPS-SECURE02;Service: Group Policy;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1490844744] SERVICE ALERT: PNCDB-2012;CPU Load;OK;SOFT;2;CPU Load 83% (15 min average)
[1490844769] SERVICE ALERT: DPS-SECURE02;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1490844799] HOST ALERT: DPS-SECURE02;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490844802] SERVICE ALERT: DPS-SECURE02;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1490844804] SERVICE ALERT: DPS-SECURE02;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1490844804] SERVICE ALERT: DPS-SECURE02;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1490844812] SERVICE ALERT: DPS-SECURE02;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1490844821] SERVICE ALERT: DPS-SECURE02;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1490844829] SERVICE ALERT: DPS-SECURE02;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1490844832] HOST ALERT: DPS-SECURE02;DOWN;HARD;2;PING CRITICAL - Packet loss = 100%
[1490844862] SERVICE ALERT: DPS-SECURE02;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1490844862] SERVICE ALERT: DPS-SECURE02;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1490844864] SERVICE ALERT: DPS-SECURE02;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1490844864] SERVICE ALERT: DPS-SECURE02;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1490844865] SERVICE ALERT: SCCM-PRIM01;C:\ Drive Space;UNKNOWN;SOFT;1;NSClient - ERROR: Invalid return from command: check_drivesize
[1490844867] SERVICE ALERT: DPS-SECURE02;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1490844872] SERVICE ALERT: DPS-SECURE02;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1490844881] SERVICE ALERT: DPS-SECURE02;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1490844885] SERVICE ALERT: OFFICEKMS;CPU Load;OK;HARD;3;CPU Load 1% (15 min average)
[1490844889] SERVICE ALERT: DPS-SECURE02;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1490844902] SERVICE ALERT: SCCM-PRIM01;Service: Group Policy;CRITICAL;SOFT;1;gpsvc: Stopped
[1490844922] SERVICE ALERT: DPS-SECURE02;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1490844922] SERVICE ALERT: DPS-SECURE02;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1490844924] SERVICE ALERT: DPS-SECURE02;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1490844924] SERVICE ALERT: DPS-SECURE02;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1490844925] SERVICE ALERT: SCCM-PRIM01;C:\ Drive Space;OK;SOFT;2;c:\ - total: 194.90 Gb - used: 123.02 Gb (63%) - free 71.88 Gb (37%)
[1490844927] SERVICE ALERT: DPS-SECURE02;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1490844941] HOST ALERT: DPS-SECURE02;UP;HARD;2;PING OK - Packet loss = 75%, RTA = 0.30 ms
[1490844962] SERVICE ALERT: DPS-SECURE02;Service: DNS Client;OK;SOFT;3;Dnscache: Started
[1490844967] SERVICE ALERT: DPS-SECURE02;NSClient++ Version;OK;SOFT;3;NSClient++ 0.4.4.15 2015-11-25
[1490844982] SERVICE ALERT: SCCM-PRIM01;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1490844985] HOST ALERT: SCCM-PRIM01;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.64)
[1490845025] SERVICE ALERT: SCCM-PRIM01;Service: Group Policy;CRITICAL;HARD;3;connect to address 164.76.4.64 and port 12489: No route to host
[1490845028] HOST ALERT: SCCM-PRIM01;DOWN;HARD;2;CRITICAL - Host Unreachable (164.76.4.64)
[1490845142] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 21% (15 min average)
[1490845146] HOST ALERT: SCCM-PRIM01;UP;HARD;2;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1490845452] SERVICE ALERT: DPS-SECURE02;Service: RPC;OK;HARD;3;RpcSs: Started
[1490845461] SERVICE ALERT: DPS-SECURE02;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1490845469] SERVICE ALERT: DPS-SECURE02;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1490845502] SERVICE ALERT: DPS-SECURE02;Service: Windows Time;OK;HARD;3;W32Time: Started
[1490845504] SERVICE ALERT: DPS-SECURE02;Service: WMI;OK;HARD;3;Winmgmt: Started
[1490845504] SERVICE ALERT: DPS-SECURE02;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1490845622] SERVICE ALERT: SCCM-PRIM01;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1490845742] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1490845802] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1490845862] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)