[1462076500] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 84% (15 min average)
[1462077100] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462077220] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462077340] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462078540] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462080340] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1462080940] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462081060] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462081180] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462082736] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462082856] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462082976] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462083580] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462084180] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462087176] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462087776] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462088976] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462088976] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1462088980] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1462089576] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462089580] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462089700] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462089820] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1462090176] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1462090420] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 42% (15 min average)
[1462090776] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462090896] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462091016] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1462092819] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462093416] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 84% (15 min average)
[1462094020] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462094140] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462094260] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462094616] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462094736] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 90% (15 min average)
[1462095336] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462095456] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1462096060] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 70% (15 min average)
[1462096660] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462096780] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1462096900] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 93% (15 min average)
[1462097500] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462097620] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462097740] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462098940] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462099540] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 91% (15 min average)
[1462100140] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462100260] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462100380] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462102180] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 64% (15 min average)
[1462102780] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462102900] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462103020] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1462105055] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1462105420] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 79% (15 min average)
[1462108420] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462108540] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462108660] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 92% (15 min average)
[1462109260] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462109380] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462109500] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1462110100] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 90% (15 min average)
[1462110700] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1462110820] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1462110940] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462112140] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 79% (15 min average)
[1462113340] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462113455] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462113460] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462113575] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 77% (15 min average)
[1462113580] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1462115380] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 80% (15 min average)
[1462115980] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1462116100] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462116220] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462117420] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 65% (15 min average)
[1462118620] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1462118740] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462118860] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1462121860] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 46% (15 min average)
[1462123060] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462123180] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1462123300] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1462128100] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1462129300] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462129420] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1462130620] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462130740] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462130860] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 91% (15 min average)
[1462131460] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462131580] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462131700] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462132300] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 91% (15 min average)
[1462132898] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462133018] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 66% (15 min average)
[1462133618] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462133738] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 91% (15 min average)
[1462134938] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1462135058] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462135178] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1462136978] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 79% (15 min average)
[1462137578] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462137698] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462137818] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 77% (15 min average)
[1462138418] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462138538] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462138658] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 54% (15 min average)
[1462139258] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462139378] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462139498] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1462144898] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 73% (15 min average)
[1462145498] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462145618] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462145738] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462147899] SERVICE ALERT: PNCDB-2012;D:\ Drive Space;WARNING;SOFT;1;d:\ - total: 350.00 Gb - used: 299.36 Gb (86%) - free 50.64 Gb (14%)
[1462148019] SERVICE ALERT: PNCDB-2012;D:\ Drive Space;WARNING;SOFT;2;d:\ - total: 350.00 Gb - used: 299.36 Gb (86%) - free 50.64 Gb (14%)
[1462148138] SERVICE ALERT: PNCDB-2012;D:\ Drive Space;WARNING;HARD;3;d:\ - total: 350.00 Gb - used: 300.34 Gb (86%) - free 49.66 Gb (14%)
[1462148738] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 85% (15 min average)
[1462149338] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462149458] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462149578] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462151717] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;SOFT;1;e:\ - total: 499.87 Gb - used: 400.90 Gb (80%) - free 98.98 Gb (20%)
[1462151837] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;SOFT;2;e:\ - total: 499.87 Gb - used: 400.90 Gb (80%) - free 98.98 Gb (20%)
[1462151957] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;HARD;3;e:\ - total: 499.87 Gb - used: 400.90 Gb (80%) - free 98.98 Gb (20%)
[1462153157] SERVICE ALERT: AD-RMAD;E:\ Drive Space;OK;HARD;3;e:\ - total: 499.87 Gb - used: 390.24 Gb (78%) - free 109.63 Gb (22%)
[1462154374] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462154378] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (15 min average)
[1462154494] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1462154614] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 52% (15 min average)
[1462154978] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462155098] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462155218] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1462157613] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1462157740] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1462157854] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462158655] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;SOFT;1;WARNING: Replication with ADDC5 failed!
[1462158774] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;SOFT;2;WARNING: Replication with ADDC5 failed!
[1462158818] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462158895] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;HARD;3;WARNING: Replication with ADDC5 failed!
[1462159418] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1462160853] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 78% (15 min average)