[1455447682] SERVICE ALERT: ADWU1;Windows Update Services Event 10042;UNKNOWN;SOFT;1;Failed to process command : EvtQuery failed: 1717: The interface is unknown.
[1455447706] SERVICE ALERT: ADWU1;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1455447802] SERVICE ALERT: ADWU1;Windows Update Services Event 10042;OK;SOFT;2;OK: No entries found
[1455447806] SERVICE ALERT: ADWU1;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 1 minute(s)
[1455448222] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;1;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1455448342] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;2;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1455448462] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;HARD;3;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1455450325] SERVICE ALERT: ADWU1;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1455450445] SERVICE ALERT: ADWU1;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1455458062] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;OK;HARD;3;OK: No entries found
[1455471096] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1455471156] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1455471217] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1455471691] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1455471752] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1455471811] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1455481102] HOST ALERT: hallethirdsouth;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1455481169] HOST ALERT: hallethirdsouth;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1455481192] HOST ALERT: hallethirdsouth;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1455481228] HOST ALERT: hallethirdsouth;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1455481259] HOST ALERT: hallethirdsouth;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1455481318] HOST ALERT: hallethirdsouth;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1455481349] HOST ALERT: hallethirdsouth;DOWN;SOFT;3;(Host check timed out after 30.00 seconds)
[1455481402] HOST ALERT: hallethirdsouth;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1455481408] HOST ALERT: hallethirdsouth;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1455481465] HOST ALERT: hallethirdsouth;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1455481498] HOST ALERT: hallethirdsouth;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1455481555] HOST ALERT: hallethirdsouth;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1455481588] HOST ALERT: hallethirdsouth;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1455481645] HOST ALERT: hallethirdsouth;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1455481678] HOST ALERT: hallethirdsouth;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1455481735] HOST ALERT: hallethirdsouth;DOWN;SOFT;8;(Host check timed out after 30.00 seconds)
[1455481768] HOST ALERT: hallethirdsouth;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1455481825] HOST ALERT: hallethirdsouth;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1455481828] HOST ALERT: hallethirdsouth;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1455481915] HOST ALERT: hallethirdsouth;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1455482010] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Clearing paper path
[1455482015] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Clearing paper path
[1455482610] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1455482616] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1455482671] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1455482675] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1455482730] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1455482734] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1455484530] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1455484534] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Checking printer
[1455485130] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1455485134] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1455485190] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1455485195] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1455485251] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1455485251] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1455485254] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1455485254] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1455485531] SERVICE ALERT: SANDBOX-DB;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1455485651] SERVICE ALERT: SANDBOX-DB;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1455485850] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Warming up
[1455485854] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Warming up
[1455497250] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1455497253] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1455501806] SERVICE ALERT: EVENTSTEST-2012;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1455501926] SERVICE ALERT: EVENTSTEST-2012;CPU Load;OK;SOFT;2;CPU Load 26% (5 min average)
[1455502451] SERVICE ALERT: SANDBOX-DB;CPU Load;WARNING;SOFT;1;CPU Load 86% (5 min average)
[1455502571] SERVICE ALERT: SANDBOX-DB;CPU Load;WARNING;SOFT;2;CPU Load 87% (5 min average)
[1455502691] SERVICE ALERT: SANDBOX-DB;CPU Load;OK;SOFT;3;CPU Load 66% (5 min average)
[1455503055] SERVICE ALERT: DB01-2012;CPU Load;WARNING;SOFT;1;CPU Load 84% (5 min average)
[1455503175] SERVICE ALERT: DB01-2012;CPU Load;OK;SOFT;2;CPU Load 76% (5 min average)
[1455503606] SERVICE ALERT: DB01-2012;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1455503726] SERVICE ALERT: DB01-2012;CPU Load;OK;SOFT;2;CPU Load 89% (5 min average)
[1455507157] SERVICE ALERT: ADWSUSSERVERS;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1455507257] SERVICE ALERT: ADWSUSSERVERS;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1455507898] SERVICE ALERT: VC;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1455507900] SERVICE ALERT: VC;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1455507998] SERVICE ALERT: VC;CPU Load;OK;SOFT;2;CPU Load 9% (5 min average)
[1455508000] SERVICE ALERT: VC;Service: Firewall;OK;SOFT;2;MpsSvc: Started