[1460261463] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 78% (5 min average)
[1460262063] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460262183] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460262303] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460264375] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;1;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1460264495] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;2;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1460264615] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;HARD;3;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1460267343] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1460267463] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 97% (5 min average)
[1460267583] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 83% (5 min average)
[1460268183] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460268303] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 92% (5 min average)
[1460268423] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 87% (5 min average)
[1460269023] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460269143] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460269263] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 96% (5 min average)
[1460271215] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;OK;HARD;3;OK: No entries found
[1460271952] SERVICE ALERT: EVENTS-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460272072] SERVICE ALERT: EVENTS-2012;CPU Load;OK;SOFT;2;CPU Load 24% (5 min average)
[1460272263] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 77% (5 min average)
[1460273463] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 91% (5 min average)
[1460273583] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 89% (5 min average)
[1460274183] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1460274303] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 91% (5 min average)
[1460274423] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 78% (5 min average)
[1460275023] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 91% (5 min average)
[1460275143] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 80% (5 min average)
[1460275503] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1460276103] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 85% (5 min average)
[1460276703] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 91% (5 min average)
[1460276823] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 87% (5 min average)
[1460278023] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1460278143] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 81% (5 min average)
[1460278743] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1460278863] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 96% (5 min average)
[1460278983] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 50% (5 min average)
[1460280183] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460280303] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1460280423] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460281402] SERVICE ALERT: ADWU1;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1460281521] SERVICE ALERT: ADWU1;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1460281623] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 75% (5 min average)
[1460281623] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1460281743] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1460282223] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460282343] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 91% (5 min average)
[1460282463] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 81% (5 min average)
[1460283063] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460283183] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1460283303] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460284411] SERVICE ALERT: INTLAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460284530] SERVICE ALERT: INTLAPP;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1460286903] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 94% (5 min average)
[1460287503] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 96% (5 min average)
[1460291103] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 90% (5 min average)
[1460291703] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 59% (5 min average)
[1460294103] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1460294223] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460294343] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1460295543] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 93% (5 min average)
[1460296143] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 51% (5 min average)
[1460296743] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460296863] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1460296983] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 97% (5 min average)
[1460299983] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 87% (5 min average)
[1460301783] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1460301903] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 94% (5 min average)
[1460302023] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 93% (5 min average)
[1460302623] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460304469] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 33% - OK!
[1460304473] HOST ALERT: markjeffersonpk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.65 ms
[1460304573] SERVICE ALERT: markjeffersonpk;Printer Status;OK;HARD;3;Ready
[1460304903] SERVICE ALERT: markjeffersonpk;Maintenance Kit HP 110V-C9152A;OK;HARD;3;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 56% - OK!
[1460304978] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1460305623] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 88% (5 min average)
[1460306223] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460306343] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460306463] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 64% (5 min average)
[1460307063] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460307183] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460307303] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460313303] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 92% (5 min average)
[1460313903] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460316903] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 92% (5 min average)
[1460317503] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460322303] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 76% (5 min average)
[1460322903] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460323023] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460323143] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460323743] SERVICE ALERT: hallefirst2;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460323803] SERVICE ALERT: hallefirst2;Printer Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460323863] SERVICE ALERT: hallefirst2;Printer Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460324343] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 63% (5 min average)
[1460324464] SERVICE ALERT: hallefirst2;Printer Status;OK;HARD;3;INFORMATION  PRINT MENU MAP  PRINT  CONFIGURATION
[1460324943] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460325063] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460325183] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 87% (5 min average)
[1460325783] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460325903] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 77% (5 min average)
[1460326503] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1460326623] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 81% (5 min average)
[1460327223] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460327343] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 62% (5 min average)
[1460327943] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460328063] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460328183] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 96% (5 min average)
[1460328783] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 71% (5 min average)
[1460329983] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460330103] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460330223] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1460332623] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1460333223] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1460337873] SERVICE ALERT: EVENTSTEST-2012;CPU Load;WARNING;SOFT;1;CPU Load 91% (5 min average)
[1460337993] SERVICE ALERT: EVENTSTEST-2012;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1460338033] SERVICE ALERT: hallegroundcolor;Cyan Cartridge HP CE261A;WARNING;SOFT;1;Cyan Cartridge HP CE261A is at 2% - WARNING!
[1460338094] SERVICE ALERT: hallegroundcolor;Cyan Cartridge HP CE261A;WARNING;SOFT;2;Cyan Cartridge HP CE261A is at 2% - WARNING!
[1460338153] SERVICE ALERT: hallegroundcolor;Cyan Cartridge HP CE261A;CRITICAL;HARD;3;Cyan Cartridge HP CE261A is at CRITICAL level!
[1460339003] SERVICE ALERT: hallefirst1;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460339063] SERVICE ALERT: hallefirst1;Printer Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460339124] SERVICE ALERT: hallefirst1;Printer Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460339723] SERVICE ALERT: hallefirst1;Printer Status;OK;HARD;3;ORDER CARTRIDGE LESS THAN 1300 PAGES 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1460341358] SERVICE ALERT: LANSWEEPER;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460341358] SERVICE ALERT: LANSWEEPER;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460341458] SERVICE ALERT: LANSWEEPER;Uptime;OK;SOFT;2;System Uptime - 24 day(s) 23 hour(s) 10 minute(s)
[1460341458] SERVICE ALERT: LANSWEEPER;C:\ Drive Space;OK;SOFT;2;c:\ - total: 79.66 Gb - used: 30.17 Gb (38%) - free 49.49 Gb (62%)
[1460341593] SERVICE ALERT: halleground2;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460341653] SERVICE ALERT: halleground2;Printer Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460341713] SERVICE ALERT: halleground2;Printer Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460342796] SERVICE ALERT: VC;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460342896] SERVICE ALERT: VC;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1460342913] SERVICE ALERT: halleground2;Printer Status;OK;HARD;3;Ready
[1460343543] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1460343663] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 91% (5 min average)
[1460343783] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 83% (5 min average)
[1460344023] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 69% (5 min average)
[1460344383] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1460344503] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 97% (5 min average)
[1460344623] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 88% (5 min average)
[1460344623] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460344743] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1460344863] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460345223] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1460345343] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1460345463] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1460345463] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460346063] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)