[1460780690] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1460780870] SERVICE ALERT: ADWU1;Windows Update Services Event 364;CRITICAL;SOFT;1;Windows Server Update Services, Content file download failed. Reason: File cert verification failure.  Source File: /d/msdownload/update/software/crup/2016/04/windows10.0-kb3154879-x86_c7b4d1ac4829121785ae79b2df7550004eca089f.cab  Destination File: C:\WSUSupdates\WsusContent\9F\C7B4D1AC4829121785AE79B2DF7550004ECA089F.cab
[1460780990] SERVICE ALERT: ADWU1;Windows Update Services Event 364;CRITICAL;SOFT;2;Windows Server Update Services, Content file download failed. Reason: File cert verification failure.  Source File: /d/msdownload/update/software/crup/2016/04/windows10.0-kb3154879-x86_c7b4d1ac4829121785ae79b2df7550004eca089f.cab  Destination File: C:\WSUSupdates\WsusContent\9F\C7B4D1AC4829121785AE79B2DF7550004ECA089F.cab
[1460781110] SERVICE ALERT: ADWU1;Windows Update Services Event 364;CRITICAL;HARD;3;Windows Server Update Services, Content file download failed. Reason: File cert verification failure.  Source File: /d/msdownload/update/software/crup/2016/04/windows10.0-kb3154879-x86_c7b4d1ac4829121785ae79b2df7550004eca089f.cab  Destination File: C:\WSUSupdates\WsusContent\9F\C7B4D1AC4829121785AE79B2DF7550004ECA089F.cab
[1460781290] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1460783294] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 364;CRITICAL;SOFT;1;Windows Server Update Services, eventlog: 1 > critical
[1460783414] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 364;CRITICAL;SOFT;2;Windows Server Update Services(error, 364, success)[Content file download failed.
[1460783534] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 364;CRITICAL;HARD;3;Windows Server Update Services(error, 364, success)[Content file download failed.
[1460786204] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1460786262] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;OK;HARD;3;OK: No entries found
[1460786324] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 76% (5 min average)
[1460786924] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1460787044] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460787164] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 81% (5 min average)
[1460787710] SERVICE ALERT: ADWU1;Windows Update Services Event 364;OK;HARD;3;OK: No entries found
[1460787764] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1460787884] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1460788004] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1460788604] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 94% (5 min average)
[1460788604] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1460788875] SERVICE ALERT: NETIDADPROXY;Service: Windows Time;OK;HARD;3;W32Time: Started
[1460789475] SERVICE ALERT: NETIDADPROXY;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1460789595] SERVICE ALERT: NETIDADPROXY;Service: Windows Time;CRITICAL;SOFT;2;W32Time: Stopped
[1460789715] SERVICE ALERT: NETIDADPROXY;Service: Windows Time;CRITICAL;HARD;3;W32Time: Stopped
[1460789804] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 83% (5 min average)
[1460790134] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 364;OK;HARD;3;Eventlog check ok
[1460790404] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1460790462] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;CRITICAL;SOFT;1;Windows Server Update Services, The server is failing to download some updates.
[1460790524] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 94% (5 min average)
[1460790582] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;CRITICAL;SOFT;2;Windows Server Update Services, The server is failing to download some updates.
[1460790644] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 77% (5 min average)
[1460790702] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;CRITICAL;HARD;3;Windows Server Update Services, The server is failing to download some updates.
[1460791244] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460791364] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 96% (5 min average)
[1460791430] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;CRITICAL;SOFT;1;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1460791484] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1460791551] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;CRITICAL;SOFT;2;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1460791670] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;CRITICAL;HARD;3;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1460792684] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 92% (5 min average)
[1460792690] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 94% (5 min average)
[1460793284] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 77% (5 min average)
[1460793290] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 97% (5 min average)
[1460795084] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1460795204] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 81% (5 min average)
[1460795690] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 82% (5 min average)
[1460797302] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;OK;HARD;3;OK: No entries found
[1460798690] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1460798804] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1460798810] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 78% (5 min average)
[1460798869] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;OK;HARD;3;Eventlog check ok
[1460798924] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 88% (5 min average)
[1460799410] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1460799524] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 91% (5 min average)
[1460799530] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 95% (5 min average)
[1460799644] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 89% (5 min average)
[1460799650] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1460800250] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 24% (5 min average)
[1460801450] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460801570] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1460801690] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460804090] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 50% (5 min average)
[1460805044] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1460805290] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1460805410] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1460805530] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 97% (5 min average)
[1460806130] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 92% (5 min average)
[1460806730] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1460809130] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 94% (5 min average)
[1460809730] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1460810330] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 94% (5 min average)
[1460810930] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 32% (5 min average)
[1460811701] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;CRITICAL;SOFT;1;Windows Server Update Services, The server is failing to download some updates.
[1460811821] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;CRITICAL;SOFT;2;Windows Server Update Services, The server is failing to download some updates.
[1460811941] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;CRITICAL;HARD;3;Windows Server Update Services, The server is failing to download some updates.
[1460812730] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1460812850] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460812970] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460813269] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;CRITICAL;SOFT;1;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1460813390] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;CRITICAL;SOFT;2;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1460813509] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;CRITICAL;HARD;3;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1460814170] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 60% (5 min average)
[1460814770] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1460814890] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460815010] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 85% (5 min average)
[1460816810] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1460816930] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460817050] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460817650] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 92% (5 min average)
[1460818250] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460818850] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 92% (5 min average)
[1460819141] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;OK;HARD;3;OK: No entries found
[1460819450] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 82% (5 min average)
[1460820108] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;OK;HARD;3;Eventlog check ok
[1460821850] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1460821970] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 95% (5 min average)
[1460822090] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1460822690] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 88% (5 min average)
[1460822949] SERVICE ALERT: INTLAPP;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1460823069] SERVICE ALERT: INTLAPP;CPU Load;OK;SOFT;2;CPU Load 16% (5 min average)
[1460823290] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460823410] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 72% (5 min average)
[1460824010] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460824129] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 79% (5 min average)
[1460825329] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460825449] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 95% (5 min average)
[1460825569] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1460827969] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1460828569] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 71% (5 min average)
[1460830369] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 91% (5 min average)
[1460830489] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460830524] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460830584] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460830609] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460830644] SERVICE ALERT: hallefirst2;Device Status;OK;SOFT;3;Clearing paper path
[1460833541] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;CRITICAL;SOFT;1;Windows Server Update Services, The server is failing to download some updates.
[1460833661] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;CRITICAL;SOFT;2;Windows Server Update Services, The server is failing to download some updates.
[1460833781] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;CRITICAL;HARD;3;Windows Server Update Services, The server is failing to download some updates.
[1460834507] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;CRITICAL;SOFT;1;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1460834627] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;CRITICAL;SOFT;2;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1460834748] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;CRITICAL;HARD;3;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1460835094] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460835155] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460835214] SERVICE ALERT: hallefirst1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460835814] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;ORDER CARTRIDGE LESS THAN 900 PAGES 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1460836008] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 62% (5 min average)
[1460836608] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460836728] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1460836848] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460838648] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 94% (5 min average)
[1460839248] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 80% (5 min average)
[1460839848] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 94% (5 min average)
[1460839968] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460840088] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1460840979] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;OK;HARD;3;OK: No entries found
[1460841288] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 86% (5 min average)
[1460841888] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460841947] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;OK;HARD;3;Eventlog check ok
[1460842008] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 86% (5 min average)
[1460842132] SERVICE ALERT: SNAP;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1460842135] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.40)
[1460842140] SERVICE ALERT: SNAP;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1460842140] SERVICE ALERT: SNAP;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1460842160] SERVICE ALERT: SNAP;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1460842163] HOST ALERT: SNAP;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.40)
[1460842214] HOST ALERT: SNAP;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1460842260] SERVICE ALERT: SNAP;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460842271] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.40)
[1460842277] SERVICE ALERT: SNAP;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460842278] SERVICE ALERT: SNAP;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1460842278] SERVICE ALERT: SNAP;Service: Group Policy;CRITICAL;SOFT;2;connect to address 164.76.102.40 and port 12489: No route to host
[1460842278] HOST ALERT: SNAP;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.40)
[1460842279] SERVICE ALERT: SNAP;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460842332] HOST ALERT: SNAP;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.41 ms
[1460842332] HOST FLAPPING ALERT: SNAP;STARTED; Host appears to have started flapping (20.1% change > 20.0% threshold)
[1460842380] SERVICE ALERT: SNAP;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460842388] SERVICE ALERT: SNAP;NSClient++ Version;CRITICAL;HARD;3;connect to address 164.76.102.40 and port 12489: No route to host
[1460842388] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.40)
[1460842390] SERVICE ALERT: SNAP;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460842397] SERVICE ALERT: SNAP;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460842398] SERVICE ALERT: SNAP;Service: Windows Time;CRITICAL;SOFT;2;connect to address 164.76.102.40 and port 12489: No route to host
[1460842398] HOST ALERT: SNAP;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.40)
[1460842398] SERVICE ALERT: SNAP;Service: Group Policy;CRITICAL;HARD;3;connect to address 164.76.102.40 and port 12489: No route to host
[1460842490] SERVICE ALERT: SNAP;C:\ Drive Space;OK;SOFT;2;c:\ - total: 1116.65 Gb - used: 35.74 Gb (3%) - free 1080.91 Gb (97%)
[1460842492] HOST ALERT: SNAP;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1460842515] SERVICE ALERT: SNAP;Service: Windows Time;OK;SOFT;3;W32Time: Started
[1460842608] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460842728] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460842848] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460842960] SERVICE ALERT: SNAP;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1460842977] SERVICE ALERT: SNAP;Service: RPC;OK;HARD;3;RpcSs: Started
[1460842979] SERVICE ALERT: SNAP;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1460842997] SERVICE ALERT: SNAP;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1460842997] HOST FLAPPING ALERT: SNAP;STOPPED; Host appears to have stopped flapping (4.7% change < 5.0% threshold)
[1460843700] SERVICE ALERT: SNAP;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1460843700] SERVICE ALERT: SNAP;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1460843735] SERVICE ALERT: SNAP;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460843746] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.40)
[1460843774] SERVICE ALERT: SNAP;Uptime;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1460843777] HOST ALERT: SNAP;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.40)
[1460843811] SERVICE ALERT: SNAP;C:\ Drive Space;CRITICAL;SOFT;2;connect to address 164.76.102.40 and port 12489: No route to host
[1460843811] SERVICE ALERT: SNAP;Service: WMI;CRITICAL;SOFT;2;connect to address 164.76.102.40 and port 12489: No route to host
[1460843814] HOST ALERT: SNAP;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.102.40)
[1460843835] SERVICE ALERT: SNAP;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1460843839] HOST ALERT: SNAP;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.40 ms
[1460843882] SERVICE ALERT: SNAP;Uptime;OK;SOFT;2;System Uptime - 7 day(s) 23 hour(s) 9 minute(s)
[1460843928] SERVICE ALERT: SNAP;Service: WMI;OK;SOFT;3;Winmgmt: Started
[1460843930] SERVICE ALERT: SNAP;C:\ Drive Space;OK;SOFT;3;c:\ - total: 1116.65 Gb - used: 35.74 Gb (3%) - free 1080.91 Gb (97%)
[1460847622] SERVICE ALERT: SNAP;CPU Load;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1460847622] SERVICE ALERT: SNAP;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1460847622] SERVICE ALERT: SNAP;Service: NetLogin;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1460847625] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.40)
[1460847689] HOST ALERT: SNAP;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1460847730] SERVICE ALERT: SNAP;Memory Usage;OK;SOFT;2;Memory usage: total:16362.77 MB - used: 1784.95 MB (11%) - free: 14577.82 MB (89%)
[1460847730] SERVICE ALERT: SNAP;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1460847737] SERVICE ALERT: SNAP;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1460847780] SERVICE ALERT: SNAP;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460847817] SERVICE ALERT: SNAP;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460847880] SERVICE ALERT: SNAP;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1460847926] SERVICE ALERT: SNAP;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1460848047] SERVICE ALERT: SNAP;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1460848175] SERVICE ALERT: SNAP;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460848187] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.40)
[1460848258] HOST ALERT: SNAP;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.40)
[1460848278] SERVICE ALERT: SNAP;Service: Windows Time;CRITICAL;HARD;3;connect to address 164.76.102.40 and port 12489: No route to host
[1460848281] HOST ALERT: SNAP;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.102.40)
[1460848332] HOST ALERT: SNAP;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.40 ms
[1460848492] SERVICE ALERT: SNAP;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1460848495] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.40)
[1460848520] SERVICE ALERT: SNAP;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1460848523] HOST ALERT: SNAP;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.40)
[1460848587] HOST ALERT: SNAP;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.46 ms
[1460848587] HOST FLAPPING ALERT: SNAP;STARTED; Host appears to have started flapping (21.4% change > 20.0% threshold)
[1460848600] SERVICE ALERT: SNAP;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1460848637] SERVICE ALERT: SNAP;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1460848890] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.40)
[1460848932] HOST ALERT: SNAP;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.41 ms
[1460849475] SERVICE ALERT: SNAP;Service: Windows Time;OK;HARD;3;W32Time: Started
[1460849537] HOST FLAPPING ALERT: SNAP;STOPPED; Host appears to have stopped flapping (4.3% change < 5.0% threshold)
[1460851649] SERVICE ALERT: SNAP;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1460851652] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.40)
[1460851686] HOST ALERT: SNAP;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.41 ms
[1460851757] SERVICE ALERT: SNAP;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1460852448] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 94% (5 min average)
[1460853048] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1460855367] SERVICE ALERT: BOEAPPTEST;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1460855379] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;CRITICAL;SOFT;1;Windows Server Update Services, The server is failing to download some updates.
[1460855487] SERVICE ALERT: BOEAPPTEST;CPU Load;OK;SOFT;2;CPU Load 12% (5 min average)
[1460855499] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;CRITICAL;SOFT;2;Windows Server Update Services, The server is failing to download some updates.
[1460855619] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;CRITICAL;HARD;3;Windows Server Update Services, The server is failing to download some updates.
[1460856347] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;CRITICAL;SOFT;1;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1460856468] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;CRITICAL;SOFT;2;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1460856587] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;CRITICAL;HARD;3;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1460858968] SERVICE ALERT: SNAP;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1460858971] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.40)
[1460859042] HOST ALERT: SNAP;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.40)
[1460859076] SERVICE ALERT: SNAP;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1460859078] HOST ALERT: SNAP;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.44 ms
[1460859646] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 56% (5 min average)
[1460859809] SERVICE ALERT: INTLTESTAPP;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460859852] SERVICE ALERT: LANSWEEPER;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460859909] SERVICE ALERT: INTLTESTAPP;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1460859952] SERVICE ALERT: LANSWEEPER;Uptime;OK;SOFT;2;System Uptime - 1 day(s) 6 hour(s) 57 minute(s)
[1460860246] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460860339] SERVICE ALERT: SNAP;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1460860339] SERVICE ALERT: SNAP;CPU Load;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1460860339] SERVICE ALERT: SNAP;Service: NetLogin;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1460860342] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.40)
[1460860347] SERVICE ALERT: SNAP;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460860348] SERVICE ALERT: SNAP;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460860366] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460860380] HOST ALERT: SNAP;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1460860447] SERVICE ALERT: SNAP;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1460860448] SERVICE ALERT: SNAP;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1460860449] SERVICE ALERT: SNAP;Memory Usage;OK;SOFT;2;Memory usage: total:16362.77 MB - used: 1788.06 MB (11%) - free: 14574.71 MB (89%)
[1460860449] SERVICE ALERT: SNAP;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1460860455] SERVICE ALERT: SNAP;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1460860486] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460862040] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460862160] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 84% (5 min average)
[1460862219] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;OK;HARD;3;OK: No entries found
[1460862760] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1460862880] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 79% (5 min average)
[1460863786] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;OK;HARD;3;Eventlog check ok
[1460865286] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 94% (5 min average)