[1460347263] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 77% (5 min average)
[1460350863] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 76% (5 min average)
[1460351463] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460351583] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460351703] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460352179] SERVICE ALERT: eateriespk;Printer Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.58.5. Is the printer off?  Is the network down?
[1460352229] SERVICE ALERT: eateriespk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.58.5. Is the printer off?  Is the network down?
[1460352239] SERVICE ALERT: eateriespk;Printer Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.58.5. Is the printer off?  Is the network down?
[1460352289] SERVICE ALERT: eateriespk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.58.5. Is the printer off?  Is the network down?
[1460352299] SERVICE ALERT: eateriespk;Printer Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.58.5. Is the printer off?  Is the network down?
[1460352349] SERVICE ALERT: eateriespk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.58.5. Is the printer off?  Is the network down?
[1460352753] SERVICE ALERT: eateriespk;Black Cartridge HP CC364X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.58.5. Is the printer off?  Is the network down?
[1460352813] SERVICE ALERT: eateriespk;Black Cartridge HP CC364X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.58.5. Is the printer off?  Is the network down?
[1460352873] SERVICE ALERT: eateriespk;Black Cartridge HP CC364X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.58.5. Is the printer off?  Is the network down?
[1460353863] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460353983] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 89% (5 min average)
[1460354583] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1460354703] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 94% (5 min average)
[1460354823] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 85% (5 min average)
[1460355423] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1460355543] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1460355663] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 89% (5 min average)
[1460356263] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1460356383] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 97% (5 min average)
[1460356503] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 97% (5 min average)
[1460357319] SERVICE ALERT: halleground2;Printer Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1460357349] SERVICE ALERT: halleground1;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1460357369] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1460357379] SERVICE ALERT: halleground2;Printer Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1460357379] SERVICE ALERT: halleground2;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1460357409] SERVICE ALERT: halleground1;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1460357429] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1460357439] SERVICE ALERT: halleground2;Printer Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1460357439] SERVICE ALERT: halleground2;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1460357469] SERVICE ALERT: halleground1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1460357469] SERVICE ALERT: halleground3;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1460357469] SERVICE ALERT: halleground1;Maintenance Kit HP 110V-C9152A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1460357489] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1460357499] SERVICE ALERT: halleground2;Maintenance Kit HP 110V-C9152A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1460357519] SERVICE ALERT: halleground3;Printer Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1460357529] SERVICE ALERT: halleground1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1460357529] SERVICE ALERT: halleground3;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1460357529] HOST ALERT: halleground2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460357539] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1460357569] HOST ALERT: halleground2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460357579] SERVICE ALERT: halleground3;Printer Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1460357579] HOST ALERT: halleground2;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.52 ms
[1460357589] SERVICE ALERT: halleground1;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1460357589] SERVICE ALERT: halleground3;Maintenance Kit HP 110V-C9152A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1460357599] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1460357639] SERVICE ALERT: halleground3;Printer Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1460357659] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1460357689] SERVICE ALERT: halleground1;Printer Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1460357703] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 93% (5 min average)
[1460357703] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 94% (5 min average)
[1460357703] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1460357746] SERVICE ALERT: halleground3;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1460357749] SERVICE ALERT: halleground1;Printer Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1460357806] SERVICE ALERT: halleground3;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1460357809] SERVICE ALERT: halleground1;Printer Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1460357839] HOST ALERT: halleground1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460357866] SERVICE ALERT: halleground3;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1460357899] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1460357903] HOST ALERT: halleground1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1460357959] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1460358020] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1460358179] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1460358240] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1460358299] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1460358303] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460358303] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 97% (5 min average)
[1460359233] HOST ALERT: halleground2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1460359254] HOST ALERT: halleground2;UP;SOFT;2;PING OK - Packet loss = 68%, RTA = 3.22 ms
[1460359503] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 90% (5 min average)
[1460359753] SERVICE ALERT: hallegroundcolor;Cyan Cartridge HP CE261A;OK;HARD;3;
[1460360103] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 82% (5 min average)
[1460360353] SERVICE ALERT: hallegroundcolor;Cyan Cartridge HP CE261A;CRITICAL;SOFT;1;Cyan Cartridge HP CE261A is at CRITICAL level!
[1460360414] SERVICE ALERT: hallegroundcolor;Cyan Cartridge HP CE261A;CRITICAL;SOFT;2;Cyan Cartridge HP CE261A is at CRITICAL level!
[1460360473] SERVICE ALERT: hallegroundcolor;Cyan Cartridge HP CE261A;CRITICAL;HARD;3;Cyan Cartridge HP CE261A is at CRITICAL level!
[1460360703] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 75% (5 min average)
[1460360703] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1460360823] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 95% (5 min average)
[1460360943] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1460361303] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460361423] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460361543] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460361543] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 89% (5 min average)
[1460362743] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 55% (5 min average)
[1460365743] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460365863] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1460365983] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460367183] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 74% (5 min average)
[1460367783] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460367903] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 82% (5 min average)
[1460369103] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1460369223] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1460369343] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460369943] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 66% (5 min average)
[1460371143] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460371263] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460371383] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 83% (5 min average)
[1460371983] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460372103] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1460372223] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1460372943] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1460373063] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 81% (5 min average)
[1460373063] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1460373663] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460373783] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 70% (5 min average)
[1460374023] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 65% (5 min average)
[1460374623] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1460374743] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460374863] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460375723] SERVICE ALERT: hallefirst1;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460375783] SERVICE ALERT: hallefirst1;Printer Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460375844] SERVICE ALERT: hallefirst1;Printer Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460376183] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1460376303] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 90% (5 min average)
[1460376423] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460376443] SERVICE ALERT: hallefirst1;Printer Status;OK;HARD;3;ORDER CARTRIDGE LESS THAN 1100 PAGES 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1460377023] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 75% (5 min average)
[1460377218] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 9% - OK!
[1460377233] SERVICE ALERT: halleground2;Printer Status;OK;HARD;3;Warming up
[1460377263] SERVICE ALERT: halleground1;Maintenance Kit HP 110V-C9152A;OK;HARD;3;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 29% - OK!
[1460377283] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 24% - OK!
[1460377293] SERVICE ALERT: halleground2;Maintenance Kit HP 110V-C9152A;OK;HARD;3;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 45% - OK!
[1460377383] SERVICE ALERT: halleground3;Maintenance Kit HP 110V-C9152A;OK;HARD;3;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 84% - OK!
[1460377383] SERVICE ALERT: halleground1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1460377433] SERVICE ALERT: halleground3;Printer Status;OK;HARD;3;Ready
[1460377453] SERVICE ALERT: halleground2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1460377493] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 56% - OK!
[1460377603] SERVICE ALERT: halleground1;Printer Status;OK;HARD;3;Ready
[1460377623] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1460377660] SERVICE ALERT: halleground3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1460377743] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 73% (5 min average)
[1460378143] HOST ALERT: VCOMPOSER;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460378160] SERVICE ALERT: VCOMPOSER;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460378190] HOST ALERT: VCOMPOSER;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460378204] SERVICE ALERT: VCOMPOSER;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460378209] SERVICE ALERT: VCOMPOSER;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460378234] HOST ALERT: VCOMPOSER;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460378243] SERVICE ALERT: VCOMPOSER;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460378246] SERVICE ALERT: VCOMPOSER;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460378273] HOST ALERT: VCOMPOSER;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460378280] SERVICE ALERT: VCOMPOSER;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460378310] HOST ALERT: VCOMPOSER;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460378317] SERVICE ALERT: VCOMPOSER;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460378324] SERVICE ALERT: VCOMPOSER;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460378329] SERVICE ALERT: VCOMPOSER;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460378335] SERVICE ALERT: VCOMPOSER;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460378343] SERVICE ALERT: VCOMPOSER;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460378347] HOST ALERT: VCOMPOSER;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1460378363] SERVICE ALERT: VCOMPOSER;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460378366] SERVICE ALERT: VCOMPOSER;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460378385] SERVICE ALERT: VCOMPOSER;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460378393] HOST ALERT: VCOMPOSER;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1460378400] SERVICE ALERT: VCOMPOSER;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460378430] HOST ALERT: VCOMPOSER;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1460378436] SERVICE ALERT: VCOMPOSER;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460378437] SERVICE ALERT: VCOMPOSER;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460378444] SERVICE ALERT: VCOMPOSER;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460378449] SERVICE ALERT: VCOMPOSER;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460378455] SERVICE ALERT: VCOMPOSER;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460378463] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 70% (5 min average)
[1460378463] SERVICE ALERT: VCOMPOSER;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460378466] HOST ALERT: VCOMPOSER;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1460378483] SERVICE ALERT: VCOMPOSER;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460378483] SERVICE ALERT: VCOMPOSER;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460378486] SERVICE ALERT: VCOMPOSER;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460378505] SERVICE ALERT: VCOMPOSER;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460378513] HOST ALERT: VCOMPOSER;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1460378523] SERVICE ALERT: VCOMPOSER;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460378556] SERVICE ALERT: VCOMPOSER;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460378557] SERVICE ALERT: VCOMPOSER;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460378575] SERVICE ALERT: VCOMPOSER;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460378583] SERVICE ALERT: VCOMPOSER;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460378603] SERVICE ALERT: VCOMPOSER;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460378606] SERVICE ALERT: SP2010;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460378615] SERVICE ALERT: VCOMPOSER;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460378615] SERVICE ALERT: SP2010;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460378617] SERVICE ALERT: SP2010DB;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460378625] SERVICE ALERT: VCOMPOSER;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460378636] HOST ALERT: SP2010;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460378643] SERVICE ALERT: VCOMPOSER;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460378647] HOST ALERT: SP2010DB;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460378663] SERVICE ALERT: SP2010;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460378663] SERVICE ALERT: VCOMPOSER;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460378664] SERVICE ALERT: SP2010DB;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460378672] HOST DOWNTIME ALERT: SP2010;STARTED; Host has entered a period of scheduled downtime
[1460378676] SERVICE ALERT: VCOMPOSER;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460378684] HOST DOWNTIME ALERT: SP2010DB;STARTED; Host has entered a period of scheduled downtime
[1460378693] HOST ALERT: SP2010;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460378694] SERVICE ALERT: SP2010DB;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460378694] HOST ALERT: SP2010DB;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460378705] SERVICE ALERT: SP2010;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460378706] SERVICE ALERT: SP2010DB;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460378723] SERVICE ALERT: VCOMPOSER;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460378726] SERVICE ALERT: SP2010;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460378735] SERVICE ALERT: VCOMPOSER;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460378735] HOST ALERT: SP2010;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460378735] SERVICE ALERT: SP2010;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460378736] HOST ALERT: SP2010DB;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460378737] SERVICE ALERT: SP2010DB;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460378748] SERVICE ALERT: SP2010;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460378754] SERVICE ALERT: SP2010DB;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460378763] SERVICE ALERT: VCOMPOSER;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460378765] HOST ALERT: SP2010;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460378766] SERVICE ALERT: SP2010DB;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460378767] HOST ALERT: SP2010DB;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460378783] SERVICE ALERT: SP2010;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460378783] SERVICE ALERT: VCOMPOSER;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460378784] SERVICE ALERT: SP2010DB;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460378802] SERVICE ALERT: SP2010;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460378802] SERVICE ALERT: SP2010DB;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460378813] HOST ALERT: SP2010;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460378814] SERVICE ALERT: SP2010DB;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460378814] HOST ALERT: SP2010DB;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460378825] SERVICE ALERT: SP2010;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460378825] SERVICE ALERT: SP2010DB;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460378846] SERVICE ALERT: SP2010;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460378850] SERVICE ALERT: SP2010DB;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460378855] SERVICE ALERT: VCOMPOSER;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460378855] HOST ALERT: SP2010;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1460378855] SERVICE ALERT: SP2010;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460378855] HOST ALERT: SP2010DB;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1460378857] SERVICE ALERT: SP2010DB;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460378861] SERVICE ALERT: SP2010;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460378868] SERVICE ALERT: SP2010;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460378874] SERVICE ALERT: SP2010DB;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460378885] HOST ALERT: SP2010;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1460378886] SERVICE ALERT: SP2010DB;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460378887] HOST ALERT: SP2010DB;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1460378903] SERVICE ALERT: SP2010;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460378903] SERVICE ALERT: VCOMPOSER;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460378904] SERVICE ALERT: SP2010DB;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460378922] SERVICE ALERT: SP2010;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460378922] SERVICE ALERT: SP2010DB;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460378933] SERVICE ALERT: SP2010;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460378933] HOST ALERT: SP2010;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1460378934] SERVICE ALERT: SP2010DB;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460378934] HOST ALERT: SP2010DB;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1460378945] SERVICE ALERT: SP2010;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460378945] SERVICE ALERT: SP2010DB;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460378967] SERVICE ALERT: SP2010;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460378970] SERVICE ALERT: SP2010;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460378970] SERVICE ALERT: SP2010DB;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460378975] HOST ALERT: SP2010;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1460378975] HOST ALERT: SP2010DB;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1460378980] SERVICE ALERT: SP2010;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460378981] SERVICE ALERT: SP2010DB;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460378981] SERVICE ALERT: SP2010;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460378988] SERVICE ALERT: SP2010;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460378994] SERVICE ALERT: SP2010DB;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460379006] SERVICE ALERT: SP2010DB;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460379010] HOST ALERT: SP2010;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1460379011] HOST ALERT: SP2010DB;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1460379024] SERVICE ALERT: SP2010DB;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460379025] SERVICE ALERT: SP2010;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460379041] SERVICE ALERT: SP2010;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460379042] SERVICE ALERT: SP2010DB;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460379053] SERVICE ALERT: SP2010;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460379083] SERVICE ALERT: SP2010;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460379083] SERVICE ALERT: SP2010DB;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460379087] SERVICE ALERT: SP2010;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460379090] SERVICE ALERT: SP2010;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460379090] SERVICE ALERT: SP2010DB;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460379100] SERVICE ALERT: SP2010;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460379101] SERVICE ALERT: SP2010DB;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460379101] SERVICE ALERT: SP2010;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460379144] SERVICE ALERT: SP2010DB;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460379145] SERVICE ALERT: SP2010;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460379161] SERVICE ALERT: SP2010;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460379163] SERVICE ALERT: SP2010DB;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460379168] SERVICE ALERT: SP2010DB;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460379173] SERVICE ALERT: SP2010;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460379203] SERVICE ALERT: SP2010;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460379203] SERVICE ALERT: SP2010DB;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460379203] SERVICE ALERT: SP2010DB;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460379207] SERVICE ALERT: SP2010;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460379210] SERVICE ALERT: SP2010;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460379220] SERVICE ALERT: SP2010;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460379221] SERVICE ALERT: SP2010DB;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460379264] SERVICE ALERT: SP2010DB;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460379265] SERVICE ALERT: SP2010;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460379281] SERVICE ALERT: SP2010;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460379283] SERVICE ALERT: SP2010DB;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460379287] SERVICE ALERT: SP2010DB;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460379323] SERVICE ALERT: SP2010;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460379323] SERVICE ALERT: SP2010DB;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460379323] SERVICE ALERT: SP2010DB;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460379401] SERVICE ALERT: SP2010;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460379403] SERVICE ALERT: SP2010DB;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460379406] SERVICE ALERT: SP2010DB;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460379443] SERVICE ALERT: SP2010DB;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460380263] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460380383] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460380429] SERVICE ALERT: owen101lab1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460380439] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460380489] SERVICE ALERT: owen101lab1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460380499] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460380503] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1460380541] SERVICE ALERT: owen101lab1;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460380549] SERVICE ALERT: owen101lab1;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460380560] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460380601] SERVICE ALERT: owen101lab1;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460380661] SERVICE ALERT: owen101lab1;Maintenance Kit HP 110V-C9152A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460380719] SERVICE ALERT: owen101lab1;Printer Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460380779] SERVICE ALERT: owen101lab1;Printer Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460380839] SERVICE ALERT: owen101lab1;Printer Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460381469] HOST ALERT: owen101lab1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460381533] HOST ALERT: owen101lab1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.54 ms
[1460381746] SERVICE ALERT: owen101lab1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1460381753] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 36% - OK!
[1460381863] HOST ALERT: owen101lab1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460381953] HOST ALERT: owen101lab1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460382043] HOST ALERT: owen101lab1;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460382133] HOST ALERT: owen101lab1;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460382223] HOST ALERT: owen101lab1;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460382303] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 92% (5 min average)
[1460382313] HOST ALERT: owen101lab1;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1460382349] SERVICE ALERT: owen101lab1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460382359] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460382379] HOST ALERT: owen101lab1;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1460382409] SERVICE ALERT: owen101lab1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460382417] HOST ALERT: owen101lab1;UP;SOFT;8;PING OK - Packet loss = 44%, RTA = 0.76 ms
[1460382418] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 36% - OK!
[1460382455] SERVICE ALERT: owen101lab1;Maintenance Kit HP 110V-C9152A;OK;HARD;3;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 56% - OK!
[1460382463] SERVICE ALERT: owen101lab1;Paper Tray Status;OK;SOFT;3;Printer Paper Is OK! All Trays Still Full
[1460382669] HOST ALERT: owen101lab1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460382695] SERVICE ALERT: owen101lab3;Printer Status;OK;HARD;3;Clearing paper path
[1460382699] HOST ALERT: owen101lab3;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.47 ms
[1460382759] HOST ALERT: owen101lab1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460382849] HOST ALERT: owen101lab1;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460382939] HOST ALERT: owen101lab1;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460383019] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460383029] HOST ALERT: owen101lab1;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460383061] SERVICE ALERT: owen101lab1;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460383069] SERVICE ALERT: owen101lab1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460383079] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460383083] SERVICE ALERT: owen101lab3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1460383091] HOST ALERT: owen101lab1;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1460383122] SERVICE ALERT: owen101lab1;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460383129] SERVICE ALERT: owen101lab1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460383139] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460383152] HOST ALERT: owen101lab1;DOWN;SOFT;7;(Host check timed out after 30.01 seconds)
[1460383181] SERVICE ALERT: owen101lab1;Maintenance Kit HP 110V-C9152A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460383189] SERVICE ALERT: owen101lab1;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460383211] HOST ALERT: owen101lab1;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1460383269] HOST ALERT: owen101lab1;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1460383295] SERVICE ALERT: owen101lab3;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460383359] HOST ALERT: owen101lab1;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1460383361] SERVICE ALERT: owen101lab3;Printer Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1460383391] HOST ALERT: owen101lab3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460383421] SERVICE ALERT: owen101lab3;Printer Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1460383451] HOST ALERT: owen101lab3;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460383503] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460383541] HOST ALERT: owen101lab3;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460383631] HOST ALERT: owen101lab3;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460383663] HOST ALERT: owen101lab1;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.67 ms
[1460383689] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1460383719] HOST ALERT: owen101lab3;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460383749] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1460383777] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1460383779] HOST ALERT: owen101lab3;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1460383809] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1460383838] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1460383839] HOST ALERT: owen101lab3;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1460383897] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1460383929] HOST ALERT: owen101lab3;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1460384019] HOST ALERT: owen101lab3;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1460384051] HOST ALERT: owen101lab3;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1460384103] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 62% (5 min average)
[1460384703] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460384823] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1460384943] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460385002] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (42.5% change >= 20.0% threshold)
[1460385002] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1460385002] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1460385002] HOST DOWNTIME ALERT: SP2010DB;STARTED; Host has entered a period of scheduled downtime
[1460385002] HOST DOWNTIME ALERT: SP2010;STARTED; Host has entered a period of scheduled downtime
[1460385663] SERVICE ALERT: hallefirst2;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460385722] SERVICE ALERT: hallefirst2;Printer Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460385782] SERVICE ALERT: hallefirst2;Printer Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460385866] HOST DOWNTIME ALERT: SP2010;STOPPED; Host has exited from a period of scheduled downtime
[1460385881] HOST DOWNTIME ALERT: SP2010DB;STOPPED; Host has exited from a period of scheduled downtime
[1460385959] HOST DOWNTIME ALERT: SP2010;STARTED; Host has entered a period of scheduled downtime
[1460385975] HOST DOWNTIME ALERT: SP2010DB;STARTED; Host has entered a period of scheduled downtime
[1460386143] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 63% (5 min average)
[1460386382] SERVICE ALERT: hallefirst2;Printer Status;OK;HARD;3;Processing job
[1460386478] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1460386537] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1460386598] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1460386743] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1460386863] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 56% (5 min average)
[1460386897] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;CRITICAL;HARD;3;Black Print Cartridge HP C8543X is at 1% - CRITICAL!
[1460386982] SERVICE ALERT: hallefirst2;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460387001] SERVICE ALERT: SP2010DB;Memory Usage;OK;HARD;3;Memory usage: total:14333.44 MB - used: 1921.49 MB (13%) - free: 12411.95 MB (87%)
[1460387005] HOST ALERT: SP2010DB;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.40 ms
[1460387042] SERVICE ALERT: hallefirst2;Printer Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460387044] SERVICE ALERT: SP2010DB;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1460387069] HOST ALERT: SP2010;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1460387102] SERVICE ALERT: hallefirst2;Printer Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460387103] SERVICE ALERT: SP2010DB;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1460387103] SERVICE ALERT: SP2010;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1460387181] SERVICE ALERT: SP2010;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1460387183] SERVICE ALERT: SP2010DB;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1460387186] SERVICE ALERT: SP2010DB;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1460387223] SERVICE ALERT: SP2010DB;C:\ Drive Space;OK;HARD;3;c:\ - total: 69.90 Gb - used: 54.65 Gb (78%) - free 15.25 Gb (22%)
[1460387226] SERVICE ALERT: SP2010;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1460387235] SERVICE ALERT: SP2010;Service: RPC;OK;HARD;3;RpcSs: Started
[1460387237] SERVICE ALERT: SP2010DB;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1460387283] SERVICE ALERT: SP2010;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1460387284] SERVICE ALERT: SP2010DB;Service: RPC;OK;HARD;3;RpcSs: Started
[1460387325] SERVICE ALERT: SP2010;Service: WMI;OK;HARD;3;Winmgmt: Started
[1460387325] SERVICE ALERT: SP2010DB;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1460387368] SERVICE ALERT: SP2010;Service: Windows Time;OK;HARD;3;W32Time: Started
[1460387374] SERVICE ALERT: SP2010DB;Service: WMI;OK;HARD;3;Winmgmt: Started
[1460387421] SERVICE ALERT: SP2010;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 13 minute(s)
[1460387422] SERVICE ALERT: SP2010DB;Service: Windows Time;OK;HARD;3;W32Time: Started
[1460387463] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460387471] SERVICE ALERT: SP2010DB;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1460387553] SERVICE ALERT: SP2010;Memory Usage;OK;HARD;3;Memory usage: total:14333.44 MB - used: 2468.76 MB (17%) - free: 11864.68 MB (83%)
[1460387583] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1460387587] SERVICE ALERT: SP2010;C:\ Drive Space;OK;HARD;3;c:\ - total: 79.90 Gb - used: 54.60 Gb (68%) - free 25.30 Gb (32%)
[1460387590] SERVICE ALERT: SP2010;VMWare Tools;UNKNOWN;HARD;3;Failed to open service manager: 1115: A system shutdown is in progress.
[1460387646] SERVICE ALERT: SP2010;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1460387703] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460387773] HOST ALERT: owen101lab1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460387797] SERVICE ALERT: marshallpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1460387843] SERVICE ALERT: hallefirst1;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460387863] HOST ALERT: owen101lab1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460387903] SERVICE ALERT: hallefirst1;Printer Status;OK;SOFT;2;MENUS  RETRIEVE JOB  INFORMATION  PAPER HANDLING
[1460387914] SERVICE ALERT: SP2010DB;CPU Load;OK;HARD;3;CPU Load 1% (5 min average)
[1460387953] HOST ALERT: owen101lab1;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460388009] HOST ALERT: owen101lab1;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460388049] SERVICE ALERT: SP2010DB;VMWare Tools;OK;HARD;3;VMTools: Started
[1460388069] HOST ALERT: owen101lab1;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460388081] SERVICE ALERT: SP2010;CPU Load;OK;HARD;3;CPU Load 7% (5 min average)
[1460388139] HOST ALERT: owen101lab1;UP;SOFT;6;PING OK - Packet loss = 54%, RTA = 0.73 ms
[1460388190] SERVICE ALERT: SP2010;VMWare Tools;OK;HARD;3;VMTools: Started
[1460388200] SERVICE ALERT: SP2010;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1460388533] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 36% - OK!
[1460388533] SERVICE FLAPPING ALERT: owen101lab1;Black Print Cartridge HP C8543X;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1460388555] HOST ALERT: owen101lab3;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.45 ms
[1460388574] SERVICE ALERT: owen101lab1;Maintenance Kit HP 110V-C9152A;OK;HARD;3;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 56% - OK!
[1460388574] SERVICE FLAPPING ALERT: owen101lab1;Maintenance Kit HP 110V-C9152A;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1460388584] SERVICE ALERT: owen101lab1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1460388584] SERVICE FLAPPING ALERT: owen101lab1;Paper Tray Status;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1460388604] SERVICE ALERT: owen101lab3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1460388633] SERVICE ALERT: owen101lab1;Printer Status;OK;HARD;3;Ready
[1460388815] SERVICE ALERT: owen101lab3;Printer Status;OK;HARD;3;TRAY 2 EMPTY PLAIN LETTER   Ready
[1460389502] SERVICE ALERT: hallefirst2;Printer Status;OK;HARD;3;Clearing paper path
[1460389502] SERVICE FLAPPING ALERT: hallefirst2;Printer Status;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1460389597] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1460389658] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1460389718] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1460390085] SERVICE ALERT: AD-RMAD;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1460390102] SERVICE ALERT: hallefirst2;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460390103] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 84% (5 min average)
[1460390162] SERVICE ALERT: hallefirst2;Printer Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460390205] SERVICE ALERT: AD-RMAD;CPU Load;OK;SOFT;2;CPU Load 58% (5 min average)
[1460390222] SERVICE ALERT: hallefirst2;Printer Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460390317] SERVICE ALERT: marshallpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1460390317] SERVICE FLAPPING ALERT: marshallpk;Paper Tray Status;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1460390343] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1460390463] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460390583] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 69% (5 min average)
[1460390703] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460390823] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460390943] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460391083] HOST ALERT: eateriespk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460391173] HOST ALERT: eateriespk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460391189] SERVICE ALERT: halleground1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1460391249] SERVICE ALERT: halleground1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1460391263] HOST ALERT: eateriespk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460391279] HOST ALERT: halleground1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460391297] HOST ALERT: eateriespk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460391309] SERVICE ALERT: halleground1;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1460391318] HOST ALERT: halleground1;UP;SOFT;2;PING OK - Packet loss = 50%, RTA = 0.68 ms
[1460391329] HOST ALERT: eateriespk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460391379] HOST ALERT: eateriespk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1460391403] SERVICE ALERT: halleground1;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460391464] SERVICE ALERT: halleground1;Printer Status;OK;SOFT;2;Clearing paper path
[1460391469] HOST ALERT: eateriespk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1460391531] HOST ALERT: halleground2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460391559] HOST ALERT: eateriespk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1460391595] HOST ALERT: halleground2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1460391649] HOST ALERT: eateriespk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1460391697] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1460391739] HOST ALERT: eateriespk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1460391904] SERVICE ALERT: halleground1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1460392270] SERVICE ALERT: hallegroundcolor;Cyan Cartridge HP CE261A;OK;HARD;3;Cyan Cartridge HP CE261A is at 100% - OK!
[1460392718] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1460392778] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1460392803] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1460392837] SERVICE ALERT: marshallpk;Paper Tray Status;OK;SOFT;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1460392863] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1460392923] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1460393222] SERVICE ALERT: hallefirst2;Printer Status;OK;HARD;3;Ready
[1460393283] SERVICE ALERT: ADWSUSSERVERS;Service: RPC;WARNING;SOFT;1;RpcSs: Error
[1460393403] SERVICE ALERT: ADWSUSSERVERS;Service: RPC;OK;SOFT;2;RpcSs: Started
[1460393583] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1460393703] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 75% (5 min average)
[1460395132] SERVICE FLAPPING ALERT: owen101lab1;Black Print Cartridge HP C8543X;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1460395143] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 94% (5 min average)
[1460395173] SERVICE FLAPPING ALERT: owen101lab1;Maintenance Kit HP 110V-C9152A;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1460395183] SERVICE FLAPPING ALERT: owen101lab1;Paper Tray Status;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1460395743] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1460396522] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;CRITICAL;HARD;3;Black Print Cartridge HP C8543X is at 1% - CRITICAL!
[1460397538] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460397592] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 36% - OK!
[1460398143] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1460398743] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460399661] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1460399674] SERVICE ALERT: VUM;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460399721] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1460399781] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1460399794] SERVICE ALERT: VUM;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1460399794] HOST ALERT: owen101labcolor;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460399853] SERVICE ALERT: owen101lab2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1460399858] HOST ALERT: owen101labcolor;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1460399898] SERVICE ALERT: owen101lab2;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1460399913] SERVICE ALERT: owen101lab2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1460399943] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 76% (5 min average)
[1460399958] SERVICE ALERT: owen101lab2;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1460399973] SERVICE ALERT: owen101lab2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1460399979] SERVICE ALERT: owen101lab1;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460399988] SERVICE ALERT: owen101lab1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460399998] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460400009] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1460400012] SERVICE ALERT: owen101lab2;Printer Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1460400018] SERVICE ALERT: owen101lab2;Maintenance Kit HP 110V-C9152A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1460400039] SERVICE ALERT: owen101lab1;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460400039] SERVICE ALERT: owen101lab1;Printer Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460400048] SERVICE ALERT: owen101lab1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460400059] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460400069] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1460400072] SERVICE ALERT: owen101lab2;Printer Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1460400099] SERVICE ALERT: owen101lab1;Maintenance Kit HP 110V-C9152A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460400099] SERVICE ALERT: owen101lab1;Printer Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460400108] SERVICE ALERT: owen101lab1;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460400118] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460400129] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1460400132] SERVICE ALERT: owen101lab2;Printer Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1460400159] SERVICE ALERT: owen101lab1;Printer Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460400219] SERVICE ALERT: owen101lab3;Printer Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1460400279] SERVICE ALERT: owen101lab3;Printer Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1460400339] SERVICE ALERT: owen101lab3;Printer Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1460400543] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460400663] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460400722] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1460400725] SERVICE ALERT: FITNESSGRAM;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460400783] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460400845] SERVICE ALERT: FITNESSGRAM;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1460401165] SERVICE ALERT: VIEW-CS;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1460401203] HOST ALERT: owen101lab2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460401213] SERVICE ALERT: VIEW-CS;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1460401218] HOST ALERT: VIEW-CS;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.36 ms
[1460401249] HOST ALERT: owen101lab2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460401293] SERVICE ALERT: owen101lab1;Maintenance Kit HP 110V-C9152A;OK;HARD;3;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 56% - OK!
[1460401303] SERVICE ALERT: owen101lab1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1460401312] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 36% - OK!
[1460401312] HOST ALERT: owen101lab2;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.69 ms
[1460401324] SERVICE ALERT: owen101lab3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1460401326] SERVICE ALERT: owen101lab2;Printer Status;OK;HARD;3;ORDER CARTRIDGE LESS THAN 1900 PAGES 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1460401352] HOST ALERT: VIEW-CS;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1460401353] SERVICE ALERT: owen101lab1;Printer Status;OK;HARD;3;TRAY 4 EMPTY PLAIN LETTER   Ready
[1460401442] HOST ALERT: VIEW-CS;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460401492] HOST ALERT: VIEW-CS;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460401533] SERVICE ALERT: owen101lab3;Printer Status;OK;HARD;3;Ready
[1460401552] HOST ALERT: VIEW-CS;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460401575] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 7% - OK!
[1460401585] HOST ALERT: VIEW-CS;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460401623] HOST ALERT: VIEW-CS;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1460401623] SERVICE ALERT: hallefirst2;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460401672] HOST ALERT: VIEW-CS;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1460401673] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1460401682] SERVICE ALERT: hallefirst2;Printer Status;OK;SOFT;2;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1460401733] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 9% - OK!
[1460401762] HOST ALERT: VIEW-CS;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1460401768] SERVICE ALERT: owen101lab2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1460401782] SERVICE ALERT: VIEW-CS;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460401812] HOST ALERT: VIEW-CS;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1460401812] SERVICE ALERT: owen101lab2;Maintenance Kit HP 110V-C9152A;OK;HARD;3;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 58% - OK!
[1460401832] SERVICE ALERT: VIEW-CS;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460401838] SERVICE FLAPPING ALERT: marshallpk;Paper Tray Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1460401862] HOST ALERT: VIEW-CS;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1460401902] SERVICE ALERT: VIEW-CS;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460401952] SERVICE ALERT: VIEW-CS;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460402022] SERVICE ALERT: VIEW-CS;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460402072] SERVICE ALERT: VIEW-CS;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460402175] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1460402235] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! 1 of 3 Trays Still Full
[1460402241] SERVICE ALERT: EMUFILES;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460402282] SERVICE FLAPPING ALERT: hallefirst2;Printer Status;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1460402314] SERVICE ALERT: SP2010DB;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460402360] SERVICE ALERT: EMUFILES;CPU Load;OK;SOFT;2;CPU Load 1% (5 min average)
[1460402372] SERVICE ALERT: markjeffersonpk;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460402432] SERVICE ALERT: markjeffersonpk;Printer Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460402434] SERVICE ALERT: SP2010DB;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460402492] SERVICE ALERT: markjeffersonpk;Printer Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460402554] SERVICE ALERT: SP2010DB;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1460403003] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (35.6% change >= 20.0% threshold)
[1460403003] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1460403003] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1460403003] HOST DOWNTIME ALERT: SP2010DB;STARTED; Host has entered a period of scheduled downtime
[1460403003] HOST DOWNTIME ALERT: SP2010;STARTED; Host has entered a period of scheduled downtime
[1460403092] SERVICE ALERT: markjeffersonpk;Printer Status;OK;HARD;3;Processing job from tray 2
[1460403631] SERVICE ALERT: SP2010DB;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.4.123 and port 12489: Connection refused
[1460403649] SERVICE ALERT: SP2010DB;VMWare Tools;CRITICAL;SOFT;1;connect to address 164.76.4.123 and port 12489: Connection refused
[1460403661] HOST ALERT: SP2010DB;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460403689] HOST ALERT: SP2010DB;UP;SOFT;2;PING OK - Packet loss = 75%, RTA = 0.33 ms
[1460403742] SERVICE ALERT: SP2010DB;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1460403769] SERVICE ALERT: SP2010DB;VMWare Tools;OK;SOFT;2;VMTools: Started
[1460404194] SERVICE ALERT: SP2010DB;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460404224] HOST ALERT: SP2010DB;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460404290] SERVICE ALERT: SP2010DB;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460404314] HOST ALERT: SP2010DB;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460404314] SERVICE ALERT: SP2010DB;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460404344] HOST ALERT: SP2010DB;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460404362] SERVICE ALERT: SP2010DB;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460404389] SERVICE ALERT: SP2010DB;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460404392] HOST ALERT: SP2010DB;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460404410] SERVICE ALERT: SP2010DB;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460404421] SERVICE ALERT: SP2010DB;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460404434] SERVICE ALERT: SP2010DB;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460404440] HOST ALERT: SP2010DB;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460404463] SERVICE ALERT: SP2010;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460404463] SERVICE ALERT: SP2010DB;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460404482] SERVICE ALERT: SP2010DB;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460404493] HOST ALERT: SP2010;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460404493] HOST ALERT: SP2010DB;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1460404509] SERVICE ALERT: SP2010DB;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460404510] HOST ALERT: SP2010;UP;SOFT;2;PING OK - Packet loss = 37%, RTA = 0.31 ms
[1460404523] SERVICE ALERT: SP2010DB;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460404530] SERVICE ALERT: SP2010DB;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460404539] HOST ALERT: SP2010DB;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1460404541] SERVICE ALERT: SP2010DB;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460404563] SERVICE ALERT: SP2010;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1460404571] HOST ALERT: SP2010DB;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1460404583] SERVICE ALERT: SP2010DB;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460404602] SERVICE ALERT: SP2010DB;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460404602] SERVICE ALERT: SP2010DB;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460404605] SERVICE ALERT: SP2010DB;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460404613] HOST ALERT: SP2010DB;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1460404623] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 1 Trays Empty
[1460404629] SERVICE ALERT: SP2010DB;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460404642] SERVICE ALERT: SP2010DB;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460404643] SERVICE ALERT: SP2010DB;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460404657] SERVICE ALERT: SP2010DB;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460404659] HOST ALERT: SP2010DB;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1460404661] SERVICE ALERT: SP2010DB;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460404682] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 1 Trays Empty
[1460404703] SERVICE ALERT: SP2010DB;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460404704] SERVICE ALERT: SP2010DB;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460404722] SERVICE ALERT: SP2010DB;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460404725] SERVICE ALERT: SP2010DB;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460404742] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 1 Trays Empty
[1460404745] SERVICE ALERT: SP2010DB;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460404762] SERVICE ALERT: SP2010DB;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460404763] SERVICE ALERT: SP2010DB;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460404777] SERVICE ALERT: SP2010DB;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460404824] SERVICE ALERT: SP2010DB;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460404836] SERVICE ALERT: SCCM-PRIM01;Service: NetLogin;CRITICAL;SOFT;1;connect to address 164.76.102.80 and port 12489: No route to host
[1460404839] SERVICE ALERT: SCCM-PRIM01;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.102.80 and port 12489: No route to host
[1460404839] HOST ALERT: SCCM-PRIM01;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.80)
[1460404842] SERVICE ALERT: SP2010DB;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460404845] SERVICE ALERT: SP2010DB;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460404865] SERVICE ALERT: SP2010DB;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460404866] SERVICE ALERT: SCCM-PRIM01;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.102.80 and port 12489: No route to host
[1460404869] HOST ALERT: SCCM-PRIM01;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.80)
[1460404882] SERVICE ALERT: SP2010DB;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460404883] HOST ALERT: eateriespk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.45 ms
[1460404886] SERVICE ALERT: SCCM-PRIM01;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.102.80 and port 12489: No route to host
[1460404889] SERVICE ALERT: SCCM-PRIM01;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.102.80 and port 12489: No route to host
[1460404889] HOST ALERT: SCCM-PRIM01;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.102.80)
[1460404897] SERVICE ALERT: SP2010DB;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460404907] HOST ALERT: SCCM-PRIM01;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.39 ms
[1460404944] SERVICE ALERT: SP2010DB;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460404944] SERVICE ALERT: SCCM-PRIM01;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1460404958] SERVICE ALERT: SCCM-PRIM01;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1460404963] SERVICE ALERT: SP2010DB;CPU Load;OK;HARD;3;CPU Load 63% (5 min average)
[1460404965] SERVICE ALERT: SP2010DB;Service: User Profile;OK;SOFT;3;ProfSvc: Started
[1460404973] HOST ALERT: SP2010DB;UP;HARD;10;PING WARNING - Packet loss = 83%, RTA = 0.45 ms
[1460404983] SERVICE ALERT: SCCM-PRIM01;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1460404983] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 83% (5 min average)
[1460405003] SERVICE ALERT: SCCM-PRIM01;Service: RPC;OK;SOFT;2;RpcSs: Started
[1460405009] SERVICE ALERT: SCCM-PRIM01;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1460405014] SERVICE ALERT: SP2010DB;Service: WMI;OK;HARD;3;Winmgmt: Started
[1460405060] SERVICE ALERT: eateriespk;Black Cartridge HP CC364X;OK;HARD;3;Black Cartridge HP CC364X is at 89% - OK!
[1460405092] SERVICE ALERT: eateriespk;Printer Status;OK;HARD;3;Processing...
[1460405110] SERVICE ALERT: SP2010DB;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 3 minute(s)
[1460405143] SERVICE ALERT: eateriespk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1460405182] SERVICE ALERT: SP2010DB;Service: Windows Time;OK;HARD;3;W32Time: Started
[1460405209] SERVICE ALERT: SP2010DB;VMWare Tools;OK;HARD;3;VMTools: Started
[1460405241] SERVICE ALERT: SP2010DB;Memory Usage;OK;HARD;3;Memory usage: total:14333.31 MB - used: 2297.62 MB (16%) - free: 12035.69 MB (84%)
[1460405283] SERVICE ALERT: SP2010DB;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1460405342] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1460405343] SERVICE ALERT: SP2010DB;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1460405422] SERVICE ALERT: SP2010DB;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1460405425] SERVICE ALERT: SP2010DB;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1460405462] SERVICE ALERT: SP2010DB;C:\ Drive Space;OK;HARD;3;c:\ - total: 69.90 Gb - used: 54.87 Gb (78%) - free 15.03 Gb (22%)
[1460405477] SERVICE ALERT: SP2010DB;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1460405524] SERVICE ALERT: SP2010DB;Service: RPC;OK;HARD;3;RpcSs: Started
[1460405546] HOST DOWNTIME ALERT: SP2010;CANCELLED; Scheduled downtime for host has been cancelled.
[1460405552] HOST DOWNTIME ALERT: SP2010DB;CANCELLED; Scheduled downtime for host has been cancelled.
[1460405583] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460405703] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1460405823] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1460407023] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 93% (5 min average)
[1460407623] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1460407838] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1460407898] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1460407957] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1460408373] SERVICE ALERT: prayharroldpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1460408433] SERVICE ALERT: prayharroldpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1460408493] SERVICE ALERT: prayharroldpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1460409757] SERVICE ALERT: marshallpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1460410293] SERVICE ALERT: prayharroldpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1460411558] SERVICE ALERT: owen101labcolor;Magenta Cartridge HP CE263A;CRITICAL;SOFT;1;Magenta Cartridge HP CE263A is at CRITICAL level!
[1460411618] SERVICE ALERT: owen101labcolor;Magenta Cartridge HP CE263A;CRITICAL;SOFT;2;Magenta Cartridge HP CE263A is at CRITICAL level!
[1460411678] SERVICE ALERT: owen101labcolor;Magenta Cartridge HP CE263A;CRITICAL;HARD;3;Magenta Cartridge HP CE263A is at CRITICAL level!
[1460411969] SERVICE ALERT: owen101labcolor;Cyan Cartridge HP CE261A;CRITICAL;SOFT;1;Cyan Cartridge HP CE261A is at CRITICAL level!
[1460412028] SERVICE ALERT: owen101labcolor;Cyan Cartridge HP CE261A;CRITICAL;SOFT;2;Cyan Cartridge HP CE261A is at CRITICAL level!
[1460412088] SERVICE ALERT: owen101labcolor;Cyan Cartridge HP CE261A;CRITICAL;HARD;3;Cyan Cartridge HP CE261A is at CRITICAL level!
[1460414742] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1460414801] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1460414823] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 90% (5 min average)
[1460414862] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1460415423] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1460416661] SERVICE ALERT: crossroadspk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1460420521] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1460420581] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 63% - OK!
[1460423340] SERVICE ALERT: TARGETX-2012;CPU Load;WARNING;SOFT;1;CPU Load 94% (5 min average)
[1460423460] SERVICE ALERT: TARGETX-2012;CPU Load;OK;SOFT;2;CPU Load 44% (5 min average)
[1460424735] SERVICE ALERT: DB01-2012;CPU Load;WARNING;SOFT;1;CPU Load 94% (5 min average)
[1460424855] SERVICE ALERT: DB01-2012;CPU Load;OK;SOFT;2;CPU Load 79% (5 min average)
[1460426823] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1460427479] SERVICE ALERT: hallefirst2;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460427539] SERVICE ALERT: hallefirst2;Printer Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460427599] SERVICE ALERT: hallefirst2;Printer Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460429241] SERVICE ALERT: VC;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460429341] SERVICE ALERT: VC;Service: RPC;OK;SOFT;2;RpcSs: Started
[1460430000] SERVICE ALERT: hallefirst2;Printer Status;OK;HARD;3;ORDER CARTRIDGE LESS THAN 1500 PAGES 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1460430300] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1460430420] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1460430540] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1460431740] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 92% (5 min average)
[1460432340] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460433540] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 85% (5 min average)
[1460433540] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)