[1453784889] SERVICE ALERT: halleground3;Paper Status Tray 5;OK;HARD;3;Tray 5 is at 40% - OK!
[1453812499] SERVICE ALERT: markjeffersonpk;Paper Status Tray 2;CRITICAL;SOFT;1;Tray 2 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453812559] SERVICE ALERT: markjeffersonpk;Paper Status Tray 2;CRITICAL;SOFT;2;Tray 2 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453812620] SERVICE ALERT: markjeffersonpk;Paper Status Tray 2;CRITICAL;HARD;3;Tray 2 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453812677] SERVICE ALERT: hallefirst3;Paper Status Tray 2;UNKNOWN;SOFT;1;Tray 2 status is UNKNOWN!
[1453812737] SERVICE ALERT: hallefirst3;Paper Status Tray 2;UNKNOWN;SOFT;2;Tray 2 status is UNKNOWN!
[1453812797] SERVICE ALERT: hallefirst3;Paper Status Tray 2;UNKNOWN;HARD;3;Tray 2 status is UNKNOWN!
[1453813087] SERVICE ALERT: studentcenterpk;Paper Status Tray 2;OK;HARD;3;Tray 2 is at 40% - OK!
[1453813637] SERVICE ALERT: studentcenterpk;Paper Status Tray 5;OK;HARD;3;Tray 5 is at 40% - OK!
[1453814016] SERVICE ALERT: marshallpk;Paper Status Tray 3;CRITICAL;SOFT;1;Tray 3 is at CRITICAL level - please refill with more paper.
[1453814077] SERVICE ALERT: marshallpk;Paper Status Tray 3;CRITICAL;SOFT;2;Tray 3 is at CRITICAL level - please refill with more paper.
[1453814136] SERVICE ALERT: marshallpk;Paper Status Tray 3;CRITICAL;HARD;3;Tray 3 is at CRITICAL level - please refill with more paper.
[1453814339] SERVICE ALERT: marshallpk;Paper Status Tray 2;OK;HARD;3;Tray 2 is OK!
[1453814736] SERVICE ALERT: marshallpk;Paper Status Tray 3;OK;HARD;3;Tray 3 is OK!
[1453816680] SERVICE ALERT: owen101lab2;Paper Status Tray 3;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.203.102! Make sure host is up and SNMP is configured properly.
[1453816699] SERVICE ALERT: owen101lab2;Maintenance Kit HP 110V-C9152A;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.203.102! Make sure host is up and SNMP is configured properly.
[1453816735] SERVICE ALERT: owen101lab2;Paper Status Tray 5;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.203.102! Make sure host is up and SNMP is configured properly.
[1453816740] SERVICE ALERT: owen101lab2;Paper Status Tray 3;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.203.102! Make sure host is up and SNMP is configured properly.
[1453816759] SERVICE ALERT: owen101lab2;Maintenance Kit HP 110V-C9152A;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.203.102! Make sure host is up and SNMP is configured properly.
[1453816795] SERVICE ALERT: owen101lab2;Paper Status Tray 5;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.203.102! Make sure host is up and SNMP is configured properly.
[1453816800] SERVICE ALERT: owen101lab2;Paper Status Tray 3;WARNING;HARD;3;WARNING: No SNMP response from 164.76.203.102! Make sure host is up and SNMP is configured properly.
[1453816819] SERVICE ALERT: owen101lab2;Maintenance Kit HP 110V-C9152A;WARNING;HARD;3;WARNING: No SNMP response from 164.76.203.102! Make sure host is up and SNMP is configured properly.
[1453816832] SERVICE ALERT: alexanderpk;Paper Status Tray 2;CRITICAL;SOFT;1;Tray 2 is at CRITICAL level - please refill with more 8.5 x 11.0 paper.
[1453816855] SERVICE ALERT: owen101lab2;Paper Status Tray 5;WARNING;HARD;3;WARNING: No SNMP response from 164.76.203.102! Make sure host is up and SNMP is configured properly.
[1453816893] SERVICE ALERT: alexanderpk;Paper Status Tray 2;CRITICAL;SOFT;2;Tray 2 is at CRITICAL level - please refill with more 8.5 x 11.0 paper.
[1453816902] SERVICE ALERT: owen101lab3;Paper Status Tray 3;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1453816952] SERVICE ALERT: alexanderpk;Paper Status Tray 2;CRITICAL;HARD;3;Tray 2 is at CRITICAL level - please refill with more 8.5 x 11.0 paper.
[1453816962] SERVICE ALERT: owen101lab3;Paper Status Tray 3;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1453816979] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1453817012] SERVICE ALERT: owen101lab3;Paper Status Tray 5;WARNING;HARD;3;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1453817022] SERVICE ALERT: owen101lab3;Paper Status Tray 3;WARNING;HARD;3;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1453817039] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1453817099] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;WARNING;HARD;3;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1453817099] SERVICE ALERT: owen101lab3;Paper Status Tray 2;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1453817102] SERVICE ALERT: owen101lab3;Maintenance Kit HP 110V-C9152A;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1453817127] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.203.102! Make sure host is up and SNMP is configured properly.
[1453817159] SERVICE ALERT: owen101lab3;Paper Status Tray 2;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1453817162] SERVICE ALERT: owen101lab3;Maintenance Kit HP 110V-C9152A;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1453817187] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.203.102! Make sure host is up and SNMP is configured properly.
[1453817219] SERVICE ALERT: owen101lab3;Paper Status Tray 2;WARNING;HARD;3;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1453817222] SERVICE ALERT: owen101lab3;Maintenance Kit HP 110V-C9152A;WARNING;HARD;3;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1453817238] SERVICE ALERT: owen101lab2;Paper Status Tray 2;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.203.102! Make sure host is up and SNMP is configured properly.
[1453817247] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;WARNING;HARD;3;WARNING: No SNMP response from 164.76.203.102! Make sure host is up and SNMP is configured properly.
[1453817298] SERVICE ALERT: owen101lab2;Paper Status Tray 2;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.203.102! Make sure host is up and SNMP is configured properly.
[1453817358] SERVICE ALERT: owen101lab2;Paper Status Tray 2;WARNING;HARD;3;WARNING: No SNMP response from 164.76.203.102! Make sure host is up and SNMP is configured properly.
[1453817813] SERVICE ALERT: owen101lab3;Paper Status Tray 2;OK;HARD;3;Tray 2 is at 10% - OK!
[1453817816] SERVICE ALERT: owen101lab3;Maintenance Kit HP 110V-C9152A;OK;HARD;3;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 72% - OK!
[1453817841] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 23% - OK!
[1453817952] SERVICE ALERT: owen101lab2;Paper Status Tray 2;OK;HARD;3;Tray 2 is at 40% - OK!
[1453817994] SERVICE ALERT: owen101lab2;Paper Status Tray 3;OK;HARD;3;Tray 3 is at 40% - OK!
[1453818013] SERVICE ALERT: owen101lab2;Maintenance Kit HP 110V-C9152A;OK;HARD;3;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 69% - OK!
[1453818049] SERVICE ALERT: owen101lab2;Paper Status Tray 5;OK;HARD;3;Tray 5 is at 40% - OK!
[1453818206] SERVICE ALERT: owen101lab3;Paper Status Tray 5;UNKNOWN;HARD;3;UNKNOWN - Tray 5 not found. Use the TRAY TEST option to determine which trays may be monitored.
[1453818216] SERVICE ALERT: owen101lab3;Paper Status Tray 3;OK;HARD;3;Tray 3 is at 40% - OK!
[1453818256] SERVICE ALERT: rooseveltpk;Paper Status Tray 3;CRITICAL;SOFT;1;Tray 3 is at CRITICAL level - please refill with more 8.5 x 11.0 paper.
[1453818293] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 45% - OK!
[1453818315] SERVICE ALERT: rooseveltpk;Paper Status Tray 3;CRITICAL;SOFT;2;Tray 3 is at CRITICAL level - please refill with more 8.5 x 11.0 paper.
[1453818376] SERVICE ALERT: rooseveltpk;Paper Status Tray 3;CRITICAL;HARD;3;Tray 3 is at CRITICAL level - please refill with more 8.5 x 11.0 paper.
[1453819753] SERVICE ALERT: markjeffersonpk;Paper Status Tray 5;UNKNOWN;SOFT;1;Tray 5 status is UNKNOWN!
[1453819814] SERVICE ALERT: markjeffersonpk;Paper Status Tray 5;CRITICAL;SOFT;2;Tray 5 is at CRITICAL level - please refill with more paper.
[1453819873] SERVICE ALERT: markjeffersonpk;Paper Status Tray 5;CRITICAL;HARD;3;Tray 5 is at CRITICAL level - please refill with more paper.
[1453820211] SERVICE ALERT: hallefirst2;Paper Status Tray 2;UNKNOWN;SOFT;1;Tray 2 status is UNKNOWN!
[1453820213] SERVICE ALERT: owen101lab3;Paper Status Tray 2;CRITICAL;SOFT;1;Tray 2 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453820240] SERVICE ALERT: markjeffersonpk;Paper Status Tray 3;OK;HARD;3;Tray 3 is at 40% - OK!
[1453820270] SERVICE ALERT: hallefirst2;Paper Status Tray 2;UNKNOWN;SOFT;2;Tray 2 status is UNKNOWN!
[1453820272] SERVICE ALERT: owen101lab3;Paper Status Tray 2;CRITICAL;SOFT;2;Tray 2 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453820330] SERVICE ALERT: hallefirst2;Paper Status Tray 2;UNKNOWN;HARD;3;Tray 2 status is UNKNOWN!
[1453820333] SERVICE ALERT: owen101lab3;Paper Status Tray 2;CRITICAL;HARD;3;Tray 2 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453820607] SERVICE ALERT: alexanderpk;Paper Status Tray 3;WARNING;HARD;3;WARNING: No SNMP response from 164.76.142.38! Make sure host is up and SNMP is configured properly.
[1453820637] HOST ALERT: alexanderpk;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1453820701] HOST ALERT: alexanderpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1453821074] SERVICE ALERT: markjeffersonpk;Paper Status Tray 5;UNKNOWN;HARD;3;Tray 5 status is UNKNOWN!
[1453821151] SERVICE ALERT: alexanderpk;Paper Status Tray 2;OK;HARD;3;Tray 2 is OK!
[1453821201] SERVICE ALERT: alexanderpk;Paper Status Tray 3;CRITICAL;HARD;3;Tray 3 is at CRITICAL level - please refill with more 8.5 x 11.0 paper.
[1453821308] SERVICE ALERT: hallesouthcommons;Paper Status Tray 3;CRITICAL;SOFT;1;Tray 3 is at CRITICAL level - please refill with more 8.5 x 11.0 paper.
[1453821368] SERVICE ALERT: hallesouthcommons;Paper Status Tray 3;CRITICAL;SOFT;2;Tray 3 is at CRITICAL level - please refill with more 8.5 x 11.0 paper.
[1453821429] SERVICE ALERT: hallesouthcommons;Paper Status Tray 3;CRITICAL;HARD;3;Tray 3 is at CRITICAL level - please refill with more 8.5 x 11.0 paper.
[1453821618] SERVICE ALERT: markjeffersonpk;Paper Status Tray 2;OK;HARD;3;Tray 2 is at 40% - OK!
[1453821673] SERVICE ALERT: markjeffersonpk;Paper Status Tray 5;OK;HARD;3;Tray 5 is at 40% - OK!
[1453822528] SERVICE ALERT: strongpk;Paper Status Tray 3;OK;HARD;3;Tray 3 is OK!
[1453822858] SERVICE ALERT: prayharroldpk;Paper Status Tray 5;CRITICAL;SOFT;1;Tray 5 is at 5% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453822919] SERVICE ALERT: prayharroldpk;Paper Status Tray 5;CRITICAL;SOFT;2;Tray 5 is at 5% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453822978] SERVICE ALERT: prayharroldpk;Paper Status Tray 5;CRITICAL;HARD;3;Tray 5 is at 5% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453823176] SERVICE ALERT: rooseveltpk;Paper Status Tray 3;OK;HARD;3;Tray 3 is OK!
[1453823286] SERVICE ALERT: halleground3;Paper Status Tray 5;CRITICAL;SOFT;1;Tray 5 is at CRITICAL level - please refill with more 11.0 x 8.5 paper.
[1453823347] SERVICE ALERT: halleground3;Paper Status Tray 5;CRITICAL;SOFT;2;Tray 5 is at CRITICAL level - please refill with more 11.0 x 8.5 paper.
[1453823406] SERVICE ALERT: halleground3;Paper Status Tray 5;CRITICAL;HARD;3;Tray 5 is at CRITICAL level - please refill with more 11.0 x 8.5 paper.
[1453824201] SERVICE ALERT: alexanderpk;Paper Status Tray 3;OK;HARD;3;Tray 3 is OK!
[1453824456] SERVICE ALERT: hallefirst2;Paper Status Tray 3;CRITICAL;SOFT;1;Tray 3 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453824517] SERVICE ALERT: hallefirst2;Paper Status Tray 3;CRITICAL;SOFT;2;Tray 3 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453824576] SERVICE ALERT: hallefirst2;Paper Status Tray 3;OK;SOFT;3;Tray 3 is at 20% - OK!
[1453825028] SERVICE ALERT: hallesouthcommons;Paper Status Tray 3;OK;HARD;3;Tray 3 is OK!
[1453825415] SERVICE ALERT: owen101lab3;Paper Status Tray 3;UNKNOWN;SOFT;1;Tray 3 status is UNKNOWN!
[1453825475] SERVICE ALERT: owen101lab3;Paper Status Tray 3;UNKNOWN;SOFT;2;Tray 3 status is UNKNOWN!
[1453825535] SERVICE ALERT: owen101lab3;Paper Status Tray 3;UNKNOWN;HARD;3;Tray 3 status is UNKNOWN!
[1453825776] SERVICE ALERT: hallefirst2;Paper Status Tray 3;CRITICAL;SOFT;1;Tray 3 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453825835] SERVICE ALERT: hallefirst2;Paper Status Tray 3;CRITICAL;SOFT;2;Tray 3 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453825894] SERVICE ALERT: hallefirst2;Paper Status Tray 3;CRITICAL;HARD;3;Tray 3 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453826136] SERVICE ALERT: owen101lab3;Paper Status Tray 3;OK;HARD;3;Tray 3 is at 10% - OK!
[1453826136] SERVICE FLAPPING ALERT: owen101lab3;Paper Status Tray 3;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1453826331] SERVICE ALERT: owen101lab3;Paper Status Tray 2;OK;HARD;3;Tray 2 is at 40% - OK!
[1453827539] SERVICE ALERT: marshallpk;Paper Status Tray 2;CRITICAL;SOFT;1;Tray 2 is at CRITICAL level - please refill with more 8.5 x 11.0 paper.
[1453827598] SERVICE ALERT: marshallpk;Paper Status Tray 2;CRITICAL;SOFT;2;Tray 2 is at CRITICAL level - please refill with more 8.5 x 11.0 paper.
[1453827658] SERVICE ALERT: marshallpk;Paper Status Tray 2;CRITICAL;HARD;3;Tray 2 is at CRITICAL level - please refill with more 8.5 x 11.0 paper.
[1453827796] SERVICE ALERT: hallefirst3;Paper Status Tray 2;CRITICAL;HARD;3;Tray 2 is at CRITICAL level - please refill with more 11.0 x 8.5 paper.
[1453828041] SERVICE ALERT: markjeffersonpk;Paper Status Tray 3;CRITICAL;SOFT;1;Tray 3 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453828100] SERVICE ALERT: markjeffersonpk;Paper Status Tray 3;CRITICAL;SOFT;2;Tray 3 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453828130] SERVICE ALERT: hallefirst2;Paper Status Tray 2;OK;HARD;3;Tray 2 is at 40% - OK!
[1453828161] SERVICE ALERT: markjeffersonpk;Paper Status Tray 3;CRITICAL;HARD;3;Tray 3 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453828294] SERVICE ALERT: hallefirst2;Paper Status Tray 3;OK;HARD;3;Tray 3 is at 40% - OK!
[1453829287] SERVICE ALERT: prayharroldpk;Paper Status Tray 3;CRITICAL;SOFT;1;Tray 3 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453829348] SERVICE ALERT: prayharroldpk;Paper Status Tray 3;CRITICAL;SOFT;2;Tray 3 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453829407] SERVICE ALERT: prayharroldpk;Paper Status Tray 3;CRITICAL;HARD;3;Tray 3 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453829855] SERVICE ALERT: eateriespk;Paper Status Tray 5;CRITICAL;SOFT;1;Tray 5 is at CRITICAL level - please refill with more 8.5 x 11.0 paper.
[1453829916] SERVICE ALERT: eateriespk;Paper Status Tray 5;CRITICAL;SOFT;2;Tray 5 is at CRITICAL level - please refill with more 8.5 x 11.0 paper.
[1453829975] SERVICE ALERT: eateriespk;Paper Status Tray 5;CRITICAL;HARD;3;Tray 5 is at CRITICAL level - please refill with more 8.5 x 11.0 paper.
[1453830428] SERVICE ALERT: prayharroldpk;Paper Status Tray 2;UNKNOWN;SOFT;1;Tray 2 status is UNKNOWN!
[1453830487] SERVICE ALERT: prayharroldpk;Paper Status Tray 2;CRITICAL;SOFT;2;Tray 2 is at CRITICAL level - please refill with more 11.0 x 8.5 paper.
[1453830547] SERVICE ALERT: prayharroldpk;Paper Status Tray 2;CRITICAL;HARD;3;Tray 2 is at CRITICAL level - please refill with more 11.0 x 8.5 paper.
[1453830777] SERVICE ALERT: prayharroldpk;Paper Status Tray 5;OK;HARD;3;Tray 5 is at 40% - OK!
[1453830904] SERVICE ALERT: marshallpk;Black Cartridge HP CC364X;WARNING;SOFT;1;Black Cartridge HP CC364X is at 4% - WARNING!
[1453830965] SERVICE ALERT: marshallpk;Black Cartridge HP CC364X;WARNING;SOFT;2;Black Cartridge HP CC364X is at 4% - WARNING!
[1453831024] SERVICE ALERT: marshallpk;Black Cartridge HP CC364X;WARNING;HARD;3;Black Cartridge HP CC364X is at 3% - WARNING!
[1453831130] SERVICE ALERT: owen101lab3;Paper Status Tray 2;CRITICAL;SOFT;1;Tray 2 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453831191] SERVICE ALERT: owen101lab3;Paper Status Tray 2;CRITICAL;SOFT;2;Tray 2 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453831250] SERVICE ALERT: owen101lab3;Paper Status Tray 2;CRITICAL;HARD;3;Tray 2 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453831813] SERVICE ALERT: prayharroldpk;Paper Status Tray 3;WARNING;HARD;3;WARNING: No SNMP response from 164.76.88.118! Make sure host is up and SNMP is configured properly.
[1453831843] HOST ALERT: prayharroldpk;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1453831907] HOST ALERT: prayharroldpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.65 ms
[1453832408] SERVICE ALERT: prayharroldpk;Paper Status Tray 3;CRITICAL;HARD;3;Tray 3 is at CRITICAL level - please refill with more 11.0 x 8.5 paper.
[1453832416] SERVICE ALERT: markjeffersonpk;Paper Status Tray 2;UNKNOWN;SOFT;1;Tray 2 status is UNKNOWN!
[1453832476] SERVICE ALERT: markjeffersonpk;Paper Status Tray 2;UNKNOWN;SOFT;2;Tray 2 status is UNKNOWN!
[1453832536] SERVICE ALERT: markjeffersonpk;Paper Status Tray 2;UNKNOWN;HARD;3;Tray 2 status is UNKNOWN!
[1453833166] SERVICE ALERT: owen101lab1;Paper Tray Status;UNKNOWN;SOFT;1;Printer needs paper!
[1453833226] SERVICE ALERT: owen101lab1;Paper Tray Status;UNKNOWN;SOFT;2;Printer needs paper!
[1453833285] SERVICE ALERT: owen101lab1;Paper Tray Status;UNKNOWN;HARD;3;Printer needs paper!
[1453833413] SERVICE ALERT: hallethirdsouth;Paper Tray Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.107.213! Make sure host is up and SNMP is configured properly.
[1453833473] SERVICE ALERT: hallethirdsouth;Paper Tray Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.107.213! Make sure host is up and SNMP is configured properly.
[1453833533] SERVICE ALERT: hallethirdsouth;Paper Tray Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.107.213! Make sure host is up and SNMP is configured properly.
[1453833557] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1453833650] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at 4% - WARNING!
[1453833711] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at 4% - WARNING!
[1453833770] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at 4% - WARNING!
[1453834158] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1453834217] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1453834277] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1453835784] SERVICE ALERT: prayharroldpk;Maintenance Kit HP 110V-C9152A;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.88.118! Make sure host is up and SNMP is configured properly.
[1453835789] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.88.118! Make sure host is up and SNMP is configured properly.
[1453835814] HOST ALERT: prayharroldpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1453835844] SERVICE ALERT: prayharroldpk;Maintenance Kit HP 110V-C9152A;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.88.118! Make sure host is up and SNMP is configured properly.
[1453835849] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.88.118! Make sure host is up and SNMP is configured properly.
[1453835854] HOST ALERT: prayharroldpk;UP;SOFT;2;PING OK - Packet loss = 50%, RTA = 0.56 ms
[1453835899] SERVICE ALERT: prayharroldpk;Maintenance Kit HP 110V-C9152A;OK;SOFT;3;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 81% - OK!
[1453835904] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;OK;SOFT;3;Black Print Cartridge HP C8543X is at 44% - OK!
[1453836286] SERVICE ALERT: owen101lab1;Paper Tray Status;OK;HARD;3;Printer Paper is OK! -  Total Trays: 2, Trays Needing Paper: 0
[1453838223] SERVICE ALERT: marshallpk;Black Cartridge HP CC364X;CRITICAL;HARD;3;Black Cartridge HP CC364X is at 1% - CRITICAL!
[1453840504] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;WARNING;SOFT;1;Cyan Cartridge HP CE261A is at 3% - WARNING!
[1453840563] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;WARNING;SOFT;2;Cyan Cartridge HP CE261A is at 3% - WARNING!
[1453840622] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;WARNING;HARD;3;Cyan Cartridge HP CE261A is at 3% - WARNING!
[1453842632] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at 5% - WARNING!
[1453842691] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at 5% - WARNING!
[1453842751] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at 5% - WARNING!
[1453843276] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1453848082] SERVICE ALERT: hallegroundcolor;Yellow Cartridge HP CE262A;WARNING;SOFT;1;Yellow Cartridge HP CE262A is at 5% - WARNING!
[1453848143] SERVICE ALERT: hallegroundcolor;Yellow Cartridge HP CE262A;WARNING;SOFT;2;Yellow Cartridge HP CE262A is at 5% - WARNING!
[1453848202] SERVICE ALERT: hallegroundcolor;Yellow Cartridge HP CE262A;WARNING;HARD;3;Yellow Cartridge HP CE262A is at 5% - WARNING!
[1453849622] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;CRITICAL;HARD;3;Cyan Cartridge HP CE261A is at 1% - CRITICAL!
[1453852035] SERVICE ALERT: strongpk;Paper Tray Status;UNKNOWN;SOFT;1;Printer needs paper!
[1453852096] SERVICE ALERT: strongpk;Paper Tray Status;UNKNOWN;SOFT;2;Printer needs paper!
[1453852156] SERVICE ALERT: strongpk;Paper Tray Status;UNKNOWN;HARD;3;Printer needs paper!
[1453854555] SERVICE ALERT: strongpk;Paper Tray Status;OK;HARD;3;Printer Paper is OK! -  Total Trays: 2, Trays Needing Paper: 1
[1453855156] SERVICE ALERT: strongpk;Paper Tray Status;UNKNOWN;SOFT;1;Printer needs paper!
[1453855215] SERVICE ALERT: strongpk;Paper Tray Status;UNKNOWN;SOFT;2;Printer needs paper!
[1453855276] SERVICE ALERT: strongpk;Paper Tray Status;UNKNOWN;HARD;3;Printer needs paper!
[1453856819] SERVICE ALERT: owen101lab3;Maintenance Kit HP 110V-C9152A;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1453856879] SERVICE ALERT: owen101lab3;Maintenance Kit HP 110V-C9152A;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1453856939] SERVICE ALERT: owen101lab3;Maintenance Kit HP 110V-C9152A;WARNING;HARD;3;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1453857003] HOST ALERT: owen101lab2;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1453857016] SERVICE ALERT: owen101lab2;Maintenance Kit HP 110V-C9152A;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.203.102! Make sure host is up and SNMP is configured properly.
[1453857035] HOST ALERT: owen101lab2;UP;SOFT;2;PING OK - Packet loss = 75%, RTA = 0.71 ms
[1453857056] SERVICE ALERT: owen101lab1;Maintenance Kit HP 110V-C9152A;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.203.101! Make sure host is up and SNMP is configured properly.
[1453857070] SERVICE ALERT: owen101lab2;Maintenance Kit HP 110V-C9152A;OK;SOFT;2;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 69% - OK!
[1453857116] SERVICE ALERT: owen101lab1;Maintenance Kit HP 110V-C9152A;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.203.101! Make sure host is up and SNMP is configured properly.
[1453857176] SERVICE ALERT: owen101lab1;Maintenance Kit HP 110V-C9152A;WARNING;HARD;3;WARNING: No SNMP response from 164.76.203.101! Make sure host is up and SNMP is configured properly.
[1453857291] SERVICE ALERT: owen101lab1;Paper Tray Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.203.101! Make sure host is up and SNMP is configured properly.
[1453857292] SERVICE ALERT: owen101lab3;Paper Tray Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1453857297] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1453857351] SERVICE ALERT: owen101lab3;Paper Tray Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1453857351] SERVICE ALERT: owen101lab1;Paper Tray Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.203.101! Make sure host is up and SNMP is configured properly.
[1453857357] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1453857410] SERVICE ALERT: owen101lab3;Paper Tray Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1453857412] SERVICE ALERT: owen101lab1;Paper Tray Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.203.101! Make sure host is up and SNMP is configured properly.
[1453857417] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;WARNING;HARD;3;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.