[1459141342] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 80% (5 min average)
[1459141942] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1459142062] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1459142182] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1459145017] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1459145137] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 95% (5 min average)
[1459145182] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 87% (5 min average)
[1459145257] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 86% (5 min average)
[1459145617] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 84.90 Gb - used: 79.04 Gb (93%) - free 5.86 Gb (7%)
[1459145737] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 84.90 Gb - used: 79.04 Gb (93%) - free 5.86 Gb (7%)
[1459145782] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459145857] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 84.90 Gb - used: 79.04 Gb (93%) - free 5.86 Gb (7%)
[1459145857] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1459145902] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1459145977] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 97% (5 min average)
[1459146022] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 83% (5 min average)
[1459146097] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 82% (5 min average)
[1459146622] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459146697] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1459146742] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1459146817] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1459146862] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1459146937] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 97% (5 min average)
[1459147537] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 92% (5 min average)
[1459148137] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 69% (5 min average)
[1459149262] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 72% (5 min average)
[1459149337] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1459149457] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 93% (5 min average)
[1459149577] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 77% (5 min average)
[1459149862] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1459149982] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 97% (5 min average)
[1459150102] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 73% (5 min average)
[1459150702] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1459150777] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1459150822] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1459150897] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 96% (5 min average)
[1459150942] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 97% (5 min average)
[1459151017] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 89% (5 min average)
[1459152742] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 62% (5 min average)
[1459155142] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1459155262] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 91% (5 min average)
[1459155382] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1459155982] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 46% (5 min average)
[1459158382] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1459158502] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 89% (5 min average)
[1459159702] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459159822] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1459159942] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1459160542] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 92% (5 min average)
[1459161142] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 79% (5 min average)
[1459161742] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1459161862] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1459161982] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1459163558] SERVICE ALERT: ADRODC-LIV;AD Replication;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 30 seconds.
[1459163654] SERVICE ALERT: ADRODC-LIV;AD Replication;OK;SOFT;2;OK: replication is up and running
[1459164641] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459164701] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459164761] SERVICE ALERT: halleground2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459165352] SERVICE ALERT: SP2010;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1459165472] SERVICE ALERT: SP2010;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1459166402] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (28.7% change >= 20.0% threshold)
[1459166402] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1459166402] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1459166402] SERVICE DOWNTIME ALERT: BOEAPPTEST;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1459167382] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 85% (5 min average)
[1459167982] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459168102] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1459168222] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 89% (5 min average)
[1459168655] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 84.90 Gb - used: 80.71 Gb (95%) - free 4.19 Gb (5%)
[1459168822] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1459168942] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1459169062] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 96% (5 min average)
[1459169662] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 80% (5 min average)
[1459170171] HOST DOWNTIME ALERT: ADWSUSSERVERS;STARTED; Host has entered a period of scheduled downtime
[1459170304] SERVICE ALERT: ADWSUSSERVERS;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459170334] HOST ALERT: ADWSUSSERVERS;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459170350] SERVICE ALERT: ADWSUSSERVERS;E:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459170355] SERVICE ALERT: ADWSUSSERVERS;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459170380] HOST ALERT: ADWSUSSERVERS;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1459170424] SERVICE ALERT: ADWSUSSERVERS;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1459170454] HOST ALERT: ADWSUSSERVERS;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1459170470] SERVICE ALERT: ADWSUSSERVERS;E:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1459170475] SERVICE ALERT: ADWSUSSERVERS;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1459170482] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 12072;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 30 seconds.
[1459170484] SERVICE ALERT: ADWSUSSERVERS;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459170484] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10022;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 30 seconds.
[1459170500] HOST ALERT: ADWSUSSERVERS;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1459170532] SERVICE ALERT: ADWSUSSERVERS;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459170544] SERVICE ALERT: ADWSUSSERVERS;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1459170562] HOST ALERT: ADWSUSSERVERS;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1459170569] SERVICE ALERT: ADWSUSSERVERS;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459170577] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10042;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 30 seconds.
[1459170590] SERVICE ALERT: ADWSUSSERVERS;E:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1459170595] SERVICE ALERT: ADWSUSSERVERS;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1459170599] HOST ALERT: ADWSUSSERVERS;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1459170602] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 12072;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 30 seconds.
[1459170604] SERVICE ALERT: ADWSUSSERVERS;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1459170604] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10022;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 30 seconds.
[1459170612] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 364;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 30 seconds.
[1459170632] HOST ALERT: ADWSUSSERVERS;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1459170652] SERVICE ALERT: ADWSUSSERVERS;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1459170665] SERVICE ALERT: ADWSUSSERVERS;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459170665] SERVICE ALERT: ADWSUSSERVERS;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459170682] HOST ALERT: ADWSUSSERVERS;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1459170689] SERVICE ALERT: ADWSUSSERVERS;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1459170697] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10042;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 30 seconds.
[1459170712] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10032;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 30 seconds.
[1459170717] SERVICE ALERT: ADWSUSSERVERS;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459170717] SERVICE ALERT: ADWSUSSERVERS;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459170719] HOST ALERT: ADWSUSSERVERS;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1459170722] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 12072;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 30 seconds.
[1459170724] SERVICE ALERT: ADWSUSSERVERS;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1459170724] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10022;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 30 seconds.
[1459170731] SERVICE ALERT: ADWSUSSERVERS;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459170732] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 364;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 30 seconds.
[1459170752] HOST ALERT: ADWSUSSERVERS;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1459170772] SERVICE ALERT: ADWSUSSERVERS;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1459170785] SERVICE ALERT: ADWSUSSERVERS;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1459170785] SERVICE ALERT: ADWSUSSERVERS;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1459170809] SERVICE ALERT: ADWSUSSERVERS;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1459170817] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10042;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 30 seconds.
[1459170832] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10032;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 30 seconds.
[1459170837] SERVICE ALERT: ADWSUSSERVERS;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1459170837] SERVICE ALERT: ADWSUSSERVERS;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1459170851] SERVICE ALERT: ADWSUSSERVERS;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1459170852] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 364;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 30 seconds.
[1459170862] SERVICE ALERT: ADWSUSSERVERS;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459170878] SERVICE ALERT: ADWSUSSERVERS;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459170887] SERVICE ALERT: ADWSUSSERVERS;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459170887] SERVICE ALERT: ADWSUSSERVERS;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459170905] SERVICE ALERT: ADWSUSSERVERS;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1459170905] SERVICE ALERT: ADWSUSSERVERS;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1459170952] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10032;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 30 seconds.
[1459170957] SERVICE ALERT: ADWSUSSERVERS;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1459170957] SERVICE ALERT: ADWSUSSERVERS;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1459170971] SERVICE ALERT: ADWSUSSERVERS;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1459170982] SERVICE ALERT: ADWSUSSERVERS;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1459170998] SERVICE ALERT: ADWSUSSERVERS;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1459171007] SERVICE ALERT: ADWSUSSERVERS;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1459171007] SERVICE ALERT: ADWSUSSERVERS;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1459171102] SERVICE ALERT: ADWSUSSERVERS;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1459171118] SERVICE ALERT: ADWSUSSERVERS;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1459171127] SERVICE ALERT: ADWSUSSERVERS;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1459171127] SERVICE ALERT: ADWSUSSERVERS;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1459171462] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1459171582] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1459171702] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1459171992] SERVICE ALERT: C4SAS-SE;CPU Load;WARNING;SOFT;1;CPU Load 94% (5 min average)
[1459172112] SERVICE ALERT: C4SAS-SE;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1459172619] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (38.6% change >= 20.0% threshold)
[1459172619] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1459172619] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1459172619] HOST DOWNTIME ALERT: ADWSUSSERVERS;STARTED; Host has entered a period of scheduled downtime
[1459172619] SERVICE DOWNTIME ALERT: BOEAPPTEST;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1459172855] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 84.90 Gb - used: 80.34 Gb (95%) - free 4.56 Gb (5%)
[1459173070] SERVICE ALERT: prayharroldpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1459173075] HOST ALERT: prayharroldpk;UP;HARD;10;PING OK - Packet loss = 16%, RTA = 0.58 ms
[1459173142] SERVICE ALERT: prayharroldpk;Device Status;OK;HARD;3;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1459173398] HOST ALERT: prayharroldpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459173488] HOST ALERT: prayharroldpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1459173578] HOST ALERT: prayharroldpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1459173668] HOST ALERT: prayharroldpk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1459173675] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459173705] HOST ALERT: prayharroldpk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1459173710] SERVICE ALERT: ADCA1;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459173735] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459173748] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459173754] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459173765] HOST ALERT: prayharroldpk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1459173795] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459173808] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459173825] HOST ALERT: prayharroldpk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1459173830] SERVICE ALERT: ADCA1;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1459173867] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459173874] SERVICE ALERT: BOEAPPTEST;CPU Load;OK;SOFT;2;CPU Load 23% (5 min average)
[1459173897] HOST ALERT: prayharroldpk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1459173986] HOST ALERT: prayharroldpk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1459174076] HOST ALERT: prayharroldpk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1459174702] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 84% (5 min average)
[1459174893] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 12072;OK;HARD;3;Eventlog check ok
[1459174894] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10022;OK;HARD;3;Eventlog check ok
[1459174897] HOST ALERT: ADWSUSSERVERS;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.36 ms
[1459174904] SERVICE ALERT: ADWSUSSERVERS;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1459174952] SERVICE ALERT: ADWSUSSERVERS;Service: RPC;OK;HARD;3;RpcSs: Started
[1459174987] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10042;OK;HARD;3;Eventlog check ok
[1459174989] SERVICE ALERT: ADWSUSSERVERS;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1459175022] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 364;OK;HARD;3;Eventlog check ok
[1459175085] SERVICE ALERT: ADWSUSSERVERS;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1459175085] SERVICE ALERT: ADWSUSSERVERS;Service: Windows Time;OK;HARD;3;W32Time: Started
[1459175122] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10032;OK;HARD;3;Eventlog check ok
[1459175137] SERVICE ALERT: ADWSUSSERVERS;C:\ Drive Space;OK;HARD;3;c:\ - total: 79.66 Gb - used: 48.44 Gb (61%) - free 31.22 Gb (39%)
[1459175137] SERVICE ALERT: ADWSUSSERVERS;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 4 minute(s)
[1459175151] SERVICE ALERT: ADWSUSSERVERS;Memory Usage;OK;HARD;3;Memory usage: total:32767.77 MB - used: 5073.70 MB (15%) - free: 27694.07 MB (85%)
[1459175282] SERVICE ALERT: ADWSUSSERVERS;Service: WMI;OK;HARD;3;Winmgmt: Started
[1459175298] SERVICE ALERT: ADWSUSSERVERS;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1459175307] SERVICE ALERT: ADWSUSSERVERS;VMWare Tools;OK;HARD;3;VMTools: Started
[1459175307] SERVICE ALERT: ADWSUSSERVERS;CPU Load;OK;HARD;3;CPU Load 35% (5 min average)
[1459175324] SERVICE ALERT: ADWSUSSERVERS;NSClient++ Version;OK;HARD;3;NSClient++ 0,4,1,73 2012-12-17
[1459175370] SERVICE ALERT: ADWSUSSERVERS;E:\ Drive Space;OK;HARD;3;e:\ - total: 99.87 Gb - used: 62.28 Gb (62%) - free 37.59 Gb (38%)
[1459175375] SERVICE ALERT: ADWSUSSERVERS;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1459175902] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1459175907] SERVICE ALERT: PROCARE;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459176022] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1459176027] SERVICE ALERT: PROCARE;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1459176142] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1459176455] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;OK;HARD;3;c:\ - total: 84.90 Gb - used: 76.03 Gb (90%) - free 8.87 Gb (10%)
[1459176462] HOST ALERT: hallefirst2;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.49 ms
[1459176475] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Warming up
[1459176753] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1459177342] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 5% - OK!
[1459177562] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;WARNING;SOFT;1;Cyan Cartridge HP CE261A is at 2% - WARNING!
[1459177622] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;WARNING;SOFT;2;Cyan Cartridge HP CE261A is at 2% - WARNING!
[1459177682] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459177683] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;WARNING;HARD;3;Cyan Cartridge HP CE261A is at 2% - WARNING!
[1459177684] HOST ALERT: hallefirst2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459177741] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459177771] HOST ALERT: hallefirst2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1459177800] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459177830] HOST ALERT: hallefirst2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1459177920] HOST ALERT: hallefirst2;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1459177949] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459177958] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459177979] HOST ALERT: hallefirst2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1459178008] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459178018] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459178038] HOST ALERT: hallefirst2;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1459178067] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459178078] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459178097] HOST ALERT: hallefirst2;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1459178187] HOST ALERT: hallefirst2;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1459178277] HOST ALERT: hallefirst2;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1459178367] HOST ALERT: hallefirst2;DOWN;HARD;10;(Host check timed out after 30.00 seconds)
[1459180341] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 90% (5 min average)
[1459180941] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 96% (5 min average)
[1459183083] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;CRITICAL;HARD;3;Cyan Cartridge HP CE261A is at 1% - CRITICAL!
[1459184015] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1459184135] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 40% (5 min average)
[1459185168] SERVICE ALERT: ADCA2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459185288] SERVICE ALERT: ADCA2;CPU Load;OK;SOFT;2;CPU Load 3% (5 min average)
[1459187541] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (5 min average)
[1459187844] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1459187903] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1459187964] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1459188141] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459188261] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1459188381] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1459189163] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1459191627] SERVICE ALERT: PROCARE;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1459191747] SERVICE ALERT: PROCARE;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1459193181] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 92% (5 min average)
[1459193502] HOST ALERT: prayharroldpk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.55 ms
[1459193589] SERVICE ALERT: prayharroldpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1459193660] SERVICE ALERT: prayharroldpk;Device Status;OK;HARD;3;Clearing paper path
[1459193762] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 30% - OK!
[1459193781] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1459193883] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;OK;HARD;3;Cyan Cartridge HP CE261A is at 100% - OK!
[1459194136] HOST ALERT: prayharroldpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459194195] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459194225] HOST ALERT: prayharroldpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1459194255] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459194265] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459194285] HOST ALERT: prayharroldpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1459194315] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459194325] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459194345] HOST ALERT: prayharroldpk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1459194368] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459194385] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459194398] HOST ALERT: prayharroldpk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1459194428] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459194458] HOST ALERT: prayharroldpk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1459194488] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459194518] HOST ALERT: prayharroldpk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1459194608] HOST ALERT: prayharroldpk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1459194698] HOST ALERT: prayharroldpk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1459194788] HOST ALERT: prayharroldpk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1459195202] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1459195203] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1459195203] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1459195203] HOST DOWNTIME ALERT: ADWSUSSERVERS;STARTED; Host has entered a period of scheduled downtime
[1459195203] SERVICE DOWNTIME ALERT: BOEAPPTEST;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1459196181] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 88% (5 min average)
[1459196781] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1459196901] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 87% (5 min average)
[1459197222] HOST ALERT: prayharroldpk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.61 ms
[1459197310] SERVICE ALERT: prayharroldpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1459197379] SERVICE ALERT: prayharroldpk;Device Status;OK;HARD;3;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1459197482] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 30% - OK!
[1459197501] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1459197621] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1459197741] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1459199909] SERVICE ALERT: NETIDADPROXY;Service: DNS Client;CRITICAL;SOFT;1;Dnscache: Stopped
[1459200029] SERVICE ALERT: NETIDADPROXY;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1459208429] SERVICE ALERT: NETIDADPROXY;Service: DNS Client;CRITICAL;SOFT;1;Dnscache: Stopped
[1459208541] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1459208549] SERVICE ALERT: NETIDADPROXY;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1459212689] SERVICE ALERT: DB01-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1459212809] SERVICE ALERT: DB01-2012;CPU Load;OK;SOFT;2;CPU Load 9% (5 min average)
[1459213381] SERVICE ALERT: IT-2012;CPU Load;WARNING;SOFT;1;CPU Load 91% (5 min average)
[1459213501] SERVICE ALERT: IT-2012;CPU Load;OK;SOFT;2;CPU Load 53% (5 min average)
[1459213510] SERVICE ALERT: C4SAS-SE;CPU Load;WARNING;SOFT;1;CPU Load 91% (5 min average)
[1459213630] SERVICE ALERT: C4SAS-SE;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1459214101] SERVICE ALERT: IT-2012;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1459214221] SERVICE ALERT: IT-2012;CPU Load;WARNING;SOFT;2;CPU Load 93% (5 min average)
[1459214341] SERVICE ALERT: IT-2012;CPU Load;WARNING;HARD;3;CPU Load 92% (5 min average)
[1459214349] SERVICE ALERT: SP2010DB;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459214468] SERVICE ALERT: SP2010DB;CPU Load;OK;SOFT;2;CPU Load 2% (5 min average)
[1459214941] SERVICE ALERT: IT-2012;CPU Load;OK;HARD;3;CPU Load 16% (5 min average)
[1459215209] SERVICE ALERT: DB01-2012;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1459215329] SERVICE ALERT: DB01-2012;CPU Load;OK;SOFT;2;CPU Load 23% (5 min average)
[1459215741] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 82% (5 min average)
[1459216941] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1459217061] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1459217181] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1459218040] SERVICE ALERT: hallegroundcolor;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 1 Trays Empty
[1459218099] SERVICE ALERT: hallegroundcolor;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 1 Trays Empty
[1459218159] SERVICE ALERT: hallegroundcolor;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 1 Trays Empty
[1459219581] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 71% (5 min average)
[1459219730] SERVICE ALERT: LANSWEEPER;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459219830] SERVICE ALERT: LANSWEEPER;VMWare Tools;OK;SOFT;2;VMTools: Started
[1459220181] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459220301] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1459220421] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1459220421] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1459221148] SERVICE ALERT: NETIDADPROXY;Service: DNS Client;CRITICAL;SOFT;1;Dnscache: Stopped
[1459221268] SERVICE ALERT: NETIDADPROXY;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1459221333] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1459221453] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 92% (5 min average)
[1459221573] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 92% (5 min average)
[1459221861] HOST ALERT: hallethirdsouth;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1459222173] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 80% (5 min average)
[1459222821] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 81% (5 min average)
[1459223421] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459223541] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1459223661] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 96% (5 min average)