[1484978062] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;1;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1484978182] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;2;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1484978302] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;HARD;3;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1484979408] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1484979528] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1484979648] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1484980200] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1484980248] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1484980287] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1484980405] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;2;Warning:  idm1  are more than 120 seconds behind!
[1484980529] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1484981448] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 91% (15 min average)
[1484981725] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1484982049] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1484982168] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1484982288] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1484983488] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 55% (15 min average)
[1484983525] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1484984088] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1484984125] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1484984208] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1484984245] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1484984328] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1484984366] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1484984366] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1484984902] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;OK;HARD;3;OK: No entries found
[1484984965] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1484984987] SERVICE ALERT: IT-2012;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1484985107] SERVICE ALERT: IT-2012;CPU Load;OK;SOFT;2;CPU Load 10% (15 min average)
[1484987328] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1484989165] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1484989765] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1484989885] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1484990005] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1484990328] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 86% (15 min average)
[1484990605] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1484991003] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (37.6% change >= 20.0% threshold)
[1484991003] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (41.7% change >= 20.0% threshold)
[1484991003] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1484991003] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1484991003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1484991003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1484991003] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1484991003] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1484991003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1484991003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1484991003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1484991003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1484991003] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1484991003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1484991003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1484991003] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1484991003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1484991003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1484991003] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1484991205] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1484992802] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (27.4% change >= 20.0% threshold)
[1484992802] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (44.5% change >= 20.0% threshold)
[1484992802] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1484992802] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1484992802] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1484992802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1484992802] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1484992802] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1484992802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1484992802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1484992802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1484992802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1484992802] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1484992802] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1484992802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1484992802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1484992802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1484992802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1484992802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1484996928] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1484997048] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1484997168] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1484997805] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1484997925] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1484998200] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service has exited from a period of scheduled downtime
[1484998358] SERVICE ALERT: INTLAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1484998478] SERVICE ALERT: INTLAPP;CPU Load;OK;SOFT;2;CPU Load 0% (15 min average)
[1484999568] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1485000168] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1485002568] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1485002726] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1485003768] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1485004968] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1485007968] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1485008088] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 26% (15 min average)
[1485009003] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (28.0% change >= 20.0% threshold)
[1485009003] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1485009003] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1485009003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1485009003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1485009003] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1485009003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1485009003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1485009003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1485009003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1485009003] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1485009003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1485009003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1485009003] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1485009003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1485009003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1485009003] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1485009775] SERVICE ALERT: CAVA2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1485009895] SERVICE ALERT: CAVA2;CPU Load;OK;SOFT;2;CPU Load 1% (15 min average)
[1485013589] SERVICE ALERT: eateriespk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1485013805] SERVICE ALERT: eateriespk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1485013865] SERVICE ALERT: eateriespk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1485013925] SERVICE ALERT: eateriespk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1485015725] SERVICE ALERT: eateriespk;Device Status;OK;HARD;3;TRAY 2 OPEN  46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1485015887] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1485019803] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1485019803] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1485019803] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1485019803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1485019803] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1485019803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1485019803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1485019803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1485019803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1485019803] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1485019803] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1485019803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1485019803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1485019803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1485019803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1485019803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1485020687] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1485020807] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1485020927] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1485022127] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 83% (15 min average)
[1485022727] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1485022847] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1485022967] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1485024098] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1485024167] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 35% (15 min average)
[1485024167] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1485025967] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1485026087] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 80% (15 min average)
[1485027287] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1485027407] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1485027527] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 76% (15 min average)
[1485032787] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1485032847] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1485032907] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1485032927] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1485033047] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 72% (15 min average)
[1485035447] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1485036647] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1485036767] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1485036887] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 75% (15 min average)
[1485037487] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1485037607] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1485037727] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 91% (15 min average)
[1485041927] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1485042047] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1485042167] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 75% (15 min average)
[1485044566] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1485044686] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 81% (15 min average)
[1485047086] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1485047206] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1485047326] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1485047926] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 43% (15 min average)
[1485048526] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1485048646] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 79% (15 min average)
[1485050446] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1485050566] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1485051166] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1485051286] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1485051406] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1485053206] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 78% (15 min average)
[1485053206] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1485053806] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1485053926] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 92% (15 min average)
[1485054526] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1485054646] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1485054766] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 79% (15 min average)
[1485055966] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1485056086] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1485057649] SERVICE ALERT: GOVGIS;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1485057769] SERVICE ALERT: GOVGIS;CPU Load;OK;SOFT;2;CPU Load 4% (15 min average)
[1485058486] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1485058606] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1485058726] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1485059926] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (15 min average)
[1485061126] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)