[1490932963] SERVICE ALERT: cava2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1490933023] SERVICE ALERT: cava2;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1490933083] SERVICE ALERT: cava2;CPU Load;OK;SOFT;3;CPU Load 0% (15 min average)
[1490933295] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 84% (15 min average)
[1490933895] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1490933955] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1490934015] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1490934601] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1490934615] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 66% (15 min average)
[1490935215] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1490935276] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1490935335] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1490935483] SERVICE ALERT: cava2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1490935543] SERVICE ALERT: cava2;CPU Load;OK;SOFT;2;CPU Load 5% (15 min average)
[1490937135] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 83% (15 min average)
[1490937735] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1490937795] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1490937855] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 88% (15 min average)
[1490938049] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1490938108] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;2;Warning:  idm1  are more than 120 seconds behind!
[1490938168] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1490938768] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1490939368] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1490939655] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1490939715] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1490939775] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1490940375] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1490940568] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1490940568] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1490941168] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1490941575] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1490942003] SERVICE ALERT: XDM;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1490942063] SERVICE ALERT: XDM;CPU Load;OK;SOFT;2;CPU Load 0% (15 min average)
[1490942175] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1490942368] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1490942775] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1490942968] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1490943568] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1490945968] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1490946029] HOST ALERT: halleground2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490946096] HOST ALERT: halleground2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1490946105] HOST ALERT: halleground1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490946144] HOST ALERT: halleground2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1490946195] HOST ALERT: halleground1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1490946234] HOST ALERT: halleground2;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1490946285] HOST ALERT: halleground1;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1490946324] HOST ALERT: halleground2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1490946375] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1490946375] HOST ALERT: halleground1;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1490946414] HOST ALERT: halleground2;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1490946465] HOST ALERT: halleground1;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1490946504] HOST ALERT: halleground2;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1490946526] HOST ALERT: halleground1;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1490946568] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1490946594] HOST ALERT: halleground2;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1490946616] HOST ALERT: halleground1;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1490946628] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1490946629] HOST ALERT: halleground2;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1490946688] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1490946696] HOST ALERT: halleground2;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1490946705] HOST ALERT: halleground1;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1490946741] HOST ALERT: halleground2;UP;HARD;10;PING WARNING - Packet loss = 82%, RTA = 0.82 ms
[1490946795] HOST ALERT: halleground1;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1490946885] HOST ALERT: halleground1;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1490946975] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1490947071] HOST ALERT: halleground2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490947161] HOST ALERT: halleground2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1490947229] HOST ALERT: halleground2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1490947296] HOST ALERT: halleground2;UP;SOFT;4;PING WARNING - Packet loss = 86%, RTA = 0.49 ms
[1490947344] HOST ALERT: halleground2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490947344] HOST FLAPPING ALERT: halleground2;STARTED; Host appears to have started flapping (22.8% change > 20.0% threshold)
[1490947434] HOST ALERT: halleground2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1490947524] HOST ALERT: halleground2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1490947614] HOST ALERT: halleground2;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1490947704] HOST ALERT: halleground2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1490947794] HOST ALERT: halleground2;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1490947830] HOST ALERT: halleground2;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1490947888] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1490947896] HOST ALERT: halleground2;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1490947900] SERVICE ALERT: ADDC5;AD Group Changes;UNKNOWN;SOFT;1;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1490947944] HOST ALERT: halleground2;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1490948034] HOST ALERT: halleground2;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1490948096] SERVICE ALERT: ADDC5;AD Group Changes;OK;SOFT;2;WARNING:  0  groups have changed!
[1490948175] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 67% (15 min average)
[1490948406] HOST ALERT: halleground2;UP;HARD;10;PING OK - Packet loss = 37%, RTA = 127.12 ms
[1490948488] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1490948496] HOST ALERT: halleground2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1490948544] HOST ALERT: halleground2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1490948634] HOST ALERT: halleground2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1490948724] HOST ALERT: halleground2;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1490948814] HOST ALERT: halleground2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1490948904] HOST ALERT: halleground2;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1490948994] HOST ALERT: halleground2;UP;SOFT;7;PING WARNING - Packet loss = 93%, RTA = 1508.68 ms
[1490949000] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service has exited from a period of scheduled downtime
[1490949002] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (44.5% change >= 20.0% threshold)
[1490949002] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (60.8% change >= 20.0% threshold)
[1490949002] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1490949002] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1490949002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1490949002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1490949002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1490949002] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1490949002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1490949002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1490949002] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1490949002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1490949002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1490949002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1490949002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1490949002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1490949002] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1490949043] HOST ALERT: halleground2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1490949043] HOST FLAPPING ALERT: halleground2;STARTED; Host appears to have started flapping (22.1% change > 20.0% threshold)
[1490949096] HOST ALERT: halleground2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1490949144] HOST ALERT: halleground2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1490949234] HOST ALERT: halleground2;DOWN;SOFT;4;(Host check timed out after 30.00 seconds)
[1490949324] HOST ALERT: halleground2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1490949414] HOST ALERT: halleground2;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1490949504] HOST ALERT: halleground2;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1490949594] HOST ALERT: halleground2;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1490949629] HOST ALERT: halleground2;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1490949696] HOST ALERT: halleground2;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1490949975] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1490950035] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1490950095] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1490950229] HOST FLAPPING ALERT: halleground2;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1490950695] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 26% (15 min average)
[1490950802] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (50.0% change >= 20.0% threshold)
[1490950802] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (52.4% change >= 20.0% threshold)
[1490950802] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1490950802] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1490950802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1490950802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1490950802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1490950802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1490950802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1490950802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1490950802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1490950802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1490950802] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1490950802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1490950802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1490950802] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1490950802] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1490951682] SERVICE ALERT: OFFICEKMS;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1490951742] SERVICE ALERT: OFFICEKMS;CPU Load;OK;SOFT;2;CPU Load 12% (15 min average)
[1490952342] SERVICE ALERT: OFFICEKMS;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1490952402] SERVICE ALERT: OFFICEKMS;CPU Load;OK;SOFT;2;CPU Load 82% (15 min average)
[1490952687] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1490952747] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1490954895] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1490954955] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1490955015] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1490957415] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1490958015] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1490959947] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1490961015] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 77% (15 min average)
[1490961615] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1490961675] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1490961735] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 87% (15 min average)
[1490961899] SERVICE ALERT: ADRODC-DET;AD Replication;WARNING;SOFT;1;WARNING: Replication with ADDC5 failed!
[1490961929] HOST ALERT: halleground3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490961960] SERVICE ALERT: ADRODC-DET;AD Replication;WARNING;SOFT;2;WARNING: Replication with ADDC5 failed!
[1490961964] HOST ALERT: halleground3;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1490962020] SERVICE ALERT: ADRODC-DET;AD Replication;WARNING;HARD;3;WARNING: Replication with ADDC5 failed!
[1490962028] HOST ALERT: halleground3;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1490962048] HOST ALERT: halleground2;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1490962088] SERVICE ALERT: halleground3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1490962090] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 9% - OK!
[1490962094] HOST ALERT: halleground1;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1490962189] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 28% - OK!
[1490962192] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 87% - OK!
[1490962260] SERVICE ALERT: halleground2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1490962269] SERVICE ALERT: halleground1;Device Status;OK;HARD;3;Ready
[1490962308] SERVICE ALERT: halleground2;Device Status;OK;HARD;3;Ready 164.76.106.52
[1490962373] SERVICE ALERT: halleground1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1490962528] SERVICE ALERT: halleground3;Device Status;OK;HARD;3;Ready
[1490963313] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1490963329] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;OK;HARD;3;Cyan Cartridge HP CE261A is at 100% - OK!
[1490965985] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5689 found since last run. oneLevelSearchOps OK. searchOps is too high! 5896 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490966045] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490966464] SERVICE ALERT: hallefirst3;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1490966465] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1490966475] SERVICE ALERT: halleground1;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1490966482] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1490966489] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1490966489] SERVICE ALERT: hallegroundcolor;Image Transfer Kit HP CE249A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1490966494] HOST ALERT: hallefirst3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490966495] HOST ALERT: halleground2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490966505] HOST ALERT: halleground1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490966512] HOST ALERT: hallefirstcolor;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490966514] SERVICE ALERT: halleground2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1490966519] HOST ALERT: markjeffersonpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490966519] HOST ALERT: hallegroundcolor;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490966524] SERVICE ALERT: hallefirst3;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1490966524] HOST ALERT: porterpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490966525] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1490966528] SERVICE ALERT: porterpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1490966535] SERVICE ALERT: halleground1;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1490966542] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1490966544] HOST ALERT: halleground2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1490966549] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1490966549] SERVICE ALERT: hallegroundcolor;Image Transfer Kit HP CE249A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1490966551] HOST ALERT: hallesecondsouth;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490966554] HOST ALERT: markjeffersonpk;UP;SOFT;2;PING OK - Packet loss = 16%, RTA = 0.70 ms
[1490966554] HOST ALERT: hallefirst3;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1490966558] HOST ALERT: porterpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1490966565] HOST ALERT: halleground1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1490966565] SERVICE ALERT: hallesecondsouth;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1490966567] HOST ALERT: hallefirst2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490966572] HOST ALERT: hallefirstcolor;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1490966574] SERVICE ALERT: halleground2;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1490966578] SERVICE ALERT: halleground1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1490966579] HOST ALERT: hallegroundcolor;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1490966584] SERVICE ALERT: hallefirst3;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1490966585] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1490966586] SERVICE ALERT: hallegroundcolor;Toner Collection Unit HP CE265A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1490966588] SERVICE ALERT: porterpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1490966589] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1490966589] SERVICE ALERT: sillpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1490966595] SERVICE ALERT: halleground1;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1490966595] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1490966595] HOST ALERT: hallesecondsouth;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1490966602] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1490966603] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;OK;SOFT;3;Black Print Cartridge HP C8543X is at 28% - OK!
[1490966604] HOST ALERT: halleground2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1490966608] HOST ALERT: halleground1;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1490966609] SERVICE ALERT: hallegroundcolor;Image Transfer Kit HP CE249A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1490966614] HOST ALERT: halleground3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490966614] HOST ALERT: hallefirst3;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1490966616] SERVICE ALERT: hallegroundcolor;Yellow Cartridge HP CE262A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1490966616] HOST ALERT: hallegroundcolor;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1490966618] HOST ALERT: porterpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1490966619] HOST ALERT: hallefirst2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1490966619] HOST ALERT: sillpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490966620] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.150.252. Is the printer off?  Is the network down?
[1490966625] SERVICE ALERT: hallesecondsouth;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1490966632] SERVICE ALERT: hallesouthcommons;Black Cartridge HP CC364X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.111.220. Is the printer off?  Is the network down?
[1490966632] HOST ALERT: hallefirstcolor;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1490966634] SERVICE ALERT: halleground2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1490966638] SERVICE ALERT: halleground1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1490966644] SERVICE ALERT: alexanderpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.142.94. Is the printer off?  Is the network down?
[1490966646] SERVICE ALERT: hallegroundcolor;Toner Collection Unit HP CE265A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1490966648] SERVICE ALERT: porterpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1490966649] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1490966649] SERVICE ALERT: sillpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1490966650] HOST ALERT: studentcenterpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490966655] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1490966655] HOST ALERT: hallesecondsouth;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1490966662] HOST ALERT: hallesouthcommons;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490966664] HOST ALERT: halleground2;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1490966668] HOST ALERT: halleground1;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1490966674] HOST ALERT: alexanderpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490966676] HOST ALERT: hallegroundcolor;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1490966676] SERVICE ALERT: hallegroundcolor;Yellow Cartridge HP CE262A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1490966678] HOST ALERT: porterpk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1490966678] SERVICE ALERT: hallegroundcolor;Cyan Cartridge HP CE261A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1490966679] HOST ALERT: hallefirst2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1490966679] HOST ALERT: sillpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1490966680] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.150.252. Is the printer off?  Is the network down?
[1490966685] SERVICE ALERT: hallesecondsouth;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1490966692] SERVICE ALERT: hallesouthcommons;Black Cartridge HP CC364X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.111.220. Is the printer off?  Is the network down?
[1490966698] SERVICE ALERT: halleground1;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1490966704] HOST ALERT: halleground3;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1490966704] SERVICE ALERT: alexanderpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.142.94. Is the printer off?  Is the network down?
[1490966704] SERVICE ALERT: porterpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1490966704] HOST ALERT: hallefirst3;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1490966706] HOST ALERT: hallegroundcolor;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1490966706] SERVICE ALERT: hallegroundcolor;Toner Collection Unit HP CE265A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1490966709] SERVICE ALERT: sillpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1490966709] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1490966709] SERVICE ALERT: sillpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1490966710] HOST ALERT: studentcenterpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1490966715] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1490966715] HOST ALERT: hallesecondsouth;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1490966717] SERVICE ALERT: hallefirstcolor;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1490966722] HOST ALERT: hallesouthcommons;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1490966722] HOST ALERT: hallefirstcolor;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1490966728] HOST ALERT: halleground1;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1490966729] SERVICE ALERT: hallegroundcolor;Magenta Cartridge HP CE263A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1490966734] SERVICE ALERT: halleground3;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1490966734] HOST ALERT: alexanderpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1490966734] HOST ALERT: porterpk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1490966736] HOST ALERT: hallegroundcolor;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1490966736] SERVICE ALERT: hallegroundcolor;Yellow Cartridge HP CE262A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1490966738] SERVICE ALERT: hallegroundcolor;Cyan Cartridge HP CE261A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1490966739] HOST ALERT: sillpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1490966739] HOST ALERT: hallefirst2;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1490966740] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.150.252. Is the printer off?  Is the network down?
[1490966742] SERVICE ALERT: hallesecondsouth;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1490966750] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1490966752] SERVICE ALERT: hallesouthcommons;Black Cartridge HP CC364X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.111.220. Is the printer off?  Is the network down?
[1490966754] HOST ALERT: hallefirst1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490966754] HOST ALERT: halleground2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1490966762] SERVICE ALERT: sillpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1490966763] SERVICE ALERT: hallesouthcommons;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.111.220. Is the printer off?  Is the network down?
[1490966764] HOST ALERT: halleground3;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1490966764] SERVICE ALERT: alexanderpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.142.94. Is the printer off?  Is the network down?
[1490966764] SERVICE ALERT: porterpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1490966766] HOST ALERT: hallegroundcolor;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1490966769] SERVICE ALERT: sillpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1490966770] HOST ALERT: studentcenterpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1490966772] HOST ALERT: hallesecondsouth;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1490966776] SERVICE ALERT: hallesouthcommons;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.111.220. Is the printer off?  Is the network down?
[1490966776] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1490966777] SERVICE ALERT: hallefirstcolor;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1490966780] HOST ALERT: hallefirst3;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1490966782] HOST ALERT: hallesouthcommons;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1490966788] SERVICE ALERT: studentcenterpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.150.252. Is the printer off?  Is the network down?
[1490966789] SERVICE ALERT: hallegroundcolor;Magenta Cartridge HP CE263A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1490966792] HOST ALERT: sillpk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1490966794] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1490966794] SERVICE ALERT: halleground3;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1490966794] HOST ALERT: alexanderpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1490966794] HOST ALERT: porterpk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1490966798] SERVICE ALERT: hallegroundcolor;Cyan Cartridge HP CE261A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1490966801] SERVICE ALERT: hallegroundcolor;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1490966802] SERVICE ALERT: hallesecondsouth;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1490966802] HOST ALERT: hallefirst1;UP;SOFT;2;PING WARNING - Packet loss = 81%, RTA = 0.71 ms
[1490966803] HOST ALERT: hallefirst2;UP;SOFT;5;PING OK - Packet loss = 0%, RTA = 0.46 ms
[1490966803] HOST ALERT: hallefirstcolor;UP;SOFT;5;PING WARNING - Packet loss = 81%, RTA = 0.30 ms
[1490966804] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 15% - OK!
[1490966805] SERVICE ALERT: hallegroundcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1490966808] HOST ALERT: hallefirst3;UP;SOFT;6;PING OK - Packet loss = 0%, RTA = 0.46 ms
[1490966818] HOST ALERT: halleground1;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1490966818] HOST ALERT: studentcenterpk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1490966819] HOST ALERT: hallegroundcolor;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1490966822] SERVICE ALERT: sillpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1490966823] SERVICE ALERT: hallesouthcommons;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.111.220. Is the printer off?  Is the network down?
[1490966824] HOST ALERT: halleground3;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1490966824] SERVICE ALERT: porterpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1490966829] SERVICE ALERT: sillpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1490966830] SERVICE ALERT: hallefirst1;Device Status;OK;SOFT;2;Ready
[1490966830] HOST ALERT: hallegroundcolor;UP;SOFT;9;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1490966831] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;SOFT;3;Printer Paper Is OK! All Trays Still Full
[1490966832] HOST ALERT: hallesecondsouth;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1490966836] SERVICE ALERT: hallesouthcommons;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.111.220. Is the printer off?  Is the network down?
[1490966843] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1490966844] SERVICE ALERT: hallegroundcolor;Magenta Cartridge HP CE263A;OK;SOFT;3;Magenta Cartridge HP CE263A is at 56% - OK!
[1490966844] HOST ALERT: halleground2;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1490966848] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 36% - OK!
[1490966848] SERVICE ALERT: studentcenterpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.150.252. Is the printer off?  Is the network down?
[1490966852] HOST ALERT: sillpk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1490966853] HOST ALERT: hallesouthcommons;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1490966854] SERVICE ALERT: halleground3;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1490966854] HOST ALERT: porterpk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1490966855] SERVICE ALERT: hallegroundcolor;Device Status;OK;SOFT;2;Ready
[1490966859] SERVICE ALERT: hallegroundcolor;Image Fuser Kit HP 110V-CE246A;OK;SOFT;2;Image Fuser Kit HP 110V-CE246A, 220V-CE247A is at 45% - OK!
[1490966862] SERVICE ALERT: hallesecondsouth;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1490966873] HOST ALERT: hallesecondsouth;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1490966878] HOST ALERT: studentcenterpk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1490966882] SERVICE ALERT: sillpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1490966883] SERVICE ALERT: hallesouthcommons;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.111.220. Is the printer off?  Is the network down?
[1490966884] HOST ALERT: alexanderpk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1490966884] HOST ALERT: halleground3;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1490966888] SERVICE ALERT: eateriespk;Device Status;OK;HARD;3;Ready
[1490966893] SERVICE ALERT: halleground3;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1490966896] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1490966896] SERVICE ALERT: hallesouthcommons;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.111.220. Is the printer off?  Is the network down?
[1490966903] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1490966908] HOST ALERT: halleground1;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1490966908] SERVICE ALERT: studentcenterpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.150.252. Is the printer off?  Is the network down?
[1490966912] HOST ALERT: sillpk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1490966913] HOST ALERT: hallesouthcommons;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1490966914] HOST ALERT: hallesecondsouth;UP;SOFT;8;PING OK - Packet loss = 58%, RTA = 0.69 ms
[1490966918] SERVICE ALERT: alexanderpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.142.94. Is the printer off?  Is the network down?
[1490966923] HOST ALERT: halleground3;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1490966934] HOST ALERT: halleground2;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1490966938] HOST ALERT: studentcenterpk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1490966944] HOST ALERT: porterpk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1490966948] HOST ALERT: alexanderpk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1490966953] SERVICE ALERT: halleground3;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1490966956] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1490966957] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;OK;SOFT;3;Black Print Cartridge HP C8543X is at 28% - OK!
[1490966978] SERVICE ALERT: alexanderpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.142.94. Is the printer off?  Is the network down?
[1490966983] HOST ALERT: halleground3;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1490966986] HOST ALERT: halleground1;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1490966994] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1490966998] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1490967003] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1490967003] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1490967003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1490967003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1490967003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1490967003] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1490967003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1490967003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1490967003] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1490967003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1490967003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1490967003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1490967003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1490967003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1490967003] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1490967008] HOST ALERT: porterpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.59 ms
[1490967013] SERVICE ALERT: halleground3;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1490967013] HOST ALERT: studentcenterpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1490967016] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1490967024] SERVICE ALERT: alexanderpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.142.94. Is the printer off?  Is the network down?
[1490967025] HOST ALERT: hallesouthcommons;UP;SOFT;2;PING OK - Packet loss = 64%, RTA = 0.47 ms
[1490967038] SERVICE ALERT: alexanderpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.142.94. Is the printer off?  Is the network down?
[1490967040] HOST ALERT: alexanderpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1490967041] HOST ALERT: halleground3;UP;SOFT;2;PING WARNING - Packet loss = 86%, RTA = 9.03 ms
[1490967044] HOST ALERT: sillpk;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1490967045] HOST ALERT: halleground1;UP;SOFT;2;PING WARNING - Packet loss = 83%, RTA = 0.78 ms
[1490967048] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 28% - OK!
[1490967052] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 87% - OK!
[1490967056] HOST ALERT: halleground2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.00 ms
[1490967078] SERVICE ALERT: alexanderpk;Device Status;OK;SOFT;2;Processing job from tray 3
[1490967082] HOST ALERT: alexanderpk;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.58 ms
[1490967134] HOST ALERT: sillpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1490967135] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1490967178] SERVICE ALERT: hallefirst3;Device Status;OK;HARD;3;Ready
[1490967180] SERVICE ALERT: halleground2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1490967189] SERVICE ALERT: halleground1;Device Status;OK;HARD;3;Ready
[1490967195] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1490967196] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;OK;HARD;3;Black Cartridge HP CE260X is at 88% - OK!
[1490967198] HOST ALERT: sillpk;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.47 ms
[1490967203] SERVICE ALERT: hallegroundcolor;Image Transfer Kit HP CE249A;OK;HARD;3;Image Transfer Kit HP CE249A is at 46% - OK!
[1490967228] SERVICE ALERT: halleground2;Device Status;OK;HARD;3;Ready 164.76.106.52
[1490967242] SERVICE ALERT: porterpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 73% - OK!
[1490967255] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1490967279] SERVICE ALERT: hallesecondsouth;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1490967292] SERVICE ALERT: halleground1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1490967300] SERVICE ALERT: hallegroundcolor;Toner Collection Unit HP CE265A;OK;HARD;3;
[1490967303] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Ready
[1490967303] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Ready
[1490967309] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1490967330] SERVICE ALERT: hallegroundcolor;Yellow Cartridge HP CE262A;OK;HARD;3;Yellow Cartridge HP CE262A is at 73% - OK!
[1490967334] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 75% - OK!
[1490967346] SERVICE ALERT: hallesouthcommons;Black Cartridge HP CC364X;OK;HARD;3;Black Cartridge HP CC364X is at 13% - OK!
[1490967358] SERVICE ALERT: alexanderpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 39% - OK!
[1490967392] SERVICE ALERT: hallegroundcolor;Cyan Cartridge HP CE261A;OK;HARD;3;Cyan Cartridge HP CE261A is at 40% - OK!
[1490967418] SERVICE ALERT: porterpk;Device Status;OK;HARD;3;Ready
[1490967423] SERVICE ALERT: sillpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 25% - OK!
[1490967448] SERVICE ALERT: halleground3;Device Status;OK;HARD;3;Ready
[1490967456] SERVICE ALERT: hallesecondsouth;Device Status;OK;HARD;3;Ready
[1490967476] SERVICE ALERT: sillpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1490967477] SERVICE ALERT: hallesouthcommons;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1490967490] SERVICE ALERT: hallesouthcommons;Device Status;OK;HARD;3;Ready
[1490967502] SERVICE ALERT: studentcenterpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1490967607] SERVICE ALERT: halleground3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1490967610] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 9% - OK!
[1490967632] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1490968445] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5948 found since last run. oneLevelSearchOps OK. searchOps is too high! 6119 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490968455] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1490968455] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (24.8% change >= 20.0% threshold)
[1490968503] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1490968505] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490968563] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1490968624] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1490968672] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1490968733] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 68% - OK!
[1490969105] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5431 found since last run. oneLevelSearchOps OK. searchOps is too high! 5639 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490969165] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490969223] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Ready
[1490969223] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1490969804] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1490969864] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1490969924] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1490970786] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1490970786] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1490970786] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1490970786] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1490970786] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1490970786] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1490970786] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1490970786] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1490970786] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1490970786] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1490970786] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1490970786] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1490970786] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1490970786] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1490970786] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1490970786] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1490970977] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5078 found since last run. oneLevelSearchOps OK. searchOps is too high! 5183 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490971037] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490971462] HOST ALERT: reflin;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490971474] HOST ALERT: rmandb;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490971526] HOST ALERT: reflin;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1490971564] HOST ALERT: rmandb;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1490971565] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7205 found since last run. oneLevelSearchOps OK. searchOps is too high! 7315 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490971625] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490971653] HOST ALERT: rmandb;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1490971717] HOST ALERT: rmandb;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1490972225] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5341 found since last run. oneLevelSearchOps OK. searchOps is too high! 5656 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490972285] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490972442] HOST ALERT: shead1;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.112)
[1490972505] HOST ALERT: shead1;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.4.112)
[1490972571] HOST ALERT: shead1;UP;SOFT;3;PING OK - Packet loss = 28%, RTA = 0.53 ms
[1490972823] SERVICE ALERT: ADRODC-DET;AD Replication;OK;HARD;3;OK: replication is up and running
[1490972885] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6382 found since last run. oneLevelSearchOps OK. searchOps is too high! 6665 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490972945] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490972975] HOST ALERT: sindex2;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.111)
[1490973039] HOST ALERT: sindex2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.60 ms
[1490973255] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 22% (15 min average)
[1490973437] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5510 found since last run. oneLevelSearchOps OK. searchOps is too high! 5658 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490973496] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490974145] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5104 found since last run. oneLevelSearchOps OK. searchOps is too high! 5415 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490974205] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490974620] HOST ALERT: oms;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490974710] HOST ALERT: oms;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1490974798] HOST ALERT: oms;UP;SOFT;3;PING WARNING - Packet loss = 82%, RTA = 0.30 ms
[1490975100] HOST DOWNTIME ALERT: TARGETXTEST2012;STARTED; Host has entered a period of scheduled downtime
[1490975100] HOST DOWNTIME ALERT: SANDBOX-APP;STARTED; Host has entered a period of scheduled downtime
[1490975100] HOST DOWNTIME ALERT: MAPS01TEST;STARTED; Host has entered a period of scheduled downtime
[1490975100] HOST DOWNTIME ALERT: EVENTSTEST-2012;STARTED; Host has entered a period of scheduled downtime
[1490975100] HOST DOWNTIME ALERT: BOEAPPTEST;STARTED; Host has entered a period of scheduled downtime
[1490975100] HOST DOWNTIME ALERT: AXLIC-DEVL;STARTED; Host has entered a period of scheduled downtime
[1490975100] HOST DOWNTIME ALERT: AXIOM-TEST-2012;STARTED; Host has entered a period of scheduled downtime
[1490975100] HOST DOWNTIME ALERT: AX-QUAL;STARTED; Host has entered a period of scheduled downtime
[1490975100] HOST DOWNTIME ALERT: AX-DEVL;STARTED; Host has entered a period of scheduled downtime
[1490975100] HOST DOWNTIME ALERT: ADDC3;STARTED; Host has entered a period of scheduled downtime
[1490975296] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6371 found since last run. oneLevelSearchOps OK. searchOps is too high! 6453 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490975356] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490975956] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6030 found since last run. oneLevelSearchOps OK. searchOps is too high! 6112 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490976016] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490976616] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5908 found since last run. oneLevelSearchOps OK. searchOps is too high! 5997 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490976676] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490976854] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1490976914] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1490976974] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1490976974] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1490977574] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1490977803] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (26.0% change >= 20.0% threshold)
[1490977803] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1490977803] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1490977803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1490977803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1490977803] HOST DOWNTIME ALERT: ADDC3;STARTED; Host has entered a period of scheduled downtime
[1490977803] HOST DOWNTIME ALERT: AX-DEVL;STARTED; Host has entered a period of scheduled downtime
[1490977803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1490977803] HOST DOWNTIME ALERT: AX-QUAL;STARTED; Host has entered a period of scheduled downtime
[1490977803] HOST DOWNTIME ALERT: AXIOM-TEST-2012;STARTED; Host has entered a period of scheduled downtime
[1490977803] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1490977803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1490977803] HOST DOWNTIME ALERT: AXLIC-DEVL;STARTED; Host has entered a period of scheduled downtime
[1490977803] HOST DOWNTIME ALERT: BOEAPPTEST;STARTED; Host has entered a period of scheduled downtime
[1490977803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1490977803] HOST DOWNTIME ALERT: EVENTSTEST-2012;STARTED; Host has entered a period of scheduled downtime
[1490977803] HOST DOWNTIME ALERT: MAPS01TEST;STARTED; Host has entered a period of scheduled downtime
[1490977803] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1490977803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1490977803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1490977803] HOST DOWNTIME ALERT: SANDBOX-APP;STARTED; Host has entered a period of scheduled downtime
[1490977803] HOST DOWNTIME ALERT: TARGETXTEST2012;STARTED; Host has entered a period of scheduled downtime
[1490977803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1490977803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1490977803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1490977803] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1490978174] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 67% (15 min average)
[1490978405] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6190 found since last run. oneLevelSearchOps OK. searchOps is too high! 6453 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490978464] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490978773] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1490978833] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1490978893] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1490979076] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5126 found since last run. oneLevelSearchOps OK. searchOps is too high! 5209 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490979136] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490979493] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 80% (15 min average)
[1490980149] SERVICE ALERT: BOEAPPTEST;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1490980178] HOST ALERT: admdemo.emich.edu;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490980209] SERVICE ALERT: BOEAPPTEST;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1490980268] HOST ALERT: admdemo.emich.edu;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1490980269] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1490980271] HOST ALERT: bandemo.emich.edu;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490980320] HOST ALERT: appqualdb;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490980335] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6273 found since last run. oneLevelSearchOps OK. searchOps is too high! 6371 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490980345] HOST ALERT: apptestdb;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490980358] HOST ALERT: admdemo.emich.edu;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1490980361] HOST ALERT: bandemo.emich.edu;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1490980395] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490980410] HOST ALERT: appqualdb;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1490980435] HOST ALERT: apptestdb;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1490980448] HOST ALERT: admdemo.emich.edu;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1490980451] HOST ALERT: bandemo.emich.edu;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1490980484] HOST ALERT: darkdb;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490980500] HOST ALERT: appqualdb;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1490980500] SERVICE ALERT: ADDC3;VMWare Tools;CRITICAL;SOFT;1;connect to address 164.76.7.203 and port 12489: Connection refused
[1490980523] SERVICE ALERT: ADDC3;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1490980525] HOST ALERT: apptestdb;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1490980538] HOST ALERT: admdemo.emich.edu;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1490980541] HOST ALERT: bandemo.emich.edu;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1490980560] SERVICE ALERT: ADDC3;VMWare Tools;OK;SOFT;2;VMTools: Started
[1490980564] HOST ALERT: appqualdb;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1490980574] HOST ALERT: darkdb;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1490980578] SERVICE ALERT: ADDC3;Service: Windows Time;CRITICAL;SOFT;2;W32Time: Stopped
[1490980589] HOST ALERT: apptestdb;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1490980591] HOST ALERT: r25testdb;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490980628] HOST ALERT: admdemo.emich.edu;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1490980631] HOST ALERT: bandemo.emich.edu;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1490980633] HOST ALERT: irimtestdb;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490980638] SERVICE ALERT: ADDC3;Service: Windows Time;CRITICAL;HARD;3;W32Time: Stopped
[1490980662] HOST ALERT: boetestdb;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490980664] HOST ALERT: darkdb;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1490980681] HOST ALERT: r25testdb;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1490980718] HOST ALERT: admdemo.emich.edu;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1490980721] HOST ALERT: bandemo.emich.edu;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1490980723] HOST ALERT: irimtestdb;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1490980743] HOST ALERT: darkdb;UP;SOFT;4;PING OK - Packet loss = 75%, RTA = 0.35 ms
[1490980752] HOST ALERT: boetestdb;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1490980771] HOST ALERT: r25testdb;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1490980785] HOST ALERT: bandemo.emich.edu;UP;SOFT;7;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1490980787] HOST ALERT: irimtestdb;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1490980788] SERVICE ALERT: AXLIC-DEVL;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.7.142 and port 12489: Connection refused
[1490980793] HOST ALERT: AXLIC-DEVL;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490980806] HOST ALERT: admdemo.emich.edu;UP;SOFT;8;PING WARNING - Packet loss = 82%, RTA = 0.30 ms
[1490980816] HOST ALERT: boetestdb;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1490980818] SERVICE ALERT: AX-QUAL;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.7.143 and port 12489: Connection refused
[1490980819] SERVICE ALERT: AX-QUAL;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.7.143 and port 12489: Connection refused
[1490980820] SERVICE ALERT: AX-QUAL;VMWare Tools;CRITICAL;SOFT;1;connect to address 164.76.7.143 and port 12489: Connection refused
[1490980825] SERVICE ALERT: AX-QUAL;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1490980835] HOST ALERT: r25testdb;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1490980844] SERVICE ALERT: EVENTSTEST-2012;Service: NetLogin;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1490980848] SERVICE ALERT: AXLIC-DEVL;Service: RPC;OK;SOFT;2;RpcSs: Started
[1490980849] HOST ALERT: AXLIC-DEVL;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.36 ms
[1490980865] SERVICE ALERT: AX-QUAL;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1490980869] SERVICE ALERT: BOEAPPTEST;CPU Load;OK;HARD;3;CPU Load 89% (15 min average)
[1490980871] SERVICE ALERT: AX-QUAL;VMWare Tools;OK;SOFT;2;VMTools: Started
[1490980875] SERVICE ALERT: AX-QUAL;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1490980876] SERVICE ALERT: EVENTSTEST-2012;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.4.145 and port 12489: Connection refused
[1490980879] SERVICE ALERT: AX-QUAL;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1490980881] SERVICE ALERT: EVENTSTEST-2012;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1490980881] SERVICE ALERT: AXIOM-TEST-2012;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1490980881] SERVICE ALERT: AXIOM-TEST-2012;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1490980885] SERVICE ALERT: AXIOM-TEST-2012;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1490980885] SERVICE ALERT: AXIOM-TEST-2012;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1490980904] SERVICE ALERT: EVENTSTEST-2012;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1490980921] SERVICE ALERT: EVENTSTEST-2012;Service: RPC;OK;SOFT;2;RpcSs: Started
[1490980921] SERVICE ALERT: AXIOM-TEST-2012;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1490980921] SERVICE ALERT: AXIOM-TEST-2012;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1490980926] SERVICE ALERT: AXIOM-TEST-2012;CPU Load;OK;SOFT;2;CPU Load 80% (15 min average)
[1490980926] SERVICE ALERT: AXIOM-TEST-2012;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1490980936] SERVICE ALERT: EVENTSTEST-2012;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1490980995] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5757 found since last run. oneLevelSearchOps OK. searchOps is too high! 5761 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490981026] SERVICE ALERT: TARGETXTEST2012;Service: WMI;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1490981042] SERVICE ALERT: MAPS01TEST;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.4.162 and port 12489: Connection refused
[1490981046] SERVICE ALERT: MAPS01TEST;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.4.162 and port 12489: Connection refused
[1490981055] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490981061] HOST ALERT: MAPS01TEST;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490981087] SERVICE ALERT: TARGETXTEST2012;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1490981093] HOST ALERT: MAPS01TEST;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1490981102] SERVICE ALERT: MAPS01TEST;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1490981106] SERVICE ALERT: MAPS01TEST;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1490981238] SERVICE ALERT: ADDC3;Service: Windows Time;OK;HARD;3;W32Time: Started
[1490981260] HOST ALERT: EVENTSTEST-2012;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.145)
[1490981323] HOST ALERT: EVENTSTEST-2012;DOWN;HARD;2;CRITICAL - Host Unreachable (164.76.4.145)
[1490981385] HOST ALERT: bandevl;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490981389] HOST ALERT: banseed;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490981392] HOST ALERT: banqual;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490981447] HOST ALERT: EVENTSTEST-2012;UP;HARD;2;PING OK - Packet loss = 0%, RTA = 0.88 ms
[1490981449] HOST ALERT: bandevl;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1490981453] HOST ALERT: banseed;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1490981456] HOST ALERT: banqual;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1490981461] SERVICE ALERT: SANDBOX-APP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1490981473] SERVICE ALERT: BOEAPPTEST;Service: Group Policy;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1490981496] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.2.48 and port 12489: Connection refused
[1490981496] SERVICE ALERT: BOEAPPTEST;Service: NetLogin;CRITICAL;SOFT;1;connect to address 164.76.2.48 and port 12489: Connection refused
[1490981505] SERVICE ALERT: BOEAPPTEST;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.2.48 and port 12489: Connection refused
[1490981521] SERVICE ALERT: SANDBOX-APP;CPU Load;OK;SOFT;2;CPU Load 91% (15 min average)
[1490981527] SERVICE ALERT: TARGETXTEST2012;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1490981536] SERVICE ALERT: BOEAPPTEST;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1490981556] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;OK;SOFT;2;c:\ - total: 99.90 Gb - used: 84.53 Gb (85%) - free 15.37 Gb (15%)
[1490981556] SERVICE ALERT: BOEAPPTEST;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1490981565] SERVICE ALERT: BOEAPPTEST;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1490981587] SERVICE ALERT: TARGETXTEST2012;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1490981604] HOST ALERT: authtestdb;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490981647] SERVICE ALERT: TARGETXTEST2012;CPU Load;OK;SOFT;3;CPU Load 81% (15 min average)
[1490981655] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5263 found since last run. oneLevelSearchOps OK. searchOps is too high! 5268 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490981694] HOST ALERT: authtestdb;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1490981715] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490981763] HOST ALERT: authtestdb;UP;SOFT;3;PING OK - Packet loss = 50%, RTA = 0.28 ms
[1490981797] SERVICE ALERT: hallegroundcolor;Cyan Cartridge HP CE261A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1490981827] HOST ALERT: hallegroundcolor;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490981832] SERVICE ALERT: hallegroundcolor;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1490981847] SERVICE ALERT: hallegroundcolor;Magenta Cartridge HP CE263A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1490981857] SERVICE ALERT: hallegroundcolor;Cyan Cartridge HP CE261A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1490981860] SERVICE ALERT: hallegroundcolor;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1490981862] HOST ALERT: hallegroundcolor;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1490981864] SERVICE ALERT: hallegroundcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1490981889] SERVICE ALERT: hallegroundcolor;Black Cartridge HP CE260X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1490981892] SERVICE ALERT: hallegroundcolor;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1490981894] HOST ALERT: hallegroundcolor;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1490981907] SERVICE ALERT: hallegroundcolor;Magenta Cartridge HP CE263A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1490981917] SERVICE ALERT: hallegroundcolor;Cyan Cartridge HP CE261A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1490981920] SERVICE ALERT: hallegroundcolor;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1490981924] SERVICE ALERT: hallegroundcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1490981937] HOST ALERT: hallegroundcolor;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1490981949] SERVICE ALERT: hallegroundcolor;Black Cartridge HP CE260X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1490981952] SERVICE ALERT: hallegroundcolor;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1490981967] SERVICE ALERT: hallegroundcolor;Magenta Cartridge HP CE263A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1490981979] HOST ALERT: hallegroundcolor;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1490981980] SERVICE ALERT: hallegroundcolor;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1490981984] SERVICE ALERT: hallegroundcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1490982009] SERVICE ALERT: hallegroundcolor;Black Cartridge HP CE260X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1490982010] HOST ALERT: hallegroundcolor;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1490982100] HOST ALERT: hallegroundcolor;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1490982121] SERVICE ALERT: SANDBOX-APP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1490982181] SERVICE ALERT: SANDBOX-APP;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1490982190] HOST ALERT: hallegroundcolor;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1490982208] SERVICE ALERT: hallegroundcolor;Image Transfer Kit HP CE249A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1490982238] HOST ALERT: hallegroundcolor;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1490982241] SERVICE ALERT: SANDBOX-APP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1490982247] SERVICE ALERT: TARGETXTEST2012;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1490982268] SERVICE ALERT: hallegroundcolor;Image Transfer Kit HP CE249A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1490982298] HOST ALERT: hallegroundcolor;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1490982305] SERVICE ALERT: hallegroundcolor;Toner Collection Unit HP CE265A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1490982307] SERVICE ALERT: TARGETXTEST2012;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1490982315] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7476 found since last run. oneLevelSearchOps OK. searchOps is too high! 7571 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490982328] SERVICE ALERT: hallegroundcolor;Image Transfer Kit HP CE249A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1490982336] SERVICE ALERT: hallegroundcolor;Yellow Cartridge HP CE262A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1490982365] SERVICE ALERT: hallegroundcolor;Toner Collection Unit HP CE265A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1490982367] SERVICE ALERT: TARGETXTEST2012;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1490982375] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490982396] SERVICE ALERT: hallegroundcolor;Yellow Cartridge HP CE262A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1490982425] SERVICE ALERT: hallegroundcolor;Toner Collection Unit HP CE265A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1490982456] SERVICE ALERT: hallegroundcolor;Yellow Cartridge HP CE262A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1490982669] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1490982729] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1490982789] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1490982967] SERVICE ALERT: TARGETXTEST2012;CPU Load;OK;HARD;3;CPU Load 7% (15 min average)
[1490982975] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5812 found since last run. oneLevelSearchOps OK. searchOps is too high! 5819 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490983017] HOST ALERT: mydevlapp;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490983021] HOST ALERT: myqualapp2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490983035] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490983081] HOST ALERT: mydevlapp;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1490983085] HOST ALERT: myqualapp2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1490983091] HOST ALERT: mydevldb;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490983097] HOST ALERT: myqualdb;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1490983181] HOST ALERT: mydevldb;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1490983187] HOST ALERT: myqualdb;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1490983245] HOST ALERT: mydevldb;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.36 ms
[1490983251] HOST ALERT: myqualdb;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1490983263] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5117 found since last run. oneLevelSearchOps OK. searchOps is too high! 5373 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490983323] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490983389] SERVICE ALERT: BOEAPPTEST;CPU Load;OK;HARD;3;CPU Load 3% (15 min average)
[1490983389] SERVICE FLAPPING ALERT: BOEAPPTEST;CPU Load;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1490983441] SERVICE ALERT: SANDBOX-APP;CPU Load;OK;HARD;3;CPU Load 0% (15 min average)
[1490983923] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7104 found since last run. oneLevelSearchOps OK. searchOps is too high! 7281 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490983983] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490984292] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1490984350] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1490984410] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 93% (15 min average)
[1490984583] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5075 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490984643] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490984924] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 15% - OK!
[1490985010] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1490985070] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1490985130] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1490985243] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7192 found since last run. oneLevelSearchOps OK. searchOps is too high! 7461 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490985303] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490985903] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7954 found since last run. oneLevelSearchOps OK. searchOps is too high! 8292 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490985963] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490986330] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 71% (15 min average)
[1490986563] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5685 found since last run. oneLevelSearchOps OK. searchOps is too high! 5882 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490986623] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490988730] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1490988790] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 79% (15 min average)
[1490990823] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5341 found since last run. oneLevelSearchOps OK. searchOps is too high! 5527 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490990883] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490991483] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5524 found since last run. oneLevelSearchOps OK. searchOps is too high! 5628 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490991543] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490991790] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1490991850] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1490992743] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5597 found since last run. oneLevelSearchOps OK. searchOps is too high! 5840 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490992803] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490993100] HOST DOWNTIME ALERT: EVENTSTEST-2012;STOPPED; Host has exited from a period of scheduled downtime
[1490993100] HOST DOWNTIME ALERT: AX-DEVL;STOPPED; Host has exited from a period of scheduled downtime
[1490993100] HOST DOWNTIME ALERT: ADDC3;STOPPED; Host has exited from a period of scheduled downtime
[1490993100] HOST DOWNTIME ALERT: AX-QUAL;STOPPED; Host has exited from a period of scheduled downtime
[1490993100] HOST DOWNTIME ALERT: AXIOM-TEST-2012;STOPPED; Host has exited from a period of scheduled downtime
[1490993100] HOST DOWNTIME ALERT: BOEAPPTEST;STOPPED; Host has exited from a period of scheduled downtime
[1490993100] HOST DOWNTIME ALERT: AXLIC-DEVL;STOPPED; Host has exited from a period of scheduled downtime
[1490993100] HOST DOWNTIME ALERT: MAPS01TEST;STOPPED; Host has exited from a period of scheduled downtime
[1490993100] HOST DOWNTIME ALERT: SANDBOX-APP;STOPPED; Host has exited from a period of scheduled downtime
[1490993100] HOST DOWNTIME ALERT: TARGETXTEST2012;STOPPED; Host has exited from a period of scheduled downtime
[1490993403] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5132 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490993463] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1490993650] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1490993710] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1490993770] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 68% (15 min average)
[1490994788] SERVICE FLAPPING ALERT: BOEAPPTEST;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1490995570] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1490995630] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1490995690] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 90% (15 min average)
[1490996890] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1490996950] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1490997010] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1490997610] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1490997984] SERVICE ALERT: ADDC5;AD Group Changes;UNKNOWN;SOFT;1;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1490998203] SERVICE ALERT: ADDC5;AD Group Changes;OK;SOFT;2;WARNING:  0  groups have changed!
[1490999463] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5022 found since last run. oneLevelSearchOps OK. searchOps is too high! 5227 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1490999523] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491000123] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5316 found since last run. oneLevelSearchOps OK. searchOps is too high! 5536 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491000183] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491001035] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5757 found since last run. oneLevelSearchOps OK. searchOps is too high! 5761 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491001095] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491001810] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1491001870] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 89% (15 min average)
[1491003070] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1491003130] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1491003190] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1491003790] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 42% (15 min average)
[1491004390] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1491004450] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1491004510] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1491005110] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 90% (15 min average)
[1491005217] SERVICE ALERT: ADRODC-DET;AD Replication;WARNING;SOFT;1;WARNING: Replication with ADDC5 failed!
[1491005277] SERVICE ALERT: ADRODC-DET;AD Replication;WARNING;SOFT;2;WARNING: Replication with ADDC5 failed!
[1491005337] SERVICE ALERT: ADRODC-DET;AD Replication;WARNING;HARD;3;WARNING: Replication with ADDC5 failed!
[1491005710] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1491005770] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1491005830] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1491006430] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 32% (15 min average)
[1491007030] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1491007090] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1491007150] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 71% (15 min average)
[1491009588] SERVICE ALERT: TARGETX-2012;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1491009648] SERVICE ALERT: TARGETX-2012;CPU Load;OK;SOFT;2;CPU Load 17% (15 min average)
[1491010200] HOST DOWNTIME ALERT: IMAGES;STARTED; Host has entered a period of scheduled downtime
[1491010750] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1491010810] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1491010870] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1491011241] SERVICE ALERT: IMAGES;MDT Monitor;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1491011246] HOST ALERT: IMAGES;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1491011310] HOST ALERT: IMAGES;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1491011821] SERVICE ALERT: IMAGES;MDT Monitor;UNKNOWN;HARD;3;Failed to open service MDT_Monitor: 1060: The specified service does not exist as an installed service.
[1491012048] SERVICE ALERT: TARGETX-2012;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1491012108] SERVICE ALERT: TARGETX-2012;CPU Load;OK;SOFT;2;CPU Load 84% (15 min average)
[1491012684] SERVICE ALERT: ADDC5;AD Group Changes;UNKNOWN;SOFT;1;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1491012908] SERVICE ALERT: ADDC5;AD Group Changes;OK;SOFT;2;WARNING:  0  groups have changed!
[1491013270] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 77% (15 min average)
[1491014470] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1491014530] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1491014590] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1491015190] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 73% (15 min average)
[1491015790] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1491015850] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1491015910] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1491016142] SERVICE ALERT: ADRODC-DET;AD Replication;OK;HARD;3;OK: replication is up and running
[1491017110] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 82% (15 min average)
[1491017400] HOST DOWNTIME ALERT: IMAGES;STOPPED; Host has exited from a period of scheduled downtime
[1491017710] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1491017770] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 81% (15 min average)