[1457068265] SERVICE ALERT: halleground2;Device Status;OK;HARD;3;INPUT PAPER PATH OPEN 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1457100426] SERVICE ALERT: porterpk;Device Status;OK;HARD;3;Processing job from tray 3
[1457100589] SERVICE ALERT: porterpk;Black Cartridge HP CC364X;OK;HARD;3;Black Cartridge HP CC364X is at 73% - OK!
[1457100978] SERVICE ALERT: porterpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1457101024] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;1;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1457101144] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;2;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1457101264] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;HARD;3;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1457111885] SERVICE ALERT: ADDC5;GPO Replication;WARNING;SOFT;1;Warning! 1 policies are not replicating properly!
[1457112222] SERVICE ALERT: ADDC5;GPO Replication;WARNING;SOFT;2;Warning! 1 policies are not replicating properly!
[1457112485] SERVICE ALERT: ADDC5;GPO Replication;WARNING;HARD;3;Warning! 1 policies are not replicating properly!
[1457113264] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;OK;HARD;3;OK: No entries found
[1457119696] SERVICE ALERT: ADDC5;GPO Replication;OK;HARD;3;OK! All GPOs are properly replicated.
[1457119718] SERVICE ALERT: LANSWEEPER;VMWare Tools;CRITICAL;SOFT;1;VMTools: Stopped
[1457119753] SERVICE ALERT: LANSWEEPER;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457119783] HOST ALERT: LANSWEEPER;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1457119858] SERVICE ALERT: LANSWEEPER;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457119873] SERVICE ALERT: LANSWEEPER;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457119873] HOST ALERT: LANSWEEPER;DOWN;SOFT;2;(Host check timed out after 30.00 seconds)
[1457119902] SERVICE ALERT: LANSWEEPER;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457119932] HOST ALERT: LANSWEEPER;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1457119957] SERVICE ALERT: LANSWEEPER;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457119958] SERVICE ALERT: LANSWEEPER;VMWare Tools;OK;SOFT;3;VMTools: Started
[1457119961] HOST ALERT: LANSWEEPER;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.41 ms
[1457119969] SERVICE ALERT: LANSWEEPER2;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457119973] SERVICE ALERT: LANSWEEPER;Service: Windows Time;OK;SOFT;3;W32Time: Started
[1457119999] HOST ALERT: LANSWEEPER2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1457120002] SERVICE ALERT: LANSWEEPER;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1457120066] SERVICE ALERT: LANSWEEPER;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1457120089] SERVICE ALERT: LANSWEEPER2;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457120089] HOST ALERT: LANSWEEPER2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1457120119] SERVICE ALERT: LANSWEEPER2;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457120141] SERVICE ALERT: LANSWEEPER2;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457120145] SERVICE ALERT: LANSWEEPER2;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457120149] HOST ALERT: LANSWEEPER2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1457120177] SERVICE ALERT: LANSWEEPER2;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457120207] HOST ALERT: LANSWEEPER2;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1457120209] SERVICE ALERT: LANSWEEPER2;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457120225] SERVICE ALERT: LANSWEEPER2;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457120239] SERVICE ALERT: LANSWEEPER2;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457120239] HOST ALERT: LANSWEEPER2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1457120239] SERVICE ALERT: LANSWEEPER2;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457120261] SERVICE ALERT: LANSWEEPER2;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457120265] SERVICE ALERT: LANSWEEPER2;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457120269] HOST ALERT: LANSWEEPER2;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1457120282] SERVICE ALERT: LANSWEEPER2;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457120297] SERVICE ALERT: LANSWEEPER2;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457120312] HOST ALERT: LANSWEEPER2;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1457120323] SERVICE ALERT: LANSWEEPER2;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457120345] SERVICE ALERT: LANSWEEPER2;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457120353] HOST ALERT: LANSWEEPER2;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1457120359] SERVICE ALERT: LANSWEEPER2;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457120359] SERVICE ALERT: LANSWEEPER2;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457120381] SERVICE ALERT: LANSWEEPER2;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457120385] SERVICE ALERT: LANSWEEPER2;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457120389] HOST ALERT: LANSWEEPER2;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1457120402] SERVICE ALERT: LANSWEEPER2;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457120414] SERVICE ALERT: LANSWEEPER2;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457120417] SERVICE ALERT: LANSWEEPER2;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457120432] HOST ALERT: LANSWEEPER2;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1457120443] SERVICE ALERT: LANSWEEPER2;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457120453] SERVICE ALERT: LANSWEEPER2;SSL Certificate 14 Days With Auto Renew;WARNING;SOFT;1;(No output on stdout) stderr: Could not resolve hostname LANSWEEPER2: Name or service not known
[1457120465] SERVICE ALERT: LANSWEEPER2;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457120465] SERVICE ALERT: LANSWEEPER2;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457120468] SERVICE ALERT: LANSWEEPER2;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457120479] SERVICE ALERT: LANSWEEPER2;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457120517] SERVICE ALERT: LANSWEEPER2;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457120518] SERVICE ALERT: LANSWEEPER2;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457120522] SERVICE ALERT: LANSWEEPER2;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457120534] SERVICE ALERT: LANSWEEPER2;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457120563] SERVICE ALERT: LANSWEEPER2;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457120573] SERVICE ALERT: LANSWEEPER2;SSL Certificate 14 Days With Auto Renew;WARNING;SOFT;2;(No output on stdout) stderr: Could not resolve hostname LANSWEEPER2: Name or service not known
[1457120585] SERVICE ALERT: LANSWEEPER2;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457120588] SERVICE ALERT: LANSWEEPER2;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457120637] SERVICE ALERT: LANSWEEPER2;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457120638] SERVICE ALERT: LANSWEEPER2;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457120654] SERVICE ALERT: LANSWEEPER2;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457120693] SERVICE ALERT: LANSWEEPER2;SSL Certificate 14 Days With Auto Renew;WARNING;HARD;3;(No output on stdout) stderr: Could not resolve hostname LANSWEEPER2: Name or service not known
[1457120705] SERVICE ALERT: LANSWEEPER2;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457120708] SERVICE ALERT: LANSWEEPER2;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457120757] SERVICE ALERT: LANSWEEPER2;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457120758] SERVICE ALERT: LANSWEEPER2;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457121344] SERVICE ALERT: LANSWEEPER;CPU Load;WARNING;SOFT;1;CPU Load 91% (5 min average)
[1457121464] SERVICE ALERT: LANSWEEPER;CPU Load;OK;SOFT;2;CPU Load 16% (5 min average)
[1457121661] SERVICE ALERT: studentcenterpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1457121721] SERVICE ALERT: studentcenterpk;Device Status;OK;SOFT;2;DUPLEXER ERROR REMOVE DUPLEXER Install duplexer with power off
[1457122748] HOST ALERT: hallethirdsouth;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1457130187] SERVICE ALERT: DB01-2012;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1457130307] SERVICE ALERT: DB01-2012;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1457130427] SERVICE ALERT: DB01-2012;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1457130690] SERVICE ALERT: DB01-2012;Service: RPC;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1457130713] SERVICE ALERT: DB01-2012;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.99.225 and port 12489: Connection refused
[1457130810] SERVICE ALERT: DB01-2012;Service: RPC;OK;SOFT;2;RpcSs: Started
[1457130833] SERVICE ALERT: DB01-2012;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1457131027] SERVICE ALERT: DB01-2012;CPU Load;WARNING;HARD;3;CPU Load 93% (5 min average)
[1457131627] SERVICE ALERT: DB01-2012;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1457132133] HOST DOWNTIME ALERT: CASPER;STARTED; Host has entered a period of scheduled downtime
[1457132133] HOST DOWNTIME ALERT: ADDC6;STARTED; Host has entered a period of scheduled downtime
[1457132133] HOST DOWNTIME ALERT: ADRODC-LIV;STARTED; Host has entered a period of scheduled downtime
[1457132133] HOST DOWNTIME ALERT: CSCFILESERV;STARTED; Host has entered a period of scheduled downtime
[1457132133] HOST DOWNTIME ALERT: NETIDAD;STARTED; Host has entered a period of scheduled downtime
[1457132133] HOST DOWNTIME ALERT: DB01-2012;STARTED; Host has entered a period of scheduled downtime
[1457132133] HOST DOWNTIME ALERT: EMUPRINTPS;STARTED; Host has entered a period of scheduled downtime
[1457132133] HOST DOWNTIME ALERT: ADRODC-DET;STARTED; Host has entered a period of scheduled downtime
[1457132133] HOST DOWNTIME ALERT: ADAUTHDC1;STARTED; Host has entered a period of scheduled downtime
[1457132133] HOST DOWNTIME ALERT: ADAUTHDC2;STARTED; Host has entered a period of scheduled downtime
[1457132133] HOST DOWNTIME ALERT: ADAUTHDC3;STARTED; Host has entered a period of scheduled downtime
[1457132133] HOST DOWNTIME ALERT: ADCA1;STARTED; Host has entered a period of scheduled downtime
[1457132133] HOST DOWNTIME ALERT: ADCA2;STARTED; Host has entered a period of scheduled downtime
[1457132133] HOST DOWNTIME ALERT: ADDC1;STARTED; Host has entered a period of scheduled downtime
[1457132133] HOST DOWNTIME ALERT: ADDFS01;STARTED; Host has entered a period of scheduled downtime
[1457132133] HOST DOWNTIME ALERT: ADWU1;STARTED; Host has entered a period of scheduled downtime
[1457132402] HOST DOWNTIME ALERT: ADAUTHDC1;STARTED; Host has entered a period of scheduled downtime
[1457132402] HOST DOWNTIME ALERT: ADAUTHDC2;STARTED; Host has entered a period of scheduled downtime
[1457132402] HOST DOWNTIME ALERT: ADCA1;STARTED; Host has entered a period of scheduled downtime
[1457132402] HOST DOWNTIME ALERT: ADDFS01;STARTED; Host has entered a period of scheduled downtime
[1457132402] HOST DOWNTIME ALERT: NETIDAD;STARTED; Host has entered a period of scheduled downtime
[1457132402] HOST DOWNTIME ALERT: ADRODC-DET;STARTED; Host has entered a period of scheduled downtime
[1457132402] HOST DOWNTIME ALERT: ADCA2;STARTED; Host has entered a period of scheduled downtime
[1457132402] HOST DOWNTIME ALERT: ADRODC-LIV;STARTED; Host has entered a period of scheduled downtime
[1457132402] HOST DOWNTIME ALERT: ADWU1;STARTED; Host has entered a period of scheduled downtime
[1457132402] HOST DOWNTIME ALERT: ADAUTHDC3;STARTED; Host has entered a period of scheduled downtime
[1457132402] HOST DOWNTIME ALERT: ADDC1;STARTED; Host has entered a period of scheduled downtime
[1457132402] HOST DOWNTIME ALERT: CASPER;STARTED; Host has entered a period of scheduled downtime
[1457132402] HOST DOWNTIME ALERT: CSCFILESERV;STARTED; Host has entered a period of scheduled downtime
[1457132402] HOST DOWNTIME ALERT: ADDC6;STARTED; Host has entered a period of scheduled downtime
[1457132402] HOST DOWNTIME ALERT: DB01-2012;STARTED; Host has entered a period of scheduled downtime
[1457132402] HOST DOWNTIME ALERT: EMUPRINTPS;STARTED; Host has entered a period of scheduled downtime
[1457132827] SERVICE ALERT: DB01-2012;CPU Load;OK;HARD;3;CPU Load 1% (5 min average)
[1457132827] SERVICE FLAPPING ALERT: DB01-2012;CPU Load;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1457133105] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457133225] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 89% (5 min average)
[1457133280] SERVICE ALERT: ADCA1;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457133400] SERVICE ALERT: ADCA1;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1457133445] SERVICE ALERT: ADDC1;Service: Firewall;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1457133463] SERVICE ALERT: ADDC1;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.99.235 and port 12489: Connection refused
[1457133467] SERVICE ALERT: ADDC1;Service: Key Distribution Center;CRITICAL;SOFT;1;connect to address 164.76.99.235 and port 12489: Connection refused
[1457133489] SERVICE ALERT: ADDC1;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457133520] SERVICE ALERT: ADCA1;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1457133565] SERVICE ALERT: ADDC1;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1457133583] SERVICE ALERT: ADDC1;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1457133587] SERVICE ALERT: ADDC1;Service: Key Distribution Center;OK;SOFT;2;KDC: Started
[1457133589] SERVICE ALERT: ADDC1;CPU Load;OK;SOFT;2;CPU Load 25% (5 min average)
[1457133825] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1457133945] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 88% (5 min average)
[1457133946] SERVICE ALERT: ADDC6;Remote Loader for IDM;CRITICAL;SOFT;1;dirxml_remote.exe: not running
[1457134066] SERVICE ALERT: ADDC6;Remote Loader for IDM;CRITICAL;SOFT;2;dirxml_remote.exe: not running
[1457134186] SERVICE ALERT: ADDC6;Remote Loader for IDM;CRITICAL;HARD;3;dirxml_remote.exe: not running
[1457134406] SERVICE ALERT: ADDFS01;Service: User Profile;CRITICAL;SOFT;1;connect to address 164.76.99.199 and port 12489: Connection refused
[1457134409] SERVICE ALERT: ADWU1;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.99.241 and port 12489: Connection refused
[1457134521] SERVICE ALERT: ADDFS01;CPU Load;WARNING;SOFT;1;CPU Load 91% (5 min average)
[1457134526] SERVICE ALERT: ADDFS01;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1457134529] SERVICE ALERT: ADWU1;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1457134545] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457134632] HOST ALERT: CSCFILESERV;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.131)
[1457134632] SERVICE ALERT: CSCFILESERV;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.102.131 and port 12489: No route to host
[1457134641] SERVICE ALERT: ADDFS01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1457134665] HOST ALERT: CASPER;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1457134665] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 66% (5 min average)
[1457134679] SERVICE ALERT: ADCA2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457134717] HOST ALERT: CASPER;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.39 ms
[1457134720] SERVICE ALERT: ADCA1;CPU Load;OK;HARD;3;CPU Load 7% (5 min average)
[1457134742] SERVICE ALERT: CSCFILESERV;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1457134746] HOST ALERT: CSCFILESERV;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1457134761] SERVICE ALERT: ADDFS01;CPU Load;OK;SOFT;3;CPU Load 66% (5 min average)
[1457134786] SERVICE ALERT: ADDC6;Remote Loader for IDM;OK;HARD;3;dirxml_remote.exe: Running
[1457134799] SERVICE ALERT: ADCA2;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1457135265] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457135327] SERVICE ALERT: ADAUTHDC1;Service: DFS Namespace;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1457135376] SERVICE ALERT: ADAUTHDC1;Service: Dell Backup Agent;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457135385] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1457135388] SERVICE ALERT: ADAUTHDC3;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457135447] SERVICE ALERT: ADAUTHDC1;Service: DFS Namespace;OK;SOFT;2;Dfs: Started
[1457135451] SERVICE ALERT: ADAUTHDC1;AD Replication;WARNING;SOFT;1;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1457135476] SERVICE ALERT: ADAUTHDC1;Service: Dell Backup Agent;CRITICAL;SOFT;2;ErdAgent: Starting
[1457135488] SERVICE ALERT: ADAUTHDC3;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1457135505] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1457135569] SERVICE ALERT: ADAUTHDC1;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1457135573] SERVICE ALERT: ADAUTHDC1;AD Replication;WARNING;SOFT;2;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1457135596] SERVICE ALERT: ADAUTHDC1;Service: Dell Backup Agent;OK;SOFT;3;ErdAgent: Started
[1457135687] SERVICE ALERT: ADAUTHDC2;AD Replication;WARNING;SOFT;1;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1457135688] SERVICE ALERT: ADAUTHDC1;AD Replication;WARNING;HARD;3;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1457135689] SERVICE ALERT: ADAUTHDC1;CPU Load;WARNING;SOFT;2;CPU Load 90% (5 min average)
[1457135734] SERVICE ALERT: ADAUTHDC3;AD Replication;UNKNOWN;SOFT;1;CHECK_NRPE: Invalid packet type received from server.
[1457135798] SERVICE ALERT: ADAUTHDC3;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1457135805] SERVICE ALERT: ADAUTHDC2;AD Replication;WARNING;SOFT;2;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1457135809] SERVICE ALERT: ADAUTHDC1;CPU Load;OK;SOFT;3;CPU Load 50% (5 min average)
[1457135854] SERVICE ALERT: ADAUTHDC3;AD Replication;UNKNOWN;SOFT;2;CHECK_NRPE: Invalid packet type received from server.
[1457135918] SERVICE ALERT: ADAUTHDC3;CPU Load;OK;SOFT;2;CPU Load 50% (5 min average)
[1457135926] SERVICE ALERT: ADAUTHDC2;AD Replication;WARNING;HARD;3;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1457135969] SERVICE ALERT: ADAUTHDC3;AD Replication;UNKNOWN;HARD;3;CHECK_NRPE: Invalid packet type received from server.
[1457136225] SERVICE ALERT: EMUPRINTPS;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457136325] SERVICE ALERT: EMUPRINTPS;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1457137170] SERVICE ALERT: ADAUTHDC3;AD Replication;WARNING;HARD;3;Send-MailMessage : The remote name could not be resolved: 'meritsmtp.emich.edu'
[1457137444] SERVICE ALERT: ADRODC-DET;Service: AD Domain Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457137446] SERVICE ALERT: ADRODC-DET;Service: PXE Boot;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457137468] SERVICE ALERT: ADRODC-DET;Service: Intersite Messaging;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457137474] HOST ALERT: ADRODC-DET;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1457137477] SERVICE ALERT: ADRODC-DET;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457137483] SERVICE ALERT: ADRODC-DET;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457137507] HOST ALERT: ADRODC-DET;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1457137517] SERVICE ALERT: ADRODC-DET;Service: NetLogin;CRITICAL;SOFT;1;connect to address 198.109.113.140 and port 12489: Connection refused
[1457137517] SERVICE ALERT: ADRODC-DET;Service: Key Distribution Center;CRITICAL;SOFT;1;connect to address 198.109.113.140 and port 12489: Connection refused
[1457137519] SERVICE ALERT: ADRODC-DET;NSClient++ Version;CRITICAL;SOFT;1;connect to address 198.109.113.140 and port 12489: Connection refused
[1457137524] SERVICE ALERT: ADRODC-DET;AD Replication;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 30 seconds.
[1457137529] SERVICE ALERT: ADRODC-DET;CPU Load;CRITICAL;SOFT;1;connect to address 198.109.113.140 and port 12489: Connection refused
[1457137544] SERVICE ALERT: ADRODC-DET;Service: AD Domain Service;CRITICAL;SOFT;2;connect to address 198.109.113.140 and port 12489: Connection refused
[1457137546] SERVICE ALERT: ADRODC-DET;Service: PXE Boot;CRITICAL;SOFT;2;connect to address 198.109.113.140 and port 12489: Connection refused
[1457137547] HOST ALERT: ADRODC-DET;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1457137577] SERVICE ALERT: ADRODC-DET;Service: Intersite Messaging;OK;SOFT;2;IsmServ: Started
[1457137577] SERVICE ALERT: ADRODC-DET;Memory Usage;OK;SOFT;2;Memory usage: total:37586.86 MB - used: 1995.24 MB (5%) - free: 35591.62 MB (95%)
[1457137581] HOST ALERT: ADRODC-DET;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 8.16 ms
[1457137583] SERVICE ALERT: ADRODC-DET;C:\ Drive Space;OK;SOFT;2;c:\ - total: 1115.16 Gb - used: 268.09 Gb (24%) - free 847.06 Gb (76%)
[1457137628] SERVICE ALERT: ADRODC-DET;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1457137628] SERVICE ALERT: ADRODC-DET;Service: Key Distribution Center;OK;SOFT;2;KDC: Started
[1457137634] SERVICE ALERT: ADRODC-DET;AD Replication;OK;SOFT;2;OK: replication is up and running
[1457137636] SERVICE ALERT: ADRODC-DET;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1457137649] SERVICE ALERT: ADRODC-DET;CPU Load;OK;SOFT;2;CPU Load 14% (5 min average)
[1457137664] SERVICE ALERT: ADRODC-DET;Service: AD Domain Service;OK;SOFT;3;NTDS: Started
[1457137666] SERVICE ALERT: ADRODC-DET;Service: PXE Boot;OK;SOFT;3;WDSServer: Started
[1457138188] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;CRITICAL;SOFT;1;Black Cartridge HP CE260X is at CRITICAL level!
[1457138247] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;CRITICAL;SOFT;2;Black Cartridge HP CE260X is at CRITICAL level!
[1457138308] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;CRITICAL;HARD;3;Black Cartridge HP CE260X is at CRITICAL level!
[1457138373] SERVICE ALERT: ADRODC-DET;Service: DFS Namespace;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457138403] HOST ALERT: ADRODC-DET;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1457138417] SERVICE ALERT: ADRODC-DET;Service: MDT Monitor;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457138443] SERVICE ALERT: ADRODC-DET;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457138447] HOST ALERT: ADRODC-DET;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1457138466] SERVICE ALERT: ADRODC-DET;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457138469] SERVICE ALERT: ADRODC-DET;Service: DNS Service;CRITICAL;SOFT;1;connect to address 198.109.113.140 and port 12489: Connection refused
[1457138470] SERVICE ALERT: ADRODC-DET;Service: User Profile;CRITICAL;SOFT;1;connect to address 198.109.113.140 and port 12489: Connection refused
[1457138473] SERVICE ALERT: ADRODC-DET;Service: DFS Namespace;CRITICAL;SOFT;2;connect to address 198.109.113.140 and port 12489: Connection refused
[1457138478] HOST ALERT: ADRODC-DET;UP;SOFT;3;PING OK - Packet loss = 61%, RTA = 8.47 ms
[1457138517] SERVICE ALERT: ADRODC-DET;Service: MDT Monitor;OK;SOFT;2;MDT_Monitor: Started
[1457138543] SERVICE ALERT: ADRODC-DET;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1457138566] SERVICE ALERT: ADRODC-DET;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1457138586] SERVICE ALERT: ADRODC-DET;Service: DNS Service;OK;SOFT;2;DNS: Started
[1457138587] SERVICE ALERT: ADRODC-DET;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1457138593] SERVICE ALERT: ADRODC-DET;Service: DFS Namespace;OK;SOFT;3;Dfs: Started
[1457139289] SERVICE ALERT: ADAUTHDC1;AD Replication;OK;HARD;3;OK: replication is up and running
[1457139529] SERVICE ALERT: ADAUTHDC2;AD Replication;OK;HARD;3;OK: replication is up and running
[1457139574] SERVICE ALERT: ADAUTHDC3;AD Replication;OK;HARD;3;OK: replication is up and running
[1457143253] SERVICE ALERT: SCCM-PRIM01;C:\ Drive Space;UNKNOWN;SOFT;1;NSClient - ERROR: Invalid return from command: check_drivesize
[1457143373] SERVICE ALERT: SCCM-PRIM01;C:\ Drive Space;UNKNOWN;SOFT;2;NSClient - ERROR: Invalid return from command: check_drivesize
[1457143460] SERVICE ALERT: SCCM-PRIM01;Service: DNS Client;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1457143513] SERVICE ALERT: SCCM-PRIM01;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457143522] SERVICE ALERT: SCCM-PRIM01;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.102.80 and port 12489: No route to host
[1457143522] HOST ALERT: SCCM-PRIM01;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.80)
[1457143565] SERVICE ALERT: SCCM-PRIM01;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.102.80 and port 12489: No route to host
[1457143568] HOST ALERT: SCCM-PRIM01;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.80)
[1457143582] SERVICE ALERT: SCCM-PRIM01;Service: DNS Client;CRITICAL;SOFT;2;connect to address 164.76.102.80 and port 12489: No route to host
[1457143585] HOST ALERT: SCCM-PRIM01;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.102.80)
[1457143615] SERVICE ALERT: SCCM-PRIM01;Service: NetLogin;CRITICAL;SOFT;1;connect to address 164.76.102.80 and port 12489: Connection refused
[1457143622] HOST ALERT: SCCM-PRIM01;UP;SOFT;4;PING OK - Packet loss = 37%, RTA = 0.48 ms
[1457143627] SERVICE FLAPPING ALERT: DB01-2012;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1457143631] SERVICE ALERT: SCCM-PRIM01;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1457143682] SERVICE ALERT: SCCM-PRIM01;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1457143699] SERVICE ALERT: SCCM-PRIM01;Service: DNS Client;OK;SOFT;3;Dnscache: Started
[1457143733] SERVICE ALERT: SCCM-PRIM01;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1457144094] SERVICE ALERT: SCCM-PRIM01;C:\ Drive Space;OK;HARD;3;c:\ - total: 194.90 Gb - used: 123.32 Gb (63%) - free 71.59 Gb (37%)
[1457145105] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 93% (5 min average)
[1457145427] SERVICE ALERT: DB01-2012;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1457145547] SERVICE ALERT: DB01-2012;CPU Load;OK;SOFT;2;CPU Load 43% (5 min average)
[1457145705] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1457146833] HOST DOWNTIME ALERT: ADAUTHDC2;STOPPED; Host has exited from a period of scheduled downtime
[1457146833] HOST DOWNTIME ALERT: ADAUTHDC1;STOPPED; Host has exited from a period of scheduled downtime
[1457146833] HOST DOWNTIME ALERT: ADDFS01;STOPPED; Host has exited from a period of scheduled downtime
[1457146833] HOST DOWNTIME ALERT: ADRODC-LIV;STOPPED; Host has exited from a period of scheduled downtime
[1457146833] HOST DOWNTIME ALERT: ADCA1;STOPPED; Host has exited from a period of scheduled downtime
[1457146833] HOST DOWNTIME ALERT: NETIDAD;STOPPED; Host has exited from a period of scheduled downtime
[1457146833] HOST DOWNTIME ALERT: ADRODC-DET;STOPPED; Host has exited from a period of scheduled downtime
[1457146833] HOST DOWNTIME ALERT: ADCA2;STOPPED; Host has exited from a period of scheduled downtime
[1457146833] HOST DOWNTIME ALERT: DB01-2012;STOPPED; Host has exited from a period of scheduled downtime
[1457146833] HOST DOWNTIME ALERT: ADWU1;STOPPED; Host has exited from a period of scheduled downtime
[1457146833] HOST DOWNTIME ALERT: ADAUTHDC3;STOPPED; Host has exited from a period of scheduled downtime
[1457146833] HOST DOWNTIME ALERT: ADDC1;STOPPED; Host has exited from a period of scheduled downtime
[1457146833] HOST DOWNTIME ALERT: EMUPRINTPS;STOPPED; Host has exited from a period of scheduled downtime
[1457146833] HOST DOWNTIME ALERT: ADDC6;STOPPED; Host has exited from a period of scheduled downtime
[1457146833] HOST DOWNTIME ALERT: CASPER;STOPPED; Host has exited from a period of scheduled downtime
[1457146833] HOST DOWNTIME ALERT: CSCFILESERV;STOPPED; Host has exited from a period of scheduled downtime
[1457147505] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 90% (5 min average)
[1457148105] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1457148105] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1457150196] HOST ALERT: ADWSUSSERVERS;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.133)
[1457150230] HOST ALERT: ADWSUSSERVERS;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.30 ms