[1458016421] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1458016476] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 46% - OK!
[1458017916] SERVICE ALERT: SP2010;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458018036] SERVICE ALERT: SP2010;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1458019312] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 88% (5 min average)
[1458019912] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1458020031] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 97% (5 min average)
[1458020149] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 89% (5 min average)
[1458020750] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458020869] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1458020989] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 95% (5 min average)
[1458022295] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 84.90 Gb - used: 78.66 Gb (93%) - free 6.24 Gb (7%)
[1458022380] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458022500] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 81% (5 min average)
[1458022789] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (5 min average)
[1458023100] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1458023220] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1458023340] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 82% (5 min average)
[1458023389] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458023509] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1458023629] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1458023940] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1458024060] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 94% (5 min average)
[1458024180] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 90% (5 min average)
[1458025380] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 76% (5 min average)
[1458025980] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1458026100] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 86% (5 min average)
[1458028427] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 93% (5 min average)
[1458028500] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1458028620] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 63% (5 min average)
[1458029027] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1458029627] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 82% (5 min average)
[1458030227] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458030347] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1458030467] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 53% (5 min average)
[1458031067] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1458031187] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 93% (5 min average)
[1458031307] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 48% (5 min average)
[1458031907] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458032027] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1458032147] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 76% (5 min average)
[1458032747] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458032867] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 93% (5 min average)
[1458032987] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 94% (5 min average)
[1458033587] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 83% (5 min average)
[1458036587] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1458036707] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1458036827] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1458037427] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 82% (5 min average)
[1458038027] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458038147] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1458038267] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1458040855] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1458041196] SERVICE ALERT: mckennypk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1458041255] SERVICE ALERT: mckennypk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1458041316] SERVICE ALERT: mckennypk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1458042716] SERVICE ALERT: owen101lab1;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1458042734] HOST DOWNTIME ALERT: ADDC6;STARTED; Host has entered a period of scheduled downtime
[1458042745] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1458042776] SERVICE ALERT: owen101lab1;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! 2 of 3 Trays Still Full
[1458042905] SERVICE ALERT: ADDC6;VMWare Tools;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1458043025] SERVICE ALERT: ADDC6;VMWare Tools;OK;SOFT;2;VMTools: Started
[1458043345] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1458043715] SERVICE ALERT: mckennypk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1458043856] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1458043878] SERVICE ALERT: INTLAPP;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 63.90 Gb - used: 54.51 Gb (85%) - free 9.39 Gb (15%)
[1458043907] SERVICE ALERT: DIRECTACCESS;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458043912] SERVICE ALERT: INTLAPP-ISSMAP;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 63.90 Gb - used: 55.08 Gb (86%) - free 8.82 Gb (14%)
[1458043916] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1458043926] SERVICE ALERT: DIRECTACCESS;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458043929] HOST ALERT: BBTS;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1458043931] HOST ALERT: BBTS_DB;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1458043935] SERVICE ALERT: MAGICAPPTEST;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458043936] SERVICE ALERT: MAGICDBPROD;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458043950] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;CRITICAL;SOFT;1;c:\ - total: 84.90 Gb - used: 80.16 Gb (94%) - free 4.74 Gb (6%)
[1458043965] HOST ALERT: MAGICAPPTEST;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1458043966] HOST ALERT: MAGICDBPROD;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1458043973] SERVICE ALERT: DIRECTACCESS;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458043976] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1458043979] HOST ALERT: BBTS_DB;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1458043979] HOST ALERT: BBTS;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1458043982] SERVICE ALERT: MAGICAPPTEST;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458043983] SERVICE ALERT: MAGICDBPROD;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458043998] SERVICE ALERT: INTLAPP;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 63.90 Gb - used: 54.51 Gb (85%) - free 9.39 Gb (15%)
[1458044012] HOST ALERT: MAGICAPPTEST;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1458044013] HOST ALERT: MAGICDBPROD;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1458044020] SERVICE ALERT: DIRECTACCESS;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458044026] HOST ALERT: BBTS_DB;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1458044026] HOST ALERT: BBTS;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1458044027] SERVICE ALERT: DIRECTACCESS;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458044029] SERVICE ALERT: MAGICAPPTEST;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458044030] SERVICE ALERT: MAGICDBPROD;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458044032] SERVICE ALERT: INTLAPP-ISSMAP;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 63.90 Gb - used: 55.08 Gb (86%) - free 8.82 Gb (14%)
[1458044046] SERVICE ALERT: DIRECTACCESS;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458044047] SERVICE ALERT: MAGICDBPROD;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458044054] SERVICE ALERT: INTLTESTDB;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 84.22 Gb - used: 70.77 Gb (84%) - free 13.45 Gb (16%)
[1458044055] SERVICE ALERT: MAGICAPPTEST;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458044056] SERVICE ALERT: MAGICDBPROD;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458044059] HOST ALERT: MAGICAPPTEST;DOWN;SOFT;3;(Host check timed out after 30.00 seconds)
[1458044060] HOST ALERT: MAGICDBPROD;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1458044065] HOST ALERT: BBTS_DB;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1458044067] SERVICE ALERT: DIRECTACCESS;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458044070] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;CRITICAL;SOFT;2;c:\ - total: 84.90 Gb - used: 80.16 Gb (94%) - free 4.74 Gb (6%)
[1458044073] HOST ALERT: BBTS;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1458044076] SERVICE ALERT: MAGICAPPTEST;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458044077] SERVICE ALERT: MAGICDBPROD;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458044093] SERVICE ALERT: DIRECTACCESS;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458044102] SERVICE ALERT: MAGICAPPTEST;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458044103] SERVICE ALERT: MAGICDBPROD;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458044106] HOST ALERT: MAGICAPPTEST;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1458044107] HOST ALERT: MAGICDBPROD;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1458044114] SERVICE ALERT: DIRECTACCESS;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458044118] SERVICE ALERT: INTLAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 63.90 Gb - used: 54.51 Gb (85%) - free 9.39 Gb (15%)
[1458044120] HOST ALERT: BBTS;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1458044120] HOST ALERT: BBTS_DB;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1458044123] SERVICE ALERT: MAGICAPPTEST;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458044124] SERVICE ALERT: MAGICDBPROD;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458044140] SERVICE ALERT: DIRECTACCESS;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458044147] SERVICE ALERT: DIRECTACCESS;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458044149] SERVICE ALERT: MAGICAPPTEST;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458044150] SERVICE ALERT: MAGICDBPROD;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458044152] SERVICE ALERT: INTLAPP-ISSMAP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 63.90 Gb - used: 55.08 Gb (86%) - free 8.82 Gb (14%)
[1458044153] HOST ALERT: MAGICAPPTEST;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1458044154] HOST ALERT: MAGICDBPROD;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1458044161] SERVICE ALERT: DIRECTACCESS;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458044166] SERVICE ALERT: DIRECTACCESS;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458044167] SERVICE ALERT: MAGICDBPROD;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458044167] HOST ALERT: BBTS;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1458044167] HOST ALERT: BBTS_DB;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1458044170] SERVICE ALERT: MAGICAPPTEST;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458044171] SERVICE ALERT: MAGICDBPROD;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458044174] SERVICE ALERT: INTLTESTDB;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 84.22 Gb - used: 70.77 Gb (84%) - free 13.45 Gb (16%)
[1458044175] SERVICE ALERT: MAGICAPPTEST;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458044175] SERVICE ALERT: MAGICAPPTEST;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458044176] SERVICE ALERT: MAGICDBPROD;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458044187] SERVICE ALERT: DIRECTACCESS;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458044190] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 84.90 Gb - used: 80.16 Gb (94%) - free 4.74 Gb (6%)
[1458044196] SERVICE ALERT: MAGICAPPTEST;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458044197] HOST ALERT: MAGICDBPROD;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1458044197] SERVICE ALERT: MAGICDBPROD;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458044200] HOST ALERT: MAGICAPPTEST;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1458044208] SERVICE ALERT: DIRECTACCESS;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458044213] SERVICE ALERT: DIRECTACCESS;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458044214] HOST ALERT: BBTS;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1458044214] HOST ALERT: BBTS_DB;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1458044217] SERVICE ALERT: MAGICAPPTEST;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458044218] SERVICE ALERT: MAGICDBPROD;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458044222] SERVICE ALERT: MAGICAPPTEST;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458044223] SERVICE ALERT: MAGICDBPROD;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458044227] HOST ALERT: MAGICDBPROD;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1458044234] SERVICE ALERT: DIRECTACCESS;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458044243] SERVICE ALERT: MAGICAPPTEST;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458044244] SERVICE ALERT: MAGICDBPROD;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458044244] SERVICE ALERT: LANSWEEPER;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1458044247] HOST ALERT: MAGICAPPTEST;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1458044255] SERVICE ALERT: DIRECTACCESS;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458044260] SERVICE ALERT: DIRECTACCESS;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458044261] HOST ALERT: BBTS;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1458044261] HOST ALERT: BBTS_DB;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1458044264] SERVICE ALERT: MAGICAPPTEST;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458044265] SERVICE ALERT: MAGICDBPROD;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458044269] SERVICE ALERT: MAGICAPPTEST;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458044270] SERVICE ALERT: MAGICDBPROD;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458044274] HOST ALERT: MAGICDBPROD;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1458044281] SERVICE ALERT: DIRECTACCESS;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458044294] SERVICE ALERT: INTLTESTDB;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 84.22 Gb - used: 70.77 Gb (84%) - free 13.45 Gb (16%)
[1458044311] SERVICE ALERT: MAGICAPPTEST;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458044312] SERVICE ALERT: MAGICDBPROD;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458044312] SERVICE ALERT: MAGICAPPTEST;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458044316] SERVICE ALERT: MAGICAPPTEST;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458044317] SERVICE ALERT: MAGICDBPROD;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458044317] HOST ALERT: BBTS;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1458044319] HOST ALERT: BBTS_DB;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1458044328] SERVICE ALERT: DIRECTACCESS;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458044337] SERVICE ALERT: MAGICAPPTEST;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458044338] SERVICE ALERT: MAGICDBPROD;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458044341] HOST ALERT: MAGICAPPTEST;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1458044342] HOST ALERT: MAGICDBPROD;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1458044349] SERVICE ALERT: DIRECTACCESS;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458044349] SERVICE ALERT: DIRECTACCESS;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458044354] SERVICE ALERT: DIRECTACCESS;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458044354] SERVICE ALERT: MAGICAPPTEST;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458044355] HOST ALERT: BBTS;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1458044355] HOST ALERT: BBTS_DB;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1458044358] SERVICE ALERT: MAGICAPPTEST;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458044359] SERVICE ALERT: MAGICDBPROD;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458044359] SERVICE ALERT: MAGICDBPROD;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458044363] SERVICE ALERT: MAGICAPPTEST;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458044364] SERVICE ALERT: MAGICDBPROD;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458044364] SERVICE ALERT: LANSWEEPER;CPU Load;OK;SOFT;2;CPU Load 85% (5 min average)
[1458044375] SERVICE ALERT: DIRECTACCESS;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458044383] SERVICE ALERT: MAGICDBPROD;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458044384] SERVICE ALERT: MAGICAPPTEST;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458044384] HOST ALERT: MAGICAPPTEST;DOWN;SOFT;3;(Host check timed out after 30.00 seconds)
[1458044385] SERVICE ALERT: MAGICDBPROD;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458044389] HOST ALERT: MAGICDBPROD;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1458044396] SERVICE ALERT: DIRECTACCESS;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458044401] SERVICE ALERT: DIRECTACCESS;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458044401] SERVICE ALERT: MAGICAPPTEST;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458044401] SERVICE ALERT: DIRECTACCESS;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458044402] HOST ALERT: BBTS_DB;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1458044402] HOST ALERT: BBTS;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1458044405] SERVICE ALERT: MAGICAPPTEST;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458044406] SERVICE ALERT: MAGICDBPROD;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458044413] SERVICE ALERT: MAGICDBPROD;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458044431] SERVICE ALERT: MAGICAPPTEST;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458044431] HOST ALERT: MAGICAPPTEST;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1458044432] SERVICE ALERT: MAGICDBPROD;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458044432] SERVICE ALERT: MAGICAPPTEST;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458044436] HOST ALERT: MAGICDBPROD;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1458044443] SERVICE ALERT: DIRECTACCESS;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458044448] SERVICE ALERT: DIRECTACCESS;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458044449] HOST ALERT: BBTS;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1458044449] HOST ALERT: BBTS_DB;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1458044452] SERVICE ALERT: MAGICAPPTEST;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458044453] SERVICE ALERT: MAGICDBPROD;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458044457] SERVICE ALERT: MAGICAPPTEST;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458044458] SERVICE ALERT: MAGICDBPROD;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458044462] HOST ALERT: MAGICAPPTEST;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1458044469] SERVICE ALERT: DIRECTACCESS;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458044474] SERVICE ALERT: MAGICAPPTEST;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458044478] SERVICE ALERT: MAGICAPPTEST;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458044479] SERVICE ALERT: MAGICDBPROD;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458044479] SERVICE ALERT: MAGICDBPROD;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458044483] HOST ALERT: MAGICDBPROD;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1458044495] SERVICE ALERT: DIRECTACCESS;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458044495] SERVICE ALERT: DIRECTACCESS;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458044503] SERVICE ALERT: MAGICDBPROD;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458044504] SERVICE ALERT: MAGICAPPTEST;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458044504] HOST ALERT: MAGICAPPTEST;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1458044505] SERVICE ALERT: MAGICDBPROD;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458044516] SERVICE ALERT: DIRECTACCESS;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458044521] SERVICE ALERT: MAGICAPPTEST;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458044521] SERVICE ALERT: DIRECTACCESS;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458044525] SERVICE ALERT: MAGICAPPTEST;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458044526] SERVICE ALERT: MAGICDBPROD;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458044532] HOST ALERT: BBTS;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1458044532] HOST ALERT: BBTS_DB;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1458044533] HOST ALERT: MAGICDBPROD;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1458044551] SERVICE ALERT: MAGICAPPTEST;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458044551] HOST ALERT: MAGICAPPTEST;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1458044552] SERVICE ALERT: MAGICAPPTEST;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458044552] SERVICE ALERT: MAGICDBPROD;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458044563] SERVICE ALERT: DIRECTACCESS;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458044572] SERVICE ALERT: MAGICAPPTEST;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458044573] SERVICE ALERT: MAGICDBPROD;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458044576] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1458044579] HOST ALERT: BBTS;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1458044579] HOST ALERT: BBTS_DB;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1458044582] HOST ALERT: MAGICAPPTEST;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1458044582] HOST ALERT: MAGICDBPROD;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1458044589] SERVICE ALERT: DIRECTACCESS;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458044598] SERVICE ALERT: MAGICAPPTEST;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458044599] SERVICE ALERT: MAGICDBPROD;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458044599] SERVICE ALERT: MAGICDBPROD;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458044615] SERVICE ALERT: DIRECTACCESS;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458044626] HOST ALERT: BBTS;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1458044626] HOST ALERT: BBTS_DB;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1458044628] HOST ALERT: MAGICAPPTEST;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1458044629] HOST ALERT: MAGICDBPROD;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1458044636] SERVICE ALERT: DIRECTACCESS;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458044641] SERVICE ALERT: DIRECTACCESS;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458044645] SERVICE ALERT: MAGICAPPTEST;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458044646] SERVICE ALERT: MAGICDBPROD;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458044665] HOST ALERT: BBTS_DB;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1458044673] HOST ALERT: BBTS;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1458044675] HOST ALERT: MAGICAPPTEST;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1458044676] HOST ALERT: MAGICDBPROD;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1458044683] SERVICE ALERT: DIRECTACCESS;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458044692] SERVICE ALERT: MAGICAPPTEST;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458044693] SERVICE ALERT: MAGICDBPROD;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458044720] HOST ALERT: BBTS;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1458044720] HOST ALERT: BBTS_DB;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1458044723] HOST ALERT: MAGICDBPROD;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1458044735] SERVICE ALERT: DIRECTACCESS;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458046443] SERVICE ALERT: C4SAS-SE;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1458046500] SERVICE ALERT: INTLTESTDB;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1458046563] SERVICE ALERT: C4SAS-SE;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1458046611] SERVICE ALERT: VCENTERSSO;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1458046620] SERVICE ALERT: INTLTESTDB;CPU Load;OK;SOFT;2;CPU Load 18% (5 min average)
[1458046730] SERVICE ALERT: VCENTERSSO;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1458047290] SERVICE ALERT: hallethirdsouth;Device Status;CRITICAL;HARD;3;(Return code of 13 is out of bounds)
[1458047294] SERVICE ALERT: INTLTESTDB;C:\ Drive Space;OK;HARD;3;c:\ - total: 84.22 Gb - used: 70.77 Gb (84%) - free 13.44 Gb (16%)
[1458047718] SERVICE ALERT: INTLAPP;C:\ Drive Space;OK;HARD;3;c:\ - total: 63.90 Gb - used: 54.51 Gb (85%) - free 9.39 Gb (15%)
[1458047752] SERVICE ALERT: INTLAPP-ISSMAP;C:\ Drive Space;OK;HARD;3;c:\ - total: 63.90 Gb - used: 55.08 Gb (86%) - free 8.82 Gb (14%)
[1458047790] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 84.90 Gb - used: 79.75 Gb (94%) - free 5.15 Gb (6%)
[1458047896] SERVICE ALERT: hallethirdsouth;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.107.213. Is the printer off?  Is the network down?
[1458048547] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1458048667] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1458048787] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 95% (5 min average)
[1458049387] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1458050587] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 64% (5 min average)
[1458050790] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;OK;HARD;3;c:\ - total: 84.90 Gb - used: 75.46 Gb (89%) - free 9.44 Gb (11%)
[1458051787] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1458051907] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 86% (5 min average)
[1458052975] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1458053035] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1458053095] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1458053106] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 94% (5 min average)
[1458053227] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1458053346] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1458053346] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1458056695] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1458056768] SERVICE ALERT: owen101lab3;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1458056828] SERVICE ALERT: owen101lab3;Device Status;OK;SOFT;2;Clearing paper path
[1458057295] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1458057355] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1458057415] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1458057428] SERVICE ALERT: owen101lab3;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1458057488] SERVICE ALERT: owen101lab3;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1458057523] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1458057549] SERVICE ALERT: owen101lab3;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1458057553] HOST ALERT: prayharroldpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1458057583] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1458057598] HOST ALERT: prayharroldpk;UP;SOFT;2;PING OK - Packet loss = 68%, RTA = 0.72 ms
[1458057637] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;OK;SOFT;3;Black Print Cartridge HP C8543X is at 5% - OK!
[1458057661] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1458057691] HOST ALERT: hallefirst2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1458057721] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1458057731] SERVICE ALERT: cava2;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1458057751] HOST ALERT: hallefirst2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1458057780] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1458057781] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1458057810] HOST ALERT: hallefirst2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1458057840] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1458057851] SERVICE ALERT: cava2;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1458057870] HOST ALERT: hallefirst2;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1458057900] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1458057930] HOST ALERT: hallefirst2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1458058020] HOST ALERT: hallefirst2;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1458058087] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1458058110] HOST ALERT: hallefirst2;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1458058147] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1458058148] SERVICE ALERT: owen101lab3;Device Status;OK;HARD;3;Warming up
[1458058177] HOST ALERT: hallefirst2;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1458058208] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1458058238] HOST ALERT: hallefirst2;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1458058327] HOST ALERT: hallefirst2;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1458058748] SERVICE ALERT: owen101lab3;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1458058801] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 29% - OK!
[1458058805] HOST ALERT: hallefirst2;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.47 ms
[1458058808] SERVICE ALERT: owen101lab3;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1458058855] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1458058868] SERVICE ALERT: owen101lab3;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1458058914] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1458058974] SERVICE ALERT: hallefirst1;Device Status;OK;SOFT;3;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1458058976] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Ready
[1458059095] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1458059176] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1458059236] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1458059295] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1458059468] SERVICE ALERT: owen101lab3;Device Status;OK;HARD;3;41.3 UNEXPECTED SIZE IN TRAY 4 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   LOAD TRAY 4 PLAIN LETTER To use another 74 72 61 79 20 70 72 65 73 73 20 1E
[1458059468] SERVICE FLAPPING ALERT: owen101lab3;Device Status;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1458059575] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1458059635] SERVICE ALERT: hallefirst2;Device Status;OK;SOFT;2;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1458060835] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1458060895] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1458060955] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1458061556] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Ready
[1458061556] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1458062815] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1458062815] SERVICE FLAPPING ALERT: markjeffersonpk;Paper Tray Status;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1458062946] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1458063084] SERVICE ALERT: hallegroundcolor;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 1 Trays Empty
[1458063144] SERVICE ALERT: hallegroundcolor;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 1 Trays Empty
[1458063204] SERVICE ALERT: hallegroundcolor;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 1 Trays Empty
[1458063415] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1458063474] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! 1 of 3 Trays Still Full
[1458064746] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 80% (5 min average)
[1458065275] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1458065335] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1458065346] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458065395] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1458065466] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1458065586] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1458065915] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1458065975] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1458066035] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1458067386] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 90% (5 min average)
[1458067986] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1458067986] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1458068155] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1458068215] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1458068275] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1458068421] SERVICE ALERT: EMUFILES;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1458068431] SERVICE ALERT: ADCA1;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1458068541] SERVICE ALERT: EMUFILES;CPU Load;OK;SOFT;2;CPU Load 12% (5 min average)
[1458068551] SERVICE ALERT: ADCA1;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1458068995] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1458069035] SERVICE ALERT: studentcenterpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1458069191] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;UNKNOWN;SOFT;1;Failed to process command : EvtQuery failed: 1717: The interface is unknown.
[1458069311] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;OK;SOFT;2;OK: No entries found
[1458069786] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 82% (5 min average)
[1458070075] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Clearing paper path
[1458070386] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1458070506] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1458070626] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1458070795] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1458070801] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1458070841] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1458070855] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1458070861] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 29% - OK!
[1458070867] SERVICE FLAPPING ALERT: owen101lab3;Device Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1458070871] HOST ALERT: prayharroldpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1458070896] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 5% - OK!
[1458070900] HOST ALERT: prayharroldpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.67 ms
[1458070915] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1458071275] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1458071335] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1458071395] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1458071515] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1458074321] SERVICE ALERT: halleground1;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1458074351] HOST ALERT: halleground1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1458074376] SERVICE ALERT: halleground1;Device Status;OK;SOFT;2;Clearing paper path
[1458074380] HOST ALERT: halleground1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.56 ms
[1458074572] SERVICE ALERT: halleground1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1458074602] HOST ALERT: halleground1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1458074632] SERVICE ALERT: halleground1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1458074662] HOST ALERT: halleground1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1458074681] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1458074692] SERVICE ALERT: halleground1;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1458074711] HOST ALERT: halleground1;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1458074741] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1458074771] HOST ALERT: halleground1;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1458074801] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1458074831] HOST ALERT: halleground1;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1458074921] HOST ALERT: halleground1;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1458074981] SERVICE ALERT: halleground1;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1458075011] HOST ALERT: halleground1;DOWN;SOFT;7;(Host check timed out after 30.00 seconds)
[1458075041] SERVICE ALERT: halleground1;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1458075071] HOST ALERT: halleground1;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1458075101] SERVICE ALERT: halleground1;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1458075131] HOST ALERT: halleground1;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1458075203] SERVICE ALERT: hallegroundcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1458075221] HOST ALERT: halleground1;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1458075714] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1458075775] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1458075835] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1458076095] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1458076637] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1458076667] HOST ALERT: hallefirstcolor;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1458076680] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1458076680] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1458076696] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1458076697] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1458076710] HOST ALERT: hallefirstcolor;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1458076740] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1458076740] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1458076756] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1458076756] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;UNKNOWN;HARD;3;UNKNOWN - No match found for 'Magenta Cartridge HP CE263A'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1458076760] HOST ALERT: hallefirstcolor;UP;SOFT;3;PING OK - Packet loss = 76%, RTA = 0.44 ms
[1458076794] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;OK;SOFT;3;
[1458076794] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;OK;SOFT;3;
[1458076815] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1458077034] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1458077351] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;OK;HARD;3;Magenta Cartridge HP CE263A is at 64% - OK!
[1458077416] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1458077635] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1458077695] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1458077755] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;SOFT;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1458078426] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 43% (5 min average)
[1458079626] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1458079746] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1458079866] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 97% (5 min average)
[1458080213] SERVICE ALERT: prayharroldpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1458080273] SERVICE ALERT: prayharroldpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1458080333] SERVICE ALERT: prayharroldpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1458080939] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1458080969] HOST ALERT: prayharroldpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1458081033] HOST ALERT: prayharroldpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.62 ms
[1458081533] SERVICE ALERT: prayharroldpk;Device Status;OK;HARD;3;ORDER CARTRIDGE LESS THAN 1200 PAGES 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1458082194] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1458085866] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 85% (5 min average)
[1458086466] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458086586] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1458086706] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1458087955] SERVICE FLAPPING ALERT: markjeffersonpk;Paper Tray Status;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1458089704] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 73% (5 min average)
[1458090304] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1458090424] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1458090544] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1458090854] SERVICE ALERT: strongpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1458090914] SERVICE ALERT: strongpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1458090973] SERVICE ALERT: strongpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1458091446] SERVICE ALERT: EVENTS-2012;CPU Load;WARNING;SOFT;1;CPU Load 91% (5 min average)
[1458091566] SERVICE ALERT: EVENTS-2012;CPU Load;OK;SOFT;2;CPU Load 87% (5 min average)
[1458096478] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10032;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1458096478] HOST ALERT: ADWSUSSERVERS;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.133)
[1458096481] SERVICE ALERT: ADWSUSSERVERS;E:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.102.133 and port 12489: No route to host
[1458096521] HOST ALERT: ADWSUSSERVERS;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1458096590] SERVICE ALERT: ADWSUSSERVERS;E:\ Drive Space;OK;SOFT;2;e:\ - total: 99.87 Gb - used: 59.45 Gb (60%) - free 40.43 Gb (40%)
[1458096594] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10032;OK;SOFT;2;Eventlog check ok
[1458097346] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1458097347] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1458097377] HOST ALERT: hallefirst1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1458097380] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1458097397] HOST ALERT: hallefirst1;UP;SOFT;2;PING OK - Packet loss = 70%, RTA = 0.76 ms
[1458097402] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 33% - OK!
[1458097441] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1458097466] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 93% (5 min average)
[1458097471] HOST ALERT: hallefirst1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1458097496] SERVICE ALERT: hallefirst1;Device Status;OK;SOFT;3;Ready
[1458097500] HOST ALERT: hallefirst1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.66 ms
[1458097500] HOST FLAPPING ALERT: hallefirst1;STARTED; Host appears to have started flapping (23.2% change > 20.0% threshold)
[1458097569] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1458097586] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 95% (5 min average)
[1458097629] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1458097689] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1458098007] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1458098067] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1458098099] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1458098127] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1458098159] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1458098186] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 92% (5 min average)
[1458098219] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1458098223] HOST FLAPPING ALERT: hallefirst1;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1458098786] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 80% (5 min average)
[1458100025] SERVICE ALERT: hallefirst3;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1458100073] SERVICE ALERT: hallefirst3;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1458100085] SERVICE ALERT: hallefirst3;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1458100133] SERVICE ALERT: hallefirst3;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1458100145] SERVICE ALERT: hallefirst3;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1458100192] SERVICE ALERT: hallefirst3;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1458100610] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1458100669] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1458100729] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?