[1482123930] SERVICE ALERT: epprogsql2;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 249.66 Gb - used: 227.78 Gb (91%) - free 21.87 Gb (9%)
[1482124013] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1482124050] SERVICE ALERT: epprogsql2;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 249.66 Gb - used: 228.45 Gb (92%) - free 21.21 Gb (8%)
[1482124143] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1482124170] SERVICE ALERT: epprogsql2;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 249.66 Gb - used: 229.11 Gb (92%) - free 20.55 Gb (8%)
[1482124263] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1482124383] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1482124613] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482124673] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482124733] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482124770] SERVICE ALERT: epprogsql2;C:\ Drive Space;OK;HARD;3;c:\ - total: 249.66 Gb - used: 221.84 Gb (89%) - free 27.82 Gb (11%)
[1482124770] SERVICE FLAPPING ALERT: epprogsql2;C:\ Drive Space;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1482125333] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1482125583] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 31% (15 min average)
[1482125933] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482125994] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482126053] SERVICE ALERT: rackhampk;Device Status;OK;SOFT;3;Open tray indicated.
[1482126183] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1482126303] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1482126423] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1482126653] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482126713] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482126773] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482127023] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1482127373] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1482127623] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1482127743] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1482127863] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1482127973] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482128006] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1482128033] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482128093] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482128124] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1482128724] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1482128844] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1482128963] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1482129000] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1482129293] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1482129401] SERVICE FLAPPING ALERT: halleground3;Device Status;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1482129563] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1482129893] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482129953] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482130013] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482130613] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1482130763] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1482130883] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;2;Warning:  idm1  are more than 120 seconds behind!
[1482131003] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1482131213] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482131273] SERVICE ALERT: rackhampk;Device Status;OK;SOFT;2;Open tray indicated.
[1482131463] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1482131603] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1482131603] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1482131874] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482131933] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482131993] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482132063] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1482132593] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1482132803] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1482133193] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482133253] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482133313] SERVICE ALERT: rackhampk;Device Status;OK;SOFT;3;Open tray indicated.
[1482133403] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1482133913] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482133974] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482134033] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482134633] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1482135063] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 88% (15 min average)
[1482135233] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482135293] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482135353] SERVICE ALERT: rackhampk;Device Status;OK;SOFT;3;Open tray indicated.
[1482135663] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1482135783] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1482135803] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1482135903] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1482135953] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482136014] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482136073] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482136170] SERVICE FLAPPING ALERT: epprogsql2;C:\ Drive Space;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1482136403] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1482136673] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1482137003] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1482137123] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1482137243] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1482137273] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482137333] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482137393] SERVICE ALERT: rackhampk;Device Status;OK;SOFT;3;Open tray indicated.
[1482137543] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 1 Trays Empty
[1482137603] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 1 Trays Empty
[1482137664] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 1 Trays Empty
[1482137843] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1482137993] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482138053] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482138114] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482138443] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1482138713] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1482138903] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1482139313] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482139373] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482139433] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482139803] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (29.9% change >= 20.0% threshold)
[1482139803] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (55.5% change >= 20.0% threshold)
[1482139803] SERVICE FLAPPING ALERT: rackhampk;Device Status;STARTED; Service appears to have started flapping (74.8% change >= 20.0% threshold)
[1482139803] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1482139803] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1482139803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1482139803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1482139803] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1482139803] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1482139803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1482139803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1482139803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1482139803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1482139803] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1482139803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1482139803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1482139803] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1482139803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1482139803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1482139803] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1482140633] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1482141233] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482141293] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482141353] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482141603] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1482141603] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (47.5% change >= 20.0% threshold)
[1482141603] SERVICE FLAPPING ALERT: rackhampk;Device Status;STARTED; Service appears to have started flapping (70.8% change >= 20.0% threshold)
[1482141603] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1482141603] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1482141603] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1482141603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1482141603] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1482141603] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1482141603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1482141603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1482141603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1482141603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1482141603] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1482141603] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1482141603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1482141603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1482141603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1482141603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1482141603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1482141903] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1482141953] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1482142023] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1482142143] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 94% (15 min average)
[1482142554] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482142613] SERVICE ALERT: rackhampk;Device Status;OK;SOFT;2;Open tray indicated.
[1482143213] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482143273] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482143333] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482143933] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1482143943] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1482144063] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1482144183] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 90% (15 min average)
[1482144533] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482144593] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482144654] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482144783] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1482144903] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1482145023] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1482145853] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1482146453] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482146513] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482146573] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482147000] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service has exited from a period of scheduled downtime
[1482147006] HOST ALERT: crossroadspk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1482147096] HOST ALERT: crossroadspk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1482147173] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1482147186] HOST ALERT: crossroadspk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1482147276] HOST ALERT: crossroadspk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1482147299] SERVICE ALERT: crossroadspk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1482147329] HOST ALERT: crossroadspk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1482147359] SERVICE ALERT: crossroadspk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1482147359] SERVICE ALERT: crossroadspk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1482147389] HOST ALERT: crossroadspk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1482147406] SERVICE ALERT: crossroadspk;Black Cartridge HP CC364X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1482147419] SERVICE ALERT: crossroadspk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1482147419] SERVICE ALERT: crossroadspk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1482147436] HOST ALERT: crossroadspk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1482147466] SERVICE ALERT: crossroadspk;Black Cartridge HP CC364X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1482147479] SERVICE ALERT: crossroadspk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1482147496] HOST ALERT: crossroadspk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1482147526] SERVICE ALERT: crossroadspk;Black Cartridge HP CC364X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1482147556] HOST ALERT: crossroadspk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1482147646] HOST ALERT: crossroadspk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1482147773] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482147833] SERVICE ALERT: rackhampk;Device Status;OK;SOFT;2;Open tray indicated.
[1482148433] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482148493] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482148553] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482149141] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1085 found since last run.
[1482149153] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1482149183] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1169 found since last run.
[1482149261] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482149303] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482149753] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482149813] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482149843] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1482149861] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1555 found since last run.
[1482149873] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482149903] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1573 found since last run.
[1482149981] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482150023] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482150581] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1726 found since last run.
[1482150623] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1746 found since last run.
[1482150701] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482150743] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482151023] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 62% (15 min average)
[1482151073] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1482151301] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1747 found since last run.
[1482151343] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1788 found since last run.
[1482151421] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482151463] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482151623] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1482151673] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482151733] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482151743] SERVICE ALERT: prayharroldpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482151743] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1482151793] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482151803] SERVICE ALERT: prayharroldpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482151863] SERVICE ALERT: prayharroldpk;Device Status;OK;SOFT;3;Clearing paper path
[1482151863] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1482152021] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1737 found since last run.
[1482152063] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1811 found since last run.
[1482152142] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482152183] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482152393] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1482152741] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1069 found since last run.
[1482152860] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482153258] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1482153317] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1482153376] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1482153593] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482153614] SERVICE ALERT: owen101lab3;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1482153654] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482153661] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1482153664] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1482153674] SERVICE ALERT: owen101lab3;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1482153713] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482153721] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1482153734] SERVICE ALERT: owen101lab3;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1482153749] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1482153781] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1482153809] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1482153836] SERVICE ALERT: owen101lab2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1482153859] SERVICE ALERT: owen101lab2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1482153863] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1482153869] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1482153896] SERVICE ALERT: owen101lab2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1482153919] SERVICE ALERT: owen101lab2;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1482153956] SERVICE ALERT: owen101lab2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1482153979] SERVICE ALERT: owen101lab2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1482153999] SERVICE ALERT: owen101lab1;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1482154006] HOST ALERT: owen101lab3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1482154023] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1482154029] HOST ALERT: owen101lab1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1482154078] HOST ALERT: owen101lab3;UP;SOFT;2;PING OK - Packet loss = 61%, RTA = 0.49 ms
[1482154119] HOST ALERT: owen101lab1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1482154169] SERVICE ALERT: owen101lab1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1482154199] HOST ALERT: owen101lab1;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1482154229] SERVICE ALERT: owen101lab1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1482154259] HOST ALERT: owen101lab1;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1482154263] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 61% (15 min average)
[1482154289] SERVICE ALERT: owen101lab1;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1482154313] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1482154313] HOST ALERT: owen101lab2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1482154319] HOST ALERT: owen101lab1;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1482154328] SERVICE ALERT: owen101lab3;Device Status;OK;HARD;3;Ready
[1482154375] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 46% - OK!
[1482154403] HOST ALERT: owen101lab2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1482154409] HOST ALERT: owen101lab1;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1482154463] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 61% - OK!
[1482154467] HOST ALERT: owen101lab2;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.53 ms
[1482154489] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1482154499] HOST ALERT: owen101lab1;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1482154529] SERVICE ALERT: strongpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.130.98. Is the printer off?  Is the network down?
[1482154549] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1482154551] SERVICE ALERT: owen101lab2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1482154570] SERVICE ALERT: owen101lab3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1482154573] SERVICE ALERT: owen101lab2;Device Status;OK;HARD;3;Ready
[1482154579] HOST ALERT: owen101lab1;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1482154583] SERVICE ALERT: strongpk;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 57% - OK!
[1482154609] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1482154629] HOST ALERT: owen101lab1;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1482154719] HOST ALERT: owen101lab1;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1482154913] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482154973] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482155033] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482155059] HOST ALERT: crossroadspk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.39 ms
[1482155156] SERVICE ALERT: owen101lab2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1482155176] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1482155179] SERVICE ALERT: owen101lab2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1482155213] SERVICE ALERT: crossroadspk;Device Status;OK;HARD;3;ORDER CARTRIDGE LESS THAN 600 PAGES 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1482155216] SERVICE ALERT: owen101lab2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1482155236] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1482155239] SERVICE ALERT: owen101lab2;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1482155260] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5714 found since last run. oneLevelSearchOps OK. searchOps is too high! 5791 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482155274] SERVICE ALERT: crossroadspk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1482155276] SERVICE ALERT: owen101lab2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1482155296] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1482155299] SERVICE ALERT: owen101lab2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1482155319] SERVICE ALERT: crossroadspk;Black Cartridge HP CC364X;OK;HARD;3;Black Cartridge HP CC364X is at 3% - OK!
[1482155380] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482155463] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1482155463] SERVICE ALERT: prayharroldpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482155523] SERVICE ALERT: prayharroldpk;Device Status;OK;SOFT;2;50 72 65 73 73 20 17 20 74 6F 20 73 65 65 20 73  74 65 70  by step information
[1482155534] SERVICE ALERT: owen101lab3;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1482155570] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1482155581] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1482155583] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1482155594] SERVICE ALERT: owen101lab3;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1482155641] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1482155654] SERVICE ALERT: owen101lab3;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1482155669] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1482155701] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1482155703] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 94% (15 min average)
[1482155729] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1482155783] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5064 found since last run. oneLevelSearchOps OK. searchOps is too high! 5541 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482155789] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1482155903] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482155980] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6102 found since last run. oneLevelSearchOps OK. searchOps is too high! 6191 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482156100] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482156233] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1482156303] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1482156423] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 88% (15 min average)
[1482156503] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7266 found since last run. oneLevelSearchOps OK. searchOps is too high! 7674 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482156623] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482156834] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482156893] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482156953] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482157023] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1482157143] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 77% (15 min average)
[1482157553] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1482157803] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1482157803] SERVICE FLAPPING ALERT: rackhampk;Device Status;STARTED; Service appears to have started flapping (76.0% change >= 20.0% threshold)
[1482157803] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1482157803] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1482157803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1482157803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1482157803] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1482157803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1482157803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1482157803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1482157803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1482157803] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1482157803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1482157803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1482157803] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1482157803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1482157803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1482157803] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1482158184] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1482158184] SERVICE FLAPPING ALERT: rackhampk;Device Status;STARTED; Service appears to have started flapping (74.0% change >= 20.0% threshold)
[1482158184] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1482158184] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1482158184] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1482158184] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1482158184] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1482158184] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1482158184] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1482158184] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1482158184] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1482158184] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1482158184] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1482158184] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1482158184] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1482158184] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1482158184] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1482158184] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1482158343] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1482158423] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6464 found since last run. oneLevelSearchOps OK. searchOps is too high! 6872 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482158463] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1482158543] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482158571] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1482158583] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1482158630] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1482158691] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1482158753] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482158813] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482158873] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482159183] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 90% (15 min average)
[1482159474] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;TRAY 3 EMPTY PLAIN LETTER   Ready
[1482159700] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6067 found since last run. oneLevelSearchOps OK. searchOps is too high! 6301 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482159783] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1482159820] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482159903] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1482160023] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1482160943] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5192 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482161004] SERVICE ALERT: hallesouthcommons;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1482161063] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482161063] SERVICE ALERT: hallesouthcommons;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1482161124] SERVICE ALERT: hallesouthcommons;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1482161663] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7306 found since last run. oneLevelSearchOps OK. searchOps is too high! 7667 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482161751] SERVICE ALERT: rooseveltpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1482161783] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482161812] SERVICE ALERT: rooseveltpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1482161872] SERVICE ALERT: rooseveltpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1482162220] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6178 found since last run. oneLevelSearchOps OK. searchOps is too high! 6435 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482162340] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482162383] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6636 found since last run. oneLevelSearchOps OK. searchOps is too high! 7118 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482162503] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482162940] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7304 found since last run. oneLevelSearchOps OK. searchOps is too high! 7553 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482163060] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482163103] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7201 found since last run. oneLevelSearchOps OK. searchOps is too high! 7668 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482163223] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482163660] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5273 found since last run. oneLevelSearchOps OK. searchOps is too high! 5496 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482163780] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482163823] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5073 found since last run. oneLevelSearchOps OK. searchOps is too high! 5507 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482163943] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482164380] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5242 found since last run. oneLevelSearchOps OK. searchOps is too high! 5432 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482164500] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482164543] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6407 found since last run. oneLevelSearchOps OK. searchOps is too high! 6909 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482164663] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482165263] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6342 found since last run. oneLevelSearchOps OK. searchOps is too high! 6672 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482165383] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482165983] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7436 found since last run. oneLevelSearchOps OK. searchOps is too high! 7937 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482166092] SERVICE ALERT: owen101lab3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1482166094] SERVICE ALERT: owen101lab2;Device Status;OK;HARD;3;Genuine HP cartridge installed
[1482166103] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482166300] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5781 found since last run. oneLevelSearchOps OK. searchOps is too high! 5985 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482166420] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482166449] SERVICE ALERT: owen101lab3;Device Status;OK;HARD;3;Ready
[1482166496] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 46% - OK!
[1482166670] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 61% - OK!
[1482166671] SERVICE ALERT: owen101lab2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1482166703] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5193 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482166823] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482166975] SERVICE ALERT: CAVA2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1482167020] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7934 found since last run. oneLevelSearchOps OK. searchOps is too high! 8113 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482167095] SERVICE ALERT: CAVA2;CPU Load;OK;SOFT;2;CPU Load 81% (15 min average)
[1482167140] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482167312] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482167372] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482167423] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6224 found since last run. oneLevelSearchOps OK. searchOps is too high! 6632 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482167433] SERVICE ALERT: halleground2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482167543] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482167723] SERVICE ALERT: hallesouthcommons;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1482167740] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7240 found since last run. oneLevelSearchOps OK. searchOps is too high! 7451 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482167860] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482167871] SERVICE ALERT: rooseveltpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1482168032] SERVICE ALERT: halleground2;Device Status;OK;HARD;3;TRAY 2 EMPTY ANY TYPE LETTER   Processing job
[1482168143] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5458 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482168263] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482168603] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1482168603] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1482168603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1482168603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1482168603] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1482168603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1482168603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1482168603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1482168603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1482168603] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1482168603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1482168603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1482168603] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1482168603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1482168603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1482168603] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1482168863] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5277 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482168889] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1482168983] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482169490] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1482169549] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1482169610] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1482169623] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1482169660] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5716 found since last run. oneLevelSearchOps OK. searchOps is too high! 5980 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482169780] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482170183] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6067 found since last run. oneLevelSearchOps OK. searchOps is too high! 6626 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482170223] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 91% (15 min average)
[1482170223] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (24.0% change >= 20.0% threshold)
[1482170303] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482170809] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1482170823] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1482170903] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6651 found since last run. oneLevelSearchOps OK. searchOps is too high! 7051 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482170943] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 79% (15 min average)
[1482170980] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5144 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482171023] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482171100] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482171543] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1482171623] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6225 found since last run. oneLevelSearchOps OK. searchOps is too high! 6650 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482171663] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1482171743] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482171783] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1482172943] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6977 found since last run. oneLevelSearchOps OK. searchOps is too high! 7299 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482173063] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482173583] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 49% (15 min average)
[1482173663] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6098 found since last run. oneLevelSearchOps OK. searchOps is too high! 6461 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482173783] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482174100] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5578 found since last run. oneLevelSearchOps OK. searchOps is too high! 5683 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482174183] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1482174220] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482174303] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 67% (15 min average)
[1482174383] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5296 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482174503] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482175103] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6495 found since last run. oneLevelSearchOps OK. searchOps is too high! 6960 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482175223] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482175823] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5081 found since last run. oneLevelSearchOps OK. searchOps is too high! 5264 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482175943] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482176103] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1482176223] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 91% (15 min average)
[1482176620] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5981 found since last run. oneLevelSearchOps OK. searchOps is too high! 6200 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482176740] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482177090] SERVICE ALERT: epprogsql2;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1482177210] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1482177330] SERVICE ALERT: epprogsql2;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1482177340] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6837 found since last run. oneLevelSearchOps OK. searchOps is too high! 7043 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482177423] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1482177460] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482177543] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1482177603] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1482177603] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1482177603] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1482177603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1482177603] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1482177603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1482177603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1482177603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1482177603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1482177603] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1482177603] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1482177603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1482177603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1482177603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1482177603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1482177603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1482177663] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 72% (15 min average)
[1482177663] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1482177721] SERVICE ALERT: prayharroldpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482177781] SERVICE ALERT: prayharroldpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482177841] SERVICE ALERT: prayharroldpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482177930] SERVICE ALERT: epprogsql2;CPU Load;OK;HARD;3;CPU Load 0% (15 min average)
[1482178102] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1482178132] HOST ALERT: prayharroldpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1482178156] SERVICE ALERT: prayharroldpk;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! 1 of 3 Trays Still Full
[1482178160] HOST ALERT: prayharroldpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.55 ms
[1482178187] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482178247] SERVICE ALERT: sillpk;Device Status;OK;SOFT;2;Clearing paper path
[1482178263] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1482178383] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1482178441] SERVICE ALERT: prayharroldpk;Device Status;OK;HARD;3;PERFORM PRINTER MAINTENANCE 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1482178503] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1482178530] SERVICE ALERT: epprogsql2;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1482178650] SERVICE ALERT: epprogsql2;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1482178660] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5780 found since last run. oneLevelSearchOps OK. searchOps is too high! 5969 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482178770] SERVICE ALERT: epprogsql2;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1482178780] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482178847] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482178907] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482178967] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482179103] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 82% (15 min average)
[1482179370] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1482179370] SERVICE FLAPPING ALERT: epprogsql2;CPU Load;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1482179380] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5030 found since last run. oneLevelSearchOps OK. searchOps is too high! 5169 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482179500] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482179543] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5321 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482179663] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482179969] SERVICE ALERT: epprogsql2;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1482180263] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5166 found since last run. oneLevelSearchOps OK. searchOps is too high! 5504 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482180303] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1482180383] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482180423] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1482180543] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1482180983] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5158 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482181103] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482181169] SERVICE ALERT: epprogsql2;CPU Load;OK;HARD;3;CPU Load 2% (15 min average)
[1482181900] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6495 found since last run. oneLevelSearchOps OK. searchOps is too high! 6544 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482182020] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482182123] SERVICE ALERT: hallesouthcommons;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1482182184] SERVICE ALERT: hallesouthcommons;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1482182244] SERVICE ALERT: hallesouthcommons;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1482182303] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5610 found since last run. oneLevelSearchOps OK. searchOps is too high! 5994 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482182423] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482182943] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1482182969] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1482183023] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5040 found since last run. oneLevelSearchOps OK. searchOps is too high! 5383 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482183089] SERVICE ALERT: epprogsql2;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1482183143] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482183209] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1482183809] SERVICE ALERT: epprogsql2;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1482184009] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1482184070] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1482184130] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1482184143] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1482184263] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1482184343] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5102 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482184367] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Remove any jammed paper.
[1482184383] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1482184463] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482184658] SERVICE ALERT: sillpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1482184718] SERVICE ALERT: sillpk;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 100% - OK!
[1482185020] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5959 found since last run. oneLevelSearchOps OK. searchOps is too high! 6118 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482185063] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6292 found since last run. oneLevelSearchOps OK. searchOps is too high! 6635 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482185140] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482185183] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482185583] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 52% (15 min average)
[1482185609] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1482185783] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5234 found since last run. oneLevelSearchOps OK. searchOps is too high! 5538 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482185903] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482186167] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482186183] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1482186209] SERVICE ALERT: epprogsql2;CPU Load;OK;HARD;3;CPU Load 84% (15 min average)
[1482186227] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482186287] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482186303] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1482186423] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1482186841] SERVICE ALERT: prayharroldpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482186901] SERVICE ALERT: prayharroldpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482186961] SERVICE ALERT: prayharroldpk;Device Status;OK;SOFT;3;RETRIEVE JOB  NO STORED JOBS
[1482187103] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5127 found since last run. oneLevelSearchOps OK. searchOps is too high! 5505 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482187157] SERVICE ALERT: prayharroldpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1482187217] SERVICE ALERT: prayharroldpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1482187223] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482187277] SERVICE ALERT: prayharroldpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1482187487] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1482187487] SERVICE FLAPPING ALERT: sillpk;Device Status;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1482187562] SERVICE ALERT: prayharroldpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482187621] SERVICE ALERT: prayharroldpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482187681] SERVICE ALERT: prayharroldpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482188200] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1482188223] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1482188260] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1482188281] SERVICE ALERT: prayharroldpk;Device Status;OK;HARD;3;CLOSE RIGHT DOOR  46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1482188281] SERVICE FLAPPING ALERT: prayharroldpk;Device Status;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1482188321] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;OK;SOFT;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1482188477] SERVICE ALERT: prayharroldpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1482188575] SERVICE ALERT: cava2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1482188695] SERVICE ALERT: cava2;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1482188815] SERVICE ALERT: cava2;CPU Load;OK;SOFT;3;CPU Load 1% (15 min average)
[1482188823] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1482189023] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5070 found since last run. oneLevelSearchOps OK. searchOps is too high! 5428 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482189143] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482189423] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 90% (15 min average)
[1482190023] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1482190143] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1482190263] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1482190863] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 69% (15 min average)
[1482191463] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1482191583] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1482191703] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1482192303] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 82% (15 min average)
[1482192940] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5433 found since last run. oneLevelSearchOps OK. searchOps is too high! 5582 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1482193060] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482193503] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1482193623] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1482193743] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1482194930] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1482195887] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482195948] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482196007] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482196044] SERVICE ALERT: hallesouthcommons;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1482196080] SERVICE ALERT: prayharroldpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482196140] SERVICE ALERT: prayharroldpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482196143] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1482196200] SERVICE ALERT: prayharroldpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482196276] SERVICE ALERT: prayharroldpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1482196336] SERVICE ALERT: prayharroldpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1482196396] SERVICE ALERT: prayharroldpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1482196743] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1482196863] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 82% (15 min average)
[1482197318] SERVICE ALERT: sillpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1482197384] SERVICE ALERT: sillpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1482197439] SERVICE ALERT: sillpk;Black Print Cartridge HP C8543X;OK;SOFT;3;Black Print Cartridge HP C8543X is at 58% - OK!
[1482197463] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1482197583] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1482197608] SERVICE FLAPPING ALERT: epprogsql2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1482197703] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 68% (15 min average)
[1482197807] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Warming up
[1482198303] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1482198423] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1482198543] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1482199637] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;SOFT;1;e:\ - total: 799.87 Gb - used: 654.08 Gb (82%) - free 145.79 Gb (18%)
[1482199743] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 79% (15 min average)
[1482199757] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;SOFT;2;e:\ - total: 799.87 Gb - used: 654.08 Gb (82%) - free 145.79 Gb (18%)
[1482199877] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;HARD;3;e:\ - total: 799.87 Gb - used: 654.08 Gb (82%) - free 145.79 Gb (18%)
[1482200343] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1482200463] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 73% (15 min average)
[1482201063] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1482201183] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1482201303] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 86% (15 min average)
[1482202200] SERVICE FLAPPING ALERT: prayharroldpk;Device Status;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1482202277] SERVICE ALERT: AD-RMAD;E:\ Drive Space;OK;HARD;3;e:\ - total: 799.87 Gb - used: 629.90 Gb (79%) - free 169.98 Gb (21%)
[1482202503] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1482202623] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1482202743] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 94% (15 min average)
[1482203343] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1482203463] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1482203583] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 92% (15 min average)
[1482204183] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1482204303] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1482204423] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1482205623] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1482206823] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1482208007] SERVICE FLAPPING ALERT: sillpk;Device Status;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1482208830] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482208889] SERVICE ALERT: halleground2;Device Status;OK;SOFT;2;Clearing paper path
[1482209824] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)