[1457760738] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 87% (5 min average)
[1457761338] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457761458] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1457761578] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1457763102] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 84.90 Gb - used: 78.78 Gb (93%) - free 6.12 Gb (7%)
[1457766790] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1457766910] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 88% (5 min average)
[1457768110] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457768230] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 88% (5 min average)
[1457768830] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1457768950] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 90% (5 min average)
[1457769070] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1457770270] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 97% (5 min average)
[1457771470] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 81% (5 min average)
[1457771778] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 94% (5 min average)
[1457772070] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1457772190] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 93% (5 min average)
[1457772310] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 69% (5 min average)
[1457772378] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1457775378] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 40% (5 min average)
[1457776578] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1457776698] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 57% (5 min average)
[1457777898] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1457778018] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1457778138] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 88% (5 min average)
[1457778738] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1457778858] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 96% (5 min average)
[1457778978] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 97% (5 min average)
[1457779578] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 38% (5 min average)
[1457780226] SERVICE ALERT: OFFICEKMS;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1457780346] SERVICE ALERT: OFFICEKMS;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1457781978] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1457782098] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1457782218] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1457783418] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 75% (5 min average)
[1457784618] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1457784738] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1457784858] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1457787258] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 94% (5 min average)
[1457788458] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1457789058] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 25% (5 min average)
[1457789658] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1457789778] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 82% (5 min average)
[1457792178] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457792298] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1457792418] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 96% (5 min average)
[1457793702] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 84.90 Gb - used: 75.72 Gb (89%) - free 9.18 Gb (11%)
[1457794218] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 59% (5 min average)
[1457794818] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1457794938] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 79% (5 min average)
[1457796738] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1457796858] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1457796978] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1457798178] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 81% (5 min average)
[1457798778] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1457798898] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 96% (5 min average)
[1457799018] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1457799618] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 75% (5 min average)
[1457800218] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1457800337] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 96% (5 min average)
[1457800457] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1457801656] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 93% (5 min average)
[1457802255] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 75% (5 min average)
[1457803455] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1457803575] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 95% (5 min average)
[1457803695] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1457804895] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 65% (5 min average)
[1457805495] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457805615] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 87% (5 min average)
[1457806215] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457806335] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1457806455] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1457807655] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 54% (5 min average)
[1457808855] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1457808975] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1457809095] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 70% (5 min average)
[1457809695] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1457809815] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1457809935] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 94% (5 min average)
[1457810535] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 64% (5 min average)
[1457811135] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1457811255] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 83% (5 min average)
[1457813055] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 94% (5 min average)
[1457813175] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 97% (5 min average)
[1457813295] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 71% (5 min average)
[1457814495] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1457814615] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1457814735] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1457816535] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 85% (5 min average)
[1457818935] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1457819055] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1457819175] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1457819775] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 78% (5 min average)
[1457820975] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1457821095] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 64% (5 min average)
[1457822295] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457822415] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 96% (5 min average)
[1457822535] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1457824335] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 80% (5 min average)
[1457824935] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1457825055] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 59% (5 min average)
[1457825655] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457825775] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1457825895] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1457827695] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 93% (5 min average)
[1457828295] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1457829495] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 93% (5 min average)
[1457830093] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1457834448] SERVICE ALERT: AXIOM-2012;CPU Load;WARNING;SOFT;1;CPU Load 94% (5 min average)
[1457834568] SERVICE ALERT: AXIOM-2012;CPU Load;OK;SOFT;2;CPU Load 9% (5 min average)
[1457834606] SERVICE ALERT: EVENTSTEST-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457834607] SERVICE ALERT: TARGETXTEST2012;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457834726] SERVICE ALERT: EVENTSTEST-2012;CPU Load;OK;SOFT;2;CPU Load 85% (5 min average)
[1457834727] SERVICE ALERT: TARGETXTEST2012;CPU Load;OK;SOFT;2;CPU Load 83% (5 min average)
[1457834893] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 77% (5 min average)
[1457835062] SERVICE ALERT: CASPER;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1457835182] SERVICE ALERT: CASPER;CPU Load;OK;SOFT;2;CPU Load 43% (5 min average)
[1457835493] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457835580] SERVICE ALERT: TARGETX-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457835613] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1457835700] SERVICE ALERT: TARGETX-2012;CPU Load;OK;SOFT;2;CPU Load 82% (5 min average)
[1457835733] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1457835926] SERVICE ALERT: EVENTSTEST-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457836046] SERVICE ALERT: EVENTSTEST-2012;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1457836108] SERVICE ALERT: AXIOM-TEST-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457836166] SERVICE ALERT: EVENTSTEST-2012;CPU Load;OK;SOFT;3;CPU Load 23% (5 min average)
[1457836228] SERVICE ALERT: AXIOM-TEST-2012;CPU Load;OK;SOFT;2;CPU Load 75% (5 min average)
[1457836368] SERVICE ALERT: AXIOM-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457836407] SERVICE ALERT: cava2;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457836488] SERVICE ALERT: AXIOM-2012;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1457836527] SERVICE ALERT: cava2;CPU Load;OK;SOFT;2;CPU Load 47% (5 min average)
[1457836527] SERVICE ALERT: TARGETXTEST2012;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457836568] SERVICE ALERT: SANDBOX-APP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457836647] SERVICE ALERT: TARGETXTEST2012;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1457836688] SERVICE ALERT: SANDBOX-APP;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1457837043] SERVICE ALERT: EVENTS-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457837163] SERVICE ALERT: EVENTS-2012;CPU Load;OK;SOFT;2;CPU Load 1% (5 min average)
[1457837343] SERVICE ALERT: OFFICEKMS;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457837463] SERVICE ALERT: OFFICEKMS;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1457837500] SERVICE ALERT: TARGETX-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457837557] SERVICE ALERT: DB01-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457837620] SERVICE ALERT: TARGETX-2012;CPU Load;OK;SOFT;2;CPU Load 1% (5 min average)
[1457837677] SERVICE ALERT: DB01-2012;CPU Load;OK;SOFT;2;CPU Load 8% (5 min average)
[1457837966] SERVICE ALERT: EVENTSTEST-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457838086] SERVICE ALERT: EVENTSTEST-2012;CPU Load;OK;SOFT;2;CPU Load 27% (5 min average)
[1457841308] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1457841428] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 77% (5 min average)