[1455098586] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;SOFT;1;e:\ - total: 99.87 Gb - used: 83.07 Gb (83%) - free 16.80 Gb (17%)
[1455098706] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;SOFT;2;e:\ - total: 99.87 Gb - used: 83.33 Gb (83%) - free 16.54 Gb (17%)
[1455098826] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;HARD;3;e:\ - total: 99.87 Gb - used: 83.33 Gb (83%) - free 16.54 Gb (17%)
[1455098887] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;SOFT;1;e:\ - total: 99.87 Gb - used: 83.33 Gb (83%) - free 16.54 Gb (17%)
[1455099007] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;SOFT;2;e:\ - total: 99.87 Gb - used: 83.33 Gb (83%) - free 16.54 Gb (17%)
[1455099127] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;HARD;3;e:\ - total: 99.87 Gb - used: 83.33 Gb (83%) - free 16.54 Gb (17%)
[1455101021] SERVICE ALERT: ADWU1;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1455101141] SERVICE ALERT: ADWU1;CPU Load;OK;SOFT;2;CPU Load 52% (5 min average)
[1455101690] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;1;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1455101810] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;2;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1455101930] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;HARD;3;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1455104721] HOST ALERT: halleground3;UP;HARD;10;PING OK - Packet loss = 61%, RTA = 0.57 ms
[1455104888] HOST ALERT: halleground3;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1455105100] SERVICE ALERT: halleground3;Device Status;OK;HARD;3;Ready
[1455105158] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 21% - OK!
[1455105160] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 21% - OK!
[1455105178] SERVICE ALERT: halleground3;Device Status;OK;HARD;3;Ready
[1455105178] SERVICE ALERT: halleground3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1455105304] SERVICE ALERT: halleground3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1455107437] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1455107497] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1455107527] SERVICE ALERT: AD-RMAD;E:\ Drive Space;OK;HARD;3;e:\ - total: 199.87 Gb - used: 83.33 Gb (42%) - free 116.54 Gb (58%)
[1455107527] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1455107557] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1455107587] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1455107648] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1455107825] SERVICE ALERT: AD-RMAD;E:\ Drive Space;OK;HARD;3;e:\ - total: 199.87 Gb - used: 83.33 Gb (42%) - free 116.54 Gb (58%)
[1455107879] SERVICE ALERT: ADWU1;C:\ Drive Space;OK;HARD;3;c:\ - total: 419.66 Gb - used: 322.85 Gb (77%) - free 96.80 Gb (23%)
[1455108157] SERVICE ALERT: markjeffersonpk;Device Status;OK;HARD;3;Open left door. Fuser is HOT.
[1455108166] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1455108247] SERVICE ALERT: markjeffersonpk;Device Status;OK;HARD;3;Close door. Reattach left accessory. 54 6F 20 65 78 69 74 20 70 72 65 73 73 20 1F 2E
[1455108766] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;CRITICAL;HARD;3;Black Print Cartridge HP C8543X is at CRITICAL level!
[1455109128] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;OK;HARD;3;OK: No entries found
[1455111362] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1455111364] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1455111412] SERVICE ALERT: hallefirst2;Paper Tray Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1455111420] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.104.180! Make sure host is up and SNMP is configured properly.
[1455111422] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1455111424] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1455111424] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1455111425] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1455111472] SERVICE ALERT: hallefirst2;Paper Tray Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1455111480] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.104.180! Make sure host is up and SNMP is configured properly.
[1455111482] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1455111484] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1455111484] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1455111485] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1455111532] SERVICE ALERT: hallefirst2;Paper Tray Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1455111540] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;WARNING;HARD;3;WARNING: No SNMP response from 164.76.104.180! Make sure host is up and SNMP is configured properly.
[1455111544] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;WARNING;HARD;3;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1455111545] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;WARNING;HARD;3;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1455111609] SERVICE ALERT: hallefirst3;Paper Tray Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.104.180! Make sure host is up and SNMP is configured properly.
[1455111668] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.104.232! Make sure host is up and SNMP is configured properly.
[1455111669] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.104.232! Make sure host is up and SNMP is configured properly.
[1455111669] SERVICE ALERT: hallefirst3;Paper Tray Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.104.180! Make sure host is up and SNMP is configured properly.
[1455111711] SERVICE ALERT: hallefirst2;Paper Tray Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1455111715] SERVICE ALERT: hallefirst1;Paper Tray Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.104.232! Make sure host is up and SNMP is configured properly.
[1455111717] SERVICE ALERT: hallefirst3;Device Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.104.180! Make sure host is up and SNMP is configured properly.
[1455111723] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.104.232! Make sure host is up and SNMP is configured properly.
[1455111728] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.104.232! Make sure host is up and SNMP is configured properly.
[1455111729] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.104.232! Make sure host is up and SNMP is configured properly.
[1455111729] SERVICE ALERT: hallefirst3;Paper Tray Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.104.180! Make sure host is up and SNMP is configured properly.
[1455111734] SERVICE ALERT: hallefirst1;Paper Tray Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.104.232! Make sure host is up and SNMP is configured properly.
[1455111771] SERVICE ALERT: hallefirst2;Paper Tray Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1455111775] SERVICE ALERT: hallefirst1;Paper Tray Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.104.232! Make sure host is up and SNMP is configured properly.
[1455111777] SERVICE ALERT: hallefirst3;Device Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.104.180! Make sure host is up and SNMP is configured properly.
[1455111784] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.104.232! Make sure host is up and SNMP is configured properly.
[1455111788] SERVICE ALERT: hallefirst1;Device Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.104.232! Make sure host is up and SNMP is configured properly.
[1455111789] SERVICE ALERT: hallefirst1;Device Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.104.232! Make sure host is up and SNMP is configured properly.
[1455111794] SERVICE ALERT: hallefirst1;Paper Tray Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.104.232! Make sure host is up and SNMP is configured properly.
[1455111797] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.104.180! Make sure host is up and SNMP is configured properly.
[1455111811] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.104.232! Make sure host is up and SNMP is configured properly.
[1455111831] SERVICE ALERT: hallefirst2;Paper Tray Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1455111835] SERVICE ALERT: hallefirst1;Paper Tray Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.104.232! Make sure host is up and SNMP is configured properly.
[1455111838] SERVICE ALERT: hallefirst3;Device Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.104.180! Make sure host is up and SNMP is configured properly.
[1455111843] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;HARD;3;WARNING: No SNMP response from 164.76.104.232! Make sure host is up and SNMP is configured properly.
[1455111854] SERVICE ALERT: hallefirst1;Paper Tray Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.104.232! Make sure host is up and SNMP is configured properly.
[1455111858] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.104.180! Make sure host is up and SNMP is configured properly.
[1455111871] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.104.232! Make sure host is up and SNMP is configured properly.
[1455111892] SERVICE ALERT: hallefirst3;Device Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.104.180! Make sure host is up and SNMP is configured properly.
[1455111917] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;WARNING;HARD;3;WARNING: No SNMP response from 164.76.104.180! Make sure host is up and SNMP is configured properly.
[1455111917] SERVICE ALERT: hallefirst3;Paper Tray Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.104.180! Make sure host is up and SNMP is configured properly.
[1455111931] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;HARD;3;WARNING: No SNMP response from 164.76.104.232! Make sure host is up and SNMP is configured properly.
[1455111952] SERVICE ALERT: hallefirst3;Device Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.104.180! Make sure host is up and SNMP is configured properly.
[1455111977] SERVICE ALERT: hallefirst3;Paper Tray Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.104.180! Make sure host is up and SNMP is configured properly.
[1455112012] SERVICE ALERT: hallefirst3;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.180! Make sure host is up and SNMP is configured properly.
[1455112037] SERVICE ALERT: hallefirst3;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.180! Make sure host is up and SNMP is configured properly.
[1455112067] HOST ALERT: hallefirst3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1455112076] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Warming up
[1455112078] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1455112078] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Processing job
[1455112083] SERVICE ALERT: hallethirdsouth;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.107.213! Make sure host is up and SNMP is configured properly.
[1455112098] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1455112098] SERVICE ALERT: hallefirst3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1455112127] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1455112131] HOST ALERT: hallefirst3;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1455112138] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 38% - OK!
[1455112140] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 38% - OK!
[1455112184] SERVICE ALERT: owen101lab3;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103.  Is the printer off?  Is the network down?
[1455112191] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103.  Is the printer off?  Is the network down?
[1455112212] SERVICE ALERT: hallethirdsouth;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.107.213.  Is the printer off?  Is the network down?
[1455112382] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;Processing job
[1455112384] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;Processing job from tray 4
[1455112387] SERVICE ALERT: hallethirdsouth;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.107.213.  Is the printer off?  Is the network down?
[1455112426] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1455112430] SERVICE ALERT: hallefirst1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1455112431] SERVICE ALERT: hallefirst3;Device Status;OK;HARD;3;TRAY 2 EMPTY PLAIN LETTER   Ready
[1455112436] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 66% - OK!
[1455112448] SERVICE ALERT: hallefirst1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1455112450] SERVICE ALERT: hallethirdsouth;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.107.213.  Is the printer off?  Is the network down?
[1455112456] SERVICE ALERT: owen101lab3;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103.  Is the printer off?  Is the network down?
[1455112510] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 42% - OK!
[1455112525] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 66% - OK!
[1455112606] SERVICE ALERT: hallefirst3;Device Status;OK;HARD;3;TRAY 2 EMPTY PLAIN LETTER   Processing job
[1455112632] SERVICE ALERT: hallefirst3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1455112692] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 42% - OK!
[1455113028] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1455113088] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1455113148] SERVICE ALERT: crossroadspk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1455113748] SERVICE ALERT: crossroadspk;Device Status;OK;HARD;3;Open tray indicated.
[1455113934] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1455113993] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1455114053] SERVICE ALERT: crossroadspk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1455114596] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1455114653] SERVICE ALERT: crossroadspk;Device Status;OK;HARD;3;LOAD TRAY 2 PLAIN LETTER 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1455114656] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1455114715] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1455114806] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1455114865] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1455114926] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1455117600] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1455117978] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1455120765] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1455121085] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1455123324] SERVICE ALERT: crossroadspk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1455123715] SERVICE ALERT: crossroadspk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1455123766] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1455124726] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;1;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1455124846] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;2;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1455124966] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;HARD;3;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1455129757] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1455129817] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1455129833] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1455129847] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1455129877] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1455129893] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1455129906] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1455129952] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1455129966] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1455130024] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1455130084] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1455130144] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1455130566] SERVICE ALERT: markjeffersonpk;Device Status;OK;HARD;3;Clearing paper path
[1455130744] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1455130793] SERVICE ALERT: strongpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.130.31. Is the printer off?  Is the network down?
[1455130800] SERVICE ALERT: strongpk;Black Cartridge HP CC364X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.130.31. Is the printer off?  Is the network down?
[1455130823] HOST ALERT: strongpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1455130853] SERVICE ALERT: strongpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.130.31. Is the printer off?  Is the network down?
[1455130860] SERVICE ALERT: strongpk;Black Cartridge HP CC364X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.130.31. Is the printer off?  Is the network down?
[1455130864] HOST ALERT: strongpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1455130872] SERVICE ALERT: strongpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.130.31. Is the printer off?  Is the network down?
[1455130883] HOST ALERT: strongpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1455130900] HOST ALERT: strongpk;UP;SOFT;2;PING WARNING - Packet loss = 82%, RTA = 0.48 ms
[1455130907] SERVICE ALERT: strongpk;Device Status;OK;SOFT;3;Ready
[1455130911] HOST ALERT: strongpk;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.66 ms
[1455130914] SERVICE ALERT: strongpk;Black Cartridge HP CC364X;OK;SOFT;3;Black Cartridge HP CC364X is at 17% - OK!
[1455130925] SERVICE ALERT: strongpk;Device Status;OK;SOFT;2;Ready
[1455131077] SERVICE ALERT: markjeffersonpk;Device Status;OK;HARD;3;Ready
[1455131153] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1455131230] HOST ALERT: strongpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1455131241] HOST ALERT: strongpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1455131298] SERVICE ALERT: strongpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.130.31. Is the printer off?  Is the network down?
[1455131320] HOST ALERT: strongpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1455131328] HOST ALERT: strongpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1455131358] SERVICE ALERT: strongpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.130.31. Is the printer off?  Is the network down?
[1455131388] HOST ALERT: strongpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1455131410] HOST ALERT: strongpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1455131418] SERVICE ALERT: strongpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.130.31. Is the printer off?  Is the network down?
[1455131448] HOST ALERT: strongpk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1455131500] HOST ALERT: strongpk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1455131513] SERVICE ALERT: strongpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.130.31. Is the printer off?  Is the network down?
[1455131520] SERVICE ALERT: strongpk;Black Cartridge HP CC364X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.130.31. Is the printer off?  Is the network down?
[1455131526] SERVICE ALERT: strongpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.130.31. Is the printer off?  Is the network down?
[1455131531] SERVICE ALERT: strongpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.130.31. Is the printer off?  Is the network down?
[1455131538] HOST ALERT: strongpk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1455131556] HOST ALERT: strongpk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1455131566] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;OK;HARD;3;OK: No entries found
[1455131573] SERVICE ALERT: strongpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.130.31. Is the printer off?  Is the network down?
[1455131580] SERVICE ALERT: strongpk;Black Cartridge HP CC364X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.130.31. Is the printer off?  Is the network down?
[1455131586] SERVICE ALERT: strongpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.130.31. Is the printer off?  Is the network down?
[1455131591] SERVICE ALERT: strongpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.130.31. Is the printer off?  Is the network down?
[1455131603] HOST ALERT: strongpk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1455131616] HOST ALERT: strongpk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1455131617] SERVICE ALERT: strongpk;Black Cartridge HP CC364X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.130.31. Is the printer off?  Is the network down?
[1455131633] SERVICE ALERT: strongpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.130.31. Is the printer off?  Is the network down?
[1455131641] SERVICE ALERT: strongpk;Black Cartridge HP CC364X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.130.31. Is the printer off?  Is the network down?
[1455131646] SERVICE ALERT: strongpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.130.31. Is the printer off?  Is the network down?
[1455131647] HOST ALERT: strongpk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1455131651] SERVICE ALERT: strongpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.130.31. Is the printer off?  Is the network down?
[1455131663] HOST ALERT: strongpk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1455131677] SERVICE ALERT: strongpk;Black Cartridge HP CC364X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.130.31. Is the printer off?  Is the network down?
[1455131681] HOST ALERT: strongpk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1455131737] SERVICE ALERT: strongpk;Black Cartridge HP CC364X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.130.31. Is the printer off?  Is the network down?
[1455131753] HOST ALERT: strongpk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1455131767] HOST ALERT: strongpk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1455131843] HOST ALERT: strongpk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1455131857] HOST ALERT: strongpk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1455131933] HOST ALERT: strongpk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1455132059] SERVICE ALERT: prayharroldpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1455132119] SERVICE ALERT: prayharroldpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1455132179] SERVICE ALERT: prayharroldpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1455132298] SERVICE ALERT: prayharroldpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1455132358] SERVICE ALERT: prayharroldpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1455132418] SERVICE ALERT: prayharroldpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1455133324] SERVICE ALERT: ADDC5;GPO Replication;WARNING;SOFT;1;Warning! 2 policies are not replicating properly!
[1455134528] SERVICE ALERT: ADDC5;GPO Replication;WARNING;SOFT;2;Warning! 2 policies are not replicating properly!
[1455134815] SERVICE ALERT: ADDC5;GPO Replication;WARNING;SOFT;1;Warning! 2 policies are not replicating properly!
[1455134819] SERVICE ALERT: prayharroldpk;Device Status;OK;HARD;3;RETRIEVE JOB  NO STORED JOBS
[1455135132] SERVICE ALERT: ADDC5;GPO Replication;WARNING;SOFT;2;Warning! 2 policies are not replicating properly!
[1455135179] SERVICE ALERT: prayharroldpk;Device Status;OK;HARD;3;Processing job from tray 4
[1455135231] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1455135261] HOST ALERT: hallefirst2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1455135270] HOST ALERT: hallefirst2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1455135291] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1455135321] HOST ALERT: hallefirst2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1455135351] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1455135360] HOST ALERT: hallefirst2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1455135381] HOST ALERT: hallefirst2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1455135414] SERVICE ALERT: ADDC5;GPO Replication;WARNING;HARD;3;Warning! 2 policies are not replicating properly!
[1455135450] HOST ALERT: hallefirst2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1455135471] HOST ALERT: hallefirst2;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1455135480] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1455135484] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1455135510] HOST ALERT: hallefirst2;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1455135514] HOST ALERT: hallefirst2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1455135530] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1455135540] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1455135544] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1455135544] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1455135545] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1455135560] HOST ALERT: hallefirst2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1455135574] HOST ALERT: hallefirst2;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1455135591] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1455135600] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1455135604] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1455135604] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1455135605] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1455135621] HOST ALERT: hallefirst2;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1455135634] HOST ALERT: hallefirst2;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1455135650] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1455135664] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1455135666] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1455135680] HOST ALERT: hallefirst2;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1455135696] HOST ALERT: hallefirst2;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1455135733] SERVICE ALERT: ADDC5;GPO Replication;WARNING;HARD;3;Warning! 2 policies are not replicating properly!
[1455135770] HOST ALERT: hallefirst2;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1455135786] HOST ALERT: hallefirst2;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1455135859] HOST ALERT: hallefirst2;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1455135876] HOST ALERT: hallefirst2;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1455135949] HOST ALERT: hallefirst2;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1455136305] SERVICE ALERT: EVENTSTEST-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1455136425] SERVICE ALERT: EVENTSTEST-2012;CPU Load;OK;SOFT;2;CPU Load 2% (5 min average)
[1455136630] SERVICE ALERT: studentcenterpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.150.252. Is the printer off?  Is the network down?
[1455136653] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.150.252. Is the printer off?  Is the network down?
[1455136660] HOST ALERT: studentcenterpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1455136683] HOST ALERT: studentcenterpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1455136690] SERVICE ALERT: studentcenterpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.150.252. Is the printer off?  Is the network down?
[1455136713] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.150.252. Is the printer off?  Is the network down?
[1455136717] HOST ALERT: studentcenterpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.65 ms
[1455136718] HOST ALERT: studentcenterpk;UP;SOFT;2;PING WARNING - Packet loss = 82%, RTA = 0.68 ms
[1455136745] SERVICE ALERT: studentcenterpk;Paper Tray Status;OK;SOFT;3;Printer Paper Is OK! All Trays Still Full
[1455136767] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;OK;SOFT;3;Black Print Cartridge HP C8543X is at 37% - OK!
[1455136794] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Ready
[1455136798] HOST ALERT: hallefirst2;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.45 ms
[1455136798] HOST ALERT: hallethirdsouth;UP;HARD;10;PING OK - Packet loss = 75%, RTA = 0.27 ms
[1455136799] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Ready
[1455136803] HOST ALERT: hallefirst2;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.44 ms
[1455136815] HOST ALERT: hallethirdsouth;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1455136843] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1455136858] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 37% - OK!
[1455136859] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 37% - OK!
[1455136890] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1455136915] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1455136949] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1455136975] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1455137010] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1455137034] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1455137107] SERVICE ALERT: halleground3;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1455137133] HOST ALERT: halleground3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1455137137] HOST ALERT: halleground3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1455137146] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1455137167] SERVICE ALERT: halleground3;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1455137197] HOST ALERT: halleground3;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1455137223] HOST ALERT: halleground3;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1455137227] SERVICE ALERT: halleground3;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1455137257] HOST ALERT: halleground3;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1455137313] HOST ALERT: halleground3;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1455137347] HOST ALERT: halleground3;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1455137403] HOST ALERT: halleground3;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1455137437] HOST ALERT: halleground3;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1455137493] HOST ALERT: halleground3;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1455137505] SERVICE ALERT: halleground3;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1455137527] HOST ALERT: halleground3;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1455137564] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1455137564] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1455137565] SERVICE ALERT: halleground3;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1455137582] SERVICE ALERT: halleground3;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1455137583] SERVICE ALERT: halleground3;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1455137583] HOST ALERT: halleground3;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1455137594] HOST ALERT: halleground3;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1455137624] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1455137624] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1455137625] SERVICE ALERT: halleground3;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1455137642] SERVICE ALERT: halleground3;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1455137643] SERVICE ALERT: halleground3;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1455137654] HOST ALERT: halleground3;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1455137654] HOST ALERT: halleground3;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1455137684] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1455137685] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1455137702] SERVICE ALERT: halleground3;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1455137703] SERVICE ALERT: halleground3;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1455137714] HOST ALERT: halleground3;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1455137715] HOST ALERT: halleground3;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1455137804] HOST ALERT: halleground3;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1455137804] HOST ALERT: halleground3;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1455137894] HOST ALERT: halleground3;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1455138607] HOST ALERT: halleground3;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1455138761] HOST ALERT: halleground3;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.52 ms
[1455138809] SERVICE ALERT: eateriespk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1455138819] SERVICE ALERT: halleground3;Device Status;OK;HARD;3;Ready
[1455138835] SERVICE ALERT: eateriespk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1455138877] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1455138878] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1455138896] SERVICE ALERT: halleground3;Device Status;OK;HARD;3;Ready
[1455138897] SERVICE ALERT: halleground3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1455139021] SERVICE ALERT: halleground3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1455139364] SERVICE ALERT: IMAGES;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1455139484] SERVICE ALERT: IMAGES;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1455140440] HOST ALERT: hallethirdsouth;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1455140530] HOST ALERT: hallethirdsouth;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1455140617] HOST ALERT: hallethirdsouth;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1455140620] HOST ALERT: hallethirdsouth;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1455140679] HOST ALERT: hallethirdsouth;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1455140710] HOST ALERT: hallethirdsouth;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1455140743] HOST ALERT: hallethirdsouth;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1455140774] HOST ALERT: hallethirdsouth;UP;SOFT;5;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1455142618] SERVICE ALERT: ADDC5;GPO Replication;OK;HARD;3;OK! All GPOs are properly replicated.
[1455142928] SERVICE ALERT: ADDC5;GPO Replication;OK;HARD;3;OK! All GPOs are properly replicated.
[1455145561] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1455145621] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1455145681] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1455145841] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1455145900] SERVICE ALERT: hallesecondsouth;Device Status;OK;SOFT;2;Clearing paper path
[1455146281] SERVICE ALERT: hallesecondsouth;Device Status;OK;HARD;3;TRAY 4 EMPTY PLAIN LETTER   Ready
[1455150481] SERVICE ALERT: eateriespk;Black Cartridge HP CC364X;WARNING;SOFT;1;Black Cartridge HP CC364X is at 5% - WARNING!
[1455150542] SERVICE ALERT: eateriespk;Black Cartridge HP CC364X;WARNING;SOFT;2;Black Cartridge HP CC364X is at 5% - WARNING!
[1455150576] SERVICE ALERT: eateriespk;Black Cartridge HP CC364X;WARNING;SOFT;1;Black Cartridge HP CC364X is at 5% - WARNING!
[1455150602] SERVICE ALERT: eateriespk;Black Cartridge HP CC364X;WARNING;HARD;3;Black Cartridge HP CC364X is at 5% - WARNING!
[1455150636] SERVICE ALERT: eateriespk;Black Cartridge HP CC364X;WARNING;SOFT;2;Black Cartridge HP CC364X is at 5% - WARNING!
[1455150696] SERVICE ALERT: eateriespk;Black Cartridge HP CC364X;WARNING;HARD;3;Black Cartridge HP CC364X is at 5% - WARNING!
[1455157833] SERVICE ALERT: DB01-2012;CPU Load;WARNING;SOFT;1;CPU Load 88% (5 min average)
[1455157899] SERVICE ALERT: DB01-2012;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1455157953] SERVICE ALERT: DB01-2012;CPU Load;CRITICAL;SOFT;2;CPU Load 93% (5 min average)
[1455158019] SERVICE ALERT: DB01-2012;CPU Load;WARNING;SOFT;2;CPU Load 91% (5 min average)
[1455158073] SERVICE ALERT: DB01-2012;CPU Load;CRITICAL;HARD;3;CPU Load 90% (5 min average)
[1455158139] SERVICE ALERT: DB01-2012;CPU Load;WARNING;HARD;3;CPU Load 93% (5 min average)
[1455158673] SERVICE ALERT: DB01-2012;CPU Load;OK;HARD;3;CPU Load 43% (5 min average)
[1455158739] SERVICE ALERT: DB01-2012;CPU Load;OK;HARD;3;CPU Load 40% (5 min average)
[1455161400] HOST DOWNTIME ALERT: REFWIN;STARTED; Host has entered a period of scheduled downtime
[1455161400] HOST DOWNTIME ALERT: MBAMDB;STARTED; Host has entered a period of scheduled downtime
[1455161400] HOST DOWNTIME ALERT: EMUSQLEXPRESS;STARTED; Host has entered a period of scheduled downtime
[1455161400] HOST DOWNTIME ALERT: EMUSCRIPT;STARTED; Host has entered a period of scheduled downtime
[1455161400] HOST DOWNTIME ALERT: MBAM;STARTED; Host has entered a period of scheduled downtime
[1455161400] HOST DOWNTIME ALERT: EMU-VALUE01;STARTED; Host has entered a period of scheduled downtime
[1455161400] HOST DOWNTIME ALERT: HDBACKUPS;STARTED; Host has entered a period of scheduled downtime
[1455161400] HOST DOWNTIME ALERT: AD-RMAD;STARTED; Host has entered a period of scheduled downtime
[1455161400] HOST DOWNTIME ALERT: OFFICEKMS;STARTED; Host has entered a period of scheduled downtime
[1455161400] HOST DOWNTIME ALERT: LANSWEEPER;STARTED; Host has entered a period of scheduled downtime
[1455161400] HOST DOWNTIME ALERT: IMAGES;STARTED; Host has entered a period of scheduled downtime
[1455161400] HOST DOWNTIME ALERT: SANDBOX-DB;STARTED; Host has entered a period of scheduled downtime