[1460694234] SERVICE ALERT: ADWU1;Windows Update Services Event 364;CRITICAL;SOFT;1;Windows Server Update Services, Content file download failed. Reason: File cert verification failure.  Source File: /d/msdownload/update/software/crup/2016/04/windows10.0-kb3154879-x86_c7b4d1ac4829121785ae79b2df7550004eca089f.cab  Destination File: C:\WSUSupdates\WsusContent\9F\C7B4D1AC4829121785AE79B2DF7550004ECA089F.cab
[1460694354] SERVICE ALERT: ADWU1;Windows Update Services Event 364;CRITICAL;SOFT;2;Windows Server Update Services, Content file download failed. Reason: File cert verification failure.  Source File: /d/msdownload/update/software/crup/2016/04/windows10.0-kb3154879-x86_c7b4d1ac4829121785ae79b2df7550004eca089f.cab  Destination File: C:\WSUSupdates\WsusContent\9F\C7B4D1AC4829121785AE79B2DF7550004ECA089F.cab
[1460694474] SERVICE ALERT: ADWU1;Windows Update Services Event 364;CRITICAL;HARD;3;Windows Server Update Services, Content file download failed. Reason: File cert verification failure.  Source File: /d/msdownload/update/software/crup/2016/04/windows10.0-kb3154879-x86_c7b4d1ac4829121785ae79b2df7550004eca089f.cab  Destination File: C:\WSUSupdates\WsusContent\9F\C7B4D1AC4829121785AE79B2DF7550004ECA089F.cab
[1460696654] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 364;CRITICAL;SOFT;1;Windows Server Update Services, eventlog: 1 > critical
[1460696774] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 364;CRITICAL;SOFT;2;Windows Server Update Services(error, 364, success)[Content file download failed.
[1460696895] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 364;CRITICAL;HARD;3;Windows Server Update Services(error, 364, success)[Content file download failed.
[1460699566] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 94% (5 min average)
[1460699686] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 95% (5 min average)
[1460699806] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 82% (5 min average)
[1460700406] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1460700526] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1460700646] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1460701034] SERVICE ALERT: INTLAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460701154] SERVICE ALERT: INTLAPP;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1460701246] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 93% (5 min average)
[1460701674] SERVICE ALERT: ADWU1;Windows Update Services Event 364;OK;HARD;3;OK: No entries found
[1460702446] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 97% (5 min average)
[1460704064] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;CRITICAL;SOFT;1;Windows Server Update Services, The server is failing to download some updates.
[1460704094] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 364;OK;HARD;3;Eventlog check ok
[1460704184] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;CRITICAL;SOFT;2;Windows Server Update Services, The server is failing to download some updates.
[1460704246] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 85% (5 min average)
[1460704246] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1460704304] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;CRITICAL;HARD;3;Windows Server Update Services, The server is failing to download some updates.
[1460704308] SERVICE ALERT: MBAM;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460704428] SERVICE ALERT: MBAM;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1460704846] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1460704966] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 97% (5 min average)
[1460705086] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 79% (5 min average)
[1460705274] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;CRITICAL;SOFT;1;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1460705394] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;CRITICAL;SOFT;2;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1460705514] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;CRITICAL;HARD;3;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1460706886] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1460707006] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 97% (5 min average)
[1460707126] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 80% (5 min average)
[1460707726] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1460707846] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 81% (5 min average)
[1460708453] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 65% (5 min average)
[1460710253] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1460710373] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 87% (5 min average)
[1460710846] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1460710904] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;OK;HARD;3;OK: No entries found
[1460710966] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 58% (5 min average)
[1460712114] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;OK;HARD;3;Eventlog check ok
[1460712772] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1460712892] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 95% (5 min average)
[1460713012] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460713904] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;CRITICAL;SOFT;1;Windows Server Update Services, The server is failing to download some updates.
[1460714024] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;CRITICAL;SOFT;2;Windows Server Update Services, The server is failing to download some updates.
[1460714144] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;CRITICAL;HARD;3;Windows Server Update Services, The server is failing to download some updates.
[1460714185] SERVICE ALERT: OFFICEKMS;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460714212] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1460714234] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;1;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1460714305] SERVICE ALERT: OFFICEKMS;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1460714354] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;2;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1460714425] SERVICE ALERT: OFFICEKMS;CPU Load;OK;SOFT;3;CPU Load 7% (5 min average)
[1460714474] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;HARD;3;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1460714812] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 97% (5 min average)
[1460714812] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1460715166] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1460716012] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 66% (5 min average)
[1460716612] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460716732] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460716852] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 97% (5 min average)
[1460717452] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 87% (5 min average)
[1460718025] SERVICE ALERT: OFFICEKMS;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1460718052] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460718145] SERVICE ALERT: OFFICEKMS;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1460718172] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460718292] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460718892] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 93% (5 min average)
[1460719492] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460720054] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 26% - OK!
[1460720077] SERVICE ALERT: hallefirst3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1460720145] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 90% - OK!
[1460720204] SERVICE ALERT: hallefirst2;Printer Status;OK;HARD;3;stack are straight. 54 6F 20 65 78 69 74 20 70 72 65 73 73 20 1F 2E
[1460720253] SERVICE ALERT: hallefirst3;Maintenance Kit HP 110V-C9152A;CRITICAL;HARD;3;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at CRITICAL level!
[1460720404] SERVICE ALERT: hallefirst3;Printer Status;OK;HARD;3;PERFORM PRINTER MAINTENANCE 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1460720454] SERVICE ALERT: hallefirst2;Maintenance Kit HP 110V-C9152A;OK;HARD;3;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 88% - OK!
[1460720534] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1460720692] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 86% (5 min average)
[1460721074] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;OK;HARD;3;OK: No entries found
[1460721344] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;OK;HARD;3;OK: No entries found
[1460721892] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1460722012] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 88% (5 min average)
[1460722604] SERVICE ALERT: hallefirst2;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460722664] SERVICE ALERT: hallefirst2;Printer Status;OK;SOFT;2;Open right  door.
[1460723212] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1460723332] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460723452] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1460724130] SERVICE ALERT: ADDC5;GPO Replication;UNKNOWN;SOFT;1;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1460724165] SERVICE ALERT: ADDC5;GPO Replication;OK;SOFT;2;Export-CSV : The process cannot access the file 'C:\Program
[1460724652] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 41% (5 min average)
[1460725852] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1460725972] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 95% (5 min average)
[1460726092] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 48% (5 min average)
[1460726514] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;CRITICAL;SOFT;1;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1460726634] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;CRITICAL;SOFT;2;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1460726692] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460726755] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;CRITICAL;HARD;3;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1460726812] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1460726932] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 59% (5 min average)
[1460728404] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (41.0% change >= 20.0% threshold)
[1460728404] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1460728404] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1460728404] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1460728404] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1460728421] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1460728480] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1460728492] SERVICE ALERT: strongpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.130.31. Is the printer off?  Is the network down?
[1460728522] HOST ALERT: strongpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460728541] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1460728552] SERVICE ALERT: strongpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.130.31. Is the printer off?  Is the network down?
[1460728582] HOST ALERT: strongpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460728612] SERVICE ALERT: strongpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.130.31. Is the printer off?  Is the network down?
[1460728629] HOST ALERT: marshallpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460728642] HOST ALERT: strongpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460728691] SERVICE ALERT: hallethirdsouth;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.107.213. Is the printer off?  Is the network down?
[1460728719] HOST ALERT: marshallpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460728732] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460728732] SERVICE ALERT: idm4.emich.edu;LDAP Stats;CRITICAL;SOFT;1;CRITICAL: Failed to connect to LDAP!
[1460728732] HOST ALERT: strongpk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460728751] SERVICE ALERT: hallethirdsouth;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.107.213. Is the printer off?  Is the network down?
[1460728809] HOST ALERT: marshallpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460728810] SERVICE ALERT: strongpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.130.31. Is the printer off?  Is the network down?
[1460728811] SERVICE ALERT: hallethirdsouth;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.107.213. Is the printer off?  Is the network down?
[1460728822] HOST ALERT: strongpk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460728852] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 72% (5 min average)
[1460728852] SERVICE ALERT: idm4.emich.edu;LDAP Stats;CRITICAL;SOFT;2;CRITICAL: Failed to connect to LDAP!
[1460728870] SERVICE ALERT: strongpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.130.31. Is the printer off?  Is the network down?
[1460728873] SERVICE ALERT: idm4.emich.edu;eDirectory Service;CRITICAL;SOFT;1;CRITICAL: Failed to run the ndsd status command successfully!
[1460728899] HOST ALERT: marshallpk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460728900] HOST ALERT: strongpk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1460728930] SERVICE ALERT: strongpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.130.31. Is the printer off?  Is the network down?
[1460728944] SERVICE ALERT: marshallpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.137.101. Is the printer off?  Is the network down?
[1460728960] HOST ALERT: strongpk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1460728972] SERVICE ALERT: idm4.emich.edu;LDAP Stats;CRITICAL;HARD;3;CRITICAL: Failed to connect to LDAP!
[1460728973] SERVICE ALERT: marshallpk;Black Cartridge HP CC364X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.137.101. Is the printer off?  Is the network down?
[1460728973] SERVICE ALERT: strongpk;Black Cartridge HP CC364X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.130.31. Is the printer off?  Is the network down?
[1460728974] HOST ALERT: marshallpk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460728993] SERVICE ALERT: idm4.emich.edu;eDirectory Service;CRITICAL;SOFT;2;CRITICAL: Failed to run the ndsd status command successfully!
[1460729001] SERVICE ALERT: marshallpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.137.101. Is the printer off?  Is the network down?
[1460729004] HOST ALERT: strongpk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1460729004] SERVICE ALERT: marshallpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.137.101. Is the printer off?  Is the network down?
[1460729031] HOST ALERT: marshallpk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1460729033] SERVICE ALERT: strongpk;Black Cartridge HP CC364X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.130.31. Is the printer off?  Is the network down?
[1460729033] SERVICE ALERT: marshallpk;Black Cartridge HP CC364X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.137.101. Is the printer off?  Is the network down?
[1460729061] SERVICE ALERT: marshallpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.137.101. Is the printer off?  Is the network down?
[1460729063] HOST ALERT: strongpk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1460729064] HOST ALERT: marshallpk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1460729064] SERVICE ALERT: marshallpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.137.101. Is the printer off?  Is the network down?
[1460729093] SERVICE ALERT: marshallpk;Black Cartridge HP CC364X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.137.101. Is the printer off?  Is the network down?
[1460729093] SERVICE ALERT: strongpk;Black Cartridge HP CC364X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.130.31. Is the printer off?  Is the network down?
[1460729094] HOST ALERT: marshallpk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1460729113] SERVICE ALERT: idm4.emich.edu;eDirectory Service;CRITICAL;HARD;3;CRITICAL: Failed to run the ndsd status command successfully!
[1460729121] SERVICE ALERT: marshallpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.137.101. Is the printer off?  Is the network down?
[1460729124] HOST ALERT: strongpk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1460729151] HOST ALERT: marshallpk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1460729241] HOST ALERT: marshallpk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1460729426] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460729452] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460729486] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460729546] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460729572] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1460729692] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 90% (5 min average)
[1460730292] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1460730941] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1460731346] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Clearing paper path
[1460731630] SERVICE ALERT: ADDC5;GPO Replication;UNKNOWN;SOFT;1;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1460731666] SERVICE ALERT: ADDC5;GPO Replication;OK;SOFT;2;Export-CSV : The process cannot access the file 'C:\Program
[1460732092] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 46% (5 min average)
[1460732141] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1460732200] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1460732261] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1460732692] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 94% (5 min average)
[1460732812] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460732860] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1460732860] SERVICE FLAPPING ALERT: alexanderpk;Paper Tray Status;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1460732932] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 74% (5 min average)
[1460733954] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;OK;HARD;3;Eventlog check ok
[1460734132] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1460734252] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 97% (5 min average)
[1460734372] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460735546] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460735606] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460735666] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460735744] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;CRITICAL;SOFT;1;Windows Server Update Services, The server is failing to download some updates.
[1460735864] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;CRITICAL;SOFT;2;Windows Server Update Services, The server is failing to download some updates.
[1460735984] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;CRITICAL;HARD;3;Windows Server Update Services, The server is failing to download some updates.
[1460736172] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 94% (5 min average)
[1460736772] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460736867] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Open right  door.
[1460736867] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1460737372] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 63% (5 min average)
[1460739130] SERVICE ALERT: ADDC5;GPO Replication;UNKNOWN;SOFT;1;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1460739167] SERVICE ALERT: ADDC5;GPO Replication;OK;SOFT;2;Export-CSV : The process cannot access the file 'C:\Program
[1460740372] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460740492] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1460740612] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1460741814] SERVICE ALERT: hallegroundcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1460742256] SERVICE ALERT: mckennypk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.152.26. Is the printer off?  Is the network down?
[1460742286] HOST ALERT: mckennypk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460742300] SERVICE ALERT: mckennypk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.152.26. Is the printer off?  Is the network down?
[1460742316] SERVICE ALERT: mckennypk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.152.26. Is the printer off?  Is the network down?
[1460742330] HOST ALERT: mckennypk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460742360] SERVICE ALERT: mckennypk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.152.26. Is the printer off?  Is the network down?
[1460742376] SERVICE ALERT: mckennypk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.152.26. Is the printer off?  Is the network down?
[1460742390] HOST ALERT: mckennypk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460742412] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 48% (5 min average)
[1460742420] SERVICE ALERT: mckennypk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.152.26. Is the printer off?  Is the network down?
[1460742450] HOST ALERT: mckennypk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460742496] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;WARNING;SOFT;1;0
[1460742540] HOST ALERT: mckennypk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460742615] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;WARNING;SOFT;2;0
[1460742630] HOST ALERT: mckennypk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1460742691] SERVICE ALERT: mckennypk;Black Cartridge HP CC364X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.152.26. Is the printer off?  Is the network down?
[1460742720] HOST ALERT: mckennypk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1460742735] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;OK;SOFT;3;0
[1460742751] SERVICE ALERT: mckennypk;Black Cartridge HP CC364X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.152.26. Is the printer off?  Is the network down?
[1460742781] HOST ALERT: mckennypk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1460742811] SERVICE ALERT: mckennypk;Black Cartridge HP CC364X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.152.26. Is the printer off?  Is the network down?
[1460742841] HOST ALERT: mckennypk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1460742931] HOST ALERT: mckennypk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1460743012] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460743132] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460743184] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;OK;HARD;3;OK: No entries found
[1460743252] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1460743514] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors is too high! 4062 found since last run. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1460743634] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1460744095] SERVICE ALERT: idm1.emich.edu;Loopback Driver;WARNING;SOFT;1;0
[1460744216] SERVICE ALERT: idm1.emich.edu;Loopback Driver;OK;SOFT;2;0
[1460744260] SERVICE FLAPPING ALERT: alexanderpk;Paper Tray Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1460744530] HOST ALERT: eateriespk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460744620] HOST ALERT: eateriespk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460744710] HOST ALERT: eateriespk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460744745] SERVICE ALERT: eateriespk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.58.5. Is the printer off?  Is the network down?
[1460744775] HOST ALERT: eateriespk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460744805] SERVICE ALERT: eateriespk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.58.5. Is the printer off?  Is the network down?
[1460744834] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors is too high! 2501 found since last run. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1460744835] HOST ALERT: eateriespk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460744865] SERVICE ALERT: eateriespk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.58.5. Is the printer off?  Is the network down?
[1460744895] HOST ALERT: eateriespk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1460744937] SERVICE ALERT: eateriespk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.58.5. Is the printer off?  Is the network down?
[1460744954] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1460744967] HOST ALERT: eateriespk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1460744982] SERVICE ALERT: eateriespk;Black Cartridge HP CC364X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.58.5. Is the printer off?  Is the network down?
[1460744997] SERVICE ALERT: eateriespk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.58.5. Is the printer off?  Is the network down?
[1460745012] HOST ALERT: eateriespk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1460745042] SERVICE ALERT: eateriespk;Black Cartridge HP CC364X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.58.5. Is the printer off?  Is the network down?
[1460745057] SERVICE ALERT: eateriespk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.58.5. Is the printer off?  Is the network down?
[1460745072] HOST ALERT: eateriespk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1460745102] SERVICE ALERT: eateriespk;Black Cartridge HP CC364X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.58.5. Is the printer off?  Is the network down?
[1460745132] HOST ALERT: eateriespk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1460745154] HOST ALERT: mckennypk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1460745205] SERVICE ALERT: mckennypk;Black Cartridge HP CC364X;OK;HARD;3;Black Cartridge HP CC364X is at 85% - OK!
[1460745370] SERVICE ALERT: mckennypk;Device Status;OK;HARD;3;ORDER MAINTENANCE KIT 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1460745414] SERVICE ALERT: mckennypk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1460745552] SERVICE ALERT: alexanderpk;Black Cartridge HP CC364X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.142.38. Is the printer off?  Is the network down?
[1460745606] SERVICE ALERT: alexanderpk;Black Cartridge HP CC364X;OK;SOFT;2;Black Cartridge HP CC364X is at 92% - OK!
[1460745865] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460745925] SERVICE ALERT: hallefirst2;Device Status;OK;SOFT;2;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1460746630] SERVICE ALERT: ADDC5;GPO Replication;UNKNOWN;SOFT;1;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1460746662] SERVICE ALERT: ADDC5;GPO Replication;OK;SOFT;2;Export-CSV : The process cannot access the file 'C:\Program
[1460747452] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1460747725] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1460748052] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 46% (5 min average)
[1460748266] HOST DOWNTIME ALERT: LANSWEEPER;STARTED; Host has entered a period of scheduled downtime
[1460748354] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;CRITICAL;SOFT;1;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1460748474] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;CRITICAL;SOFT;2;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1460748484] SERVICE ALERT: LANSWEEPER;Service: NetLogin;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1460748594] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;CRITICAL;HARD;3;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1460748604] SERVICE ALERT: LANSWEEPER;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1460748652] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460748772] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460748892] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460750612] SERVICE ALERT: idm4.emich.edu;eDirectory Service;OK;HARD;3;OK: ndsd status is good.
[1460750692] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 77% (5 min average)
[1460751292] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 94% (5 min average)
[1460751412] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460751532] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 84% (5 min average)
[1460751664] SERVICE ALERT: EMUFILES;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460751784] SERVICE ALERT: EMUFILES;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1460753932] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460754052] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1460754172] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460754724] SERVICE ALERT: VCENTERSSO;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460754844] SERVICE ALERT: VCENTERSSO;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1460755046] SERVICE ALERT: porterpk;Black Cartridge HP CC364X;WARNING;SOFT;1;Black Cartridge HP CC364X is at 2% - WARNING!
[1460755106] SERVICE ALERT: porterpk;Black Cartridge HP CC364X;WARNING;SOFT;2;Black Cartridge HP CC364X is at 2% - WARNING!
[1460755166] SERVICE ALERT: porterpk;Black Cartridge HP CC364X;WARNING;HARD;3;Black Cartridge HP CC364X is at 2% - WARNING!
[1460755193] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;OK;HARD;3;Eventlog check ok
[1460755460] HOST DOWNTIME ALERT: LANSWEEPER;STOPPED; Host has exited from a period of scheduled downtime
[1460756379] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460756439] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460756499] SERVICE ALERT: hallefirst1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460756564] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1460756684] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 80% (5 min average)
[1460757099] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;Clearing paper path
[1460757172] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 85% (5 min average)
[1460757584] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;CRITICAL;SOFT;1;Windows Server Update Services, The server is failing to download some updates.
[1460757704] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;CRITICAL;SOFT;2;Windows Server Update Services, The server is failing to download some updates.
[1460757772] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460757824] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;CRITICAL;HARD;3;Windows Server Update Services, The server is failing to download some updates.
[1460757892] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 93% (5 min average)
[1460758012] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1460758612] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 80% (5 min average)
[1460759212] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460759332] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 87% (5 min average)
[1460759932] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1460760052] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 97% (5 min average)
[1460760172] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 95% (5 min average)
[1460760772] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 87% (5 min average)
[1460761372] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460761492] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460761612] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 96% (5 min average)
[1460762812] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 53% (5 min average)
[1460763412] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460763532] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460763652] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 95% (5 min average)
[1460764050] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460764170] SERVICE ALERT: BOEAPPTEST;CPU Load;OK;SOFT;2;CPU Load 3% (5 min average)
[1460764424] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;OK;HARD;3;OK: No entries found
[1460766884] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1460767004] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 94% (5 min average)
[1460767124] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 88% (5 min average)
[1460769313] SERVICE ALERT: IT-2012;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1460769414] SERVICE ALERT: SANDBOX-DB;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1460769433] SERVICE ALERT: IT-2012;CPU Load;OK;SOFT;2;CPU Load 31% (5 min average)
[1460769534] SERVICE ALERT: SANDBOX-DB;CPU Load;OK;SOFT;2;CPU Load 85% (5 min average)
[1460769591] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;CRITICAL;SOFT;1;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1460769652] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 71% (5 min average)
[1460769711] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;CRITICAL;SOFT;2;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1460769831] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;CRITICAL;HARD;3;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1460770252] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460770372] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460770492] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1460771324] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 91% (5 min average)
[1460771444] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 80% (5 min average)
[1460773638] SERVICE ALERT: CBORD;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460773738] SERVICE ALERT: CBORD;Service: RPC;OK;SOFT;2;RpcSs: Started
[1460775644] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460775764] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 80% (5 min average)
[1460776364] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1460776484] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 93% (5 min average)
[1460776604] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 90% (5 min average)
[1460777031] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;OK;HARD;3;Eventlog check ok
[1460777690] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 93% (5 min average)
[1460777804] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 88% (5 min average)
[1460778290] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460778822] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;CRITICAL;SOFT;1;Windows Server Update Services, The server is failing to download some updates.
[1460778942] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;CRITICAL;SOFT;2;Windows Server Update Services, The server is failing to download some updates.
[1460779062] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;CRITICAL;HARD;3;Windows Server Update Services, The server is failing to download some updates.