[1462163018] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 81% (15 min average)
[1462163618] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1462163738] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462163858] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462168051] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1462168171] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1462169699] SERVICE ALERT: ADRODC-LIV;AD Replication;OK;HARD;3;OK: replication is up and running
[1462170571] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1462170691] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1462172491] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1462172611] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 90% (15 min average)
[1462175256] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1462175611] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462175731] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 84% (15 min average)
[1462177056] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 77% (15 min average)
[1462178256] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1462178376] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462178496] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1462179696] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 83% (15 min average)
[1462185096] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1462185216] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462185336] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462185336] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1462188336] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 65% (15 min average)
[1462188936] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462189056] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1462189176] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 82% (15 min average)
[1462189776] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462189896] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462190016] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462191190] SERVICE ALERT: BOEAPP;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462191221] SERVICE ALERT: ADDC11;Service: NetLogin;CRITICAL;SOFT;1;Netlogon: Stopped
[1462191290] SERVICE ALERT: BOEAPP;Service: RPC;OK;SOFT;2;RpcSs: Started
[1462191341] SERVICE ALERT: ADDC11;Service: NetLogin;CRITICAL;SOFT;2;Netlogon: Stopped
[1462191461] SERVICE ALERT: ADDC11;Service: NetLogin;CRITICAL;HARD;3;Netlogon: Stopped
[1462191783] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462191843] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462191903] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462191931] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462191966] SERVICE ALERT: ADDC22;Service: NetLogin;CRITICAL;SOFT;1;Netlogon: Stopped
[1462192023] SERVICE ALERT: ADDC33;Service: NetLogin;CRITICAL;SOFT;1;Netlogon: Stopped
[1462192051] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 44% (15 min average)
[1462192061] SERVICE ALERT: ADDC11;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1462192066] SERVICE ALERT: ADDC44;Service: NetLogin;CRITICAL;SOFT;1;Netlogon: Stopped
[1462192086] SERVICE ALERT: ADDC22;Service: NetLogin;CRITICAL;SOFT;2;Netlogon: Stopped
[1462192118] SERVICE ALERT: ADDC55;Service: NetLogin;CRITICAL;SOFT;1;Netlogon: Stopped
[1462192143] SERVICE ALERT: ADDC33;Service: NetLogin;CRITICAL;SOFT;2;Netlogon: Stopped
[1462192185] SERVICE ALERT: ADDC44;Service: NetLogin;CRITICAL;SOFT;2;Netlogon: Stopped
[1462192206] SERVICE ALERT: ADDC22;Service: NetLogin;CRITICAL;HARD;3;Netlogon: Stopped
[1462192238] SERVICE ALERT: ADDC55;Service: NetLogin;CRITICAL;SOFT;2;Netlogon: Stopped
[1462192263] SERVICE ALERT: ADDC33;Service: NetLogin;CRITICAL;HARD;3;Netlogon: Stopped
[1462192305] SERVICE ALERT: ADDC44;Service: NetLogin;CRITICAL;HARD;3;Netlogon: Stopped
[1462192358] SERVICE ALERT: ADDC55;Service: NetLogin;CRITICAL;HARD;3;Netlogon: Stopped
[1462192398] HOST DOWNTIME ALERT: ADDC11;STARTED; Host has entered a period of scheduled downtime
[1462192415] HOST DOWNTIME ALERT: ADDC22;STARTED; Host has entered a period of scheduled downtime
[1462192416] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 70% (15 min average)
[1462192429] HOST DOWNTIME ALERT: ADDC33;STARTED; Host has entered a period of scheduled downtime
[1462192444] HOST DOWNTIME ALERT: ADDC44;STARTED; Host has entered a period of scheduled downtime
[1462192459] HOST DOWNTIME ALERT: ADDC55;STARTED; Host has entered a period of scheduled downtime
[1462192472] HOST DOWNTIME ALERT: ADDC66;STARTED; Host has entered a period of scheduled downtime
[1462192661] SERVICE ALERT: ADDC11;Service: NetLogin;CRITICAL;SOFT;1;Netlogon: Stopped
[1462192663] SERVICE ALERT: ADDC66;Service: NetLogin;CRITICAL;SOFT;1;Netlogon: Stopped
[1462192781] SERVICE ALERT: ADDC11;Service: NetLogin;CRITICAL;SOFT;2;Netlogon: Stopped
[1462192783] SERVICE ALERT: ADDC66;Service: NetLogin;CRITICAL;SOFT;2;Netlogon: Stopped
[1462192806] SERVICE ALERT: ADDC22;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1462192863] SERVICE ALERT: ADDC33;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1462192901] SERVICE ALERT: ADDC11;Service: NetLogin;CRITICAL;HARD;3;Netlogon: Stopped
[1462192903] SERVICE ALERT: ADDC66;Service: NetLogin;CRITICAL;HARD;3;Netlogon: Stopped
[1462192911] SERVICE ALERT: ADDC44;Service: User Profile;CRITICAL;SOFT;1;connect to address 164.76.7.204 and port 12489: Connection refused
[1462192958] SERVICE ALERT: ADDC55;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1462193028] SERVICE ALERT: ADDC44;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1462193103] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1462193503] SERVICE ALERT: ADDC66;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1462193505] SERVICE ALERT: ADDC44;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1462193721] HOST DOWNTIME ALERT: ADDC66;CANCELLED; Scheduled downtime for host has been cancelled.
[1462193728] HOST DOWNTIME ALERT: ADDC55;CANCELLED; Scheduled downtime for host has been cancelled.
[1462193733] HOST DOWNTIME ALERT: ADDC44;CANCELLED; Scheduled downtime for host has been cancelled.
[1462193741] HOST DOWNTIME ALERT: ADDC33;CANCELLED; Scheduled downtime for host has been cancelled.
[1462193746] HOST DOWNTIME ALERT: ADDC22;CANCELLED; Scheduled downtime for host has been cancelled.
[1462193750] HOST DOWNTIME ALERT: ADDC11;CANCELLED; Scheduled downtime for host has been cancelled.
[1462194002] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1462194003] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1462194003] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1462194003] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1462194003] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1462194089] SERVICE ALERT: ADDC44;Service: Dell Backup Agent;UNKNOWN;SOFT;1;Failed to open service ErdAgent: 1060: The specified service does not exist as an installed service.
[1462194097] SERVICE ALERT: ADDC11;AD Replication;WARNING;SOFT;1;C:\NagiosMonitoring\NagiosMonitoring_AD_ReplicationCheck.ps1 : The term
[1462194102] SERVICE ALERT: ADDC11;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1462194102] SERVICE FLAPPING ALERT: ADDC11;Service: NetLogin;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1462194175] SERVICE ALERT: ADDC11;Service: Dell Backup Agent;UNKNOWN;SOFT;1;Failed to open service ErdAgent: 1060: The specified service does not exist as an installed service.
[1462194184] SERVICE ALERT: ADDC55;AD Replication;WARNING;SOFT;1;C:\NagiosMonitoring\NagiosMonitoring_AD_ReplicationCheck.ps1 : The term
[1462194197] SERVICE ALERT: ADDC44;AD Replication;WARNING;SOFT;1;C:\NagiosMonitoring\NagiosMonitoring_AD_ReplicationCheck.ps1 : The term
[1462194208] SERVICE ALERT: ADDC11;AD Replication;WARNING;SOFT;2;C:\NagiosMonitoring\NagiosMonitoring_AD_ReplicationCheck.ps1 : The term
[1462194209] SERVICE ALERT: ADDC44;Service: Dell Backup Agent;UNKNOWN;SOFT;2;Failed to open service ErdAgent: 1060: The specified service does not exist as an installed service.
[1462194262] SERVICE ALERT: ADDC55;Service: Dell Backup Agent;UNKNOWN;SOFT;1;Failed to open service ErdAgent: 1060: The specified service does not exist as an installed service.
[1462194270] SERVICE ALERT: ADDC22;AD Replication;WARNING;SOFT;1;C:\NagiosMonitoring\NagiosMonitoring_AD_ReplicationCheck.ps1 : The term
[1462194295] SERVICE ALERT: ADDC11;Service: Dell Backup Agent;UNKNOWN;SOFT;2;Failed to open service ErdAgent: 1060: The specified service does not exist as an installed service.
[1462194296] SERVICE ALERT: ADDC55;AD Replication;WARNING;SOFT;2;C:\NagiosMonitoring\NagiosMonitoring_AD_ReplicationCheck.ps1 : The term
[1462194311] SERVICE ALERT: ADDC44;AD Replication;WARNING;SOFT;2;C:\NagiosMonitoring\NagiosMonitoring_AD_ReplicationCheck.ps1 : The term
[1462194328] SERVICE ALERT: ADDC11;AD Replication;WARNING;HARD;3;C:\NagiosMonitoring\NagiosMonitoring_AD_ReplicationCheck.ps1 : The term
[1462194329] SERVICE ALERT: ADDC44;Service: Dell Backup Agent;UNKNOWN;HARD;3;Failed to open service ErdAgent: 1060: The specified service does not exist as an installed service.
[1462194349] SERVICE ALERT: ADDC22;Service: Dell Backup Agent;UNKNOWN;SOFT;1;Failed to open service ErdAgent: 1060: The specified service does not exist as an installed service.
[1462194382] SERVICE ALERT: ADDC55;Service: Dell Backup Agent;UNKNOWN;SOFT;2;Failed to open service ErdAgent: 1060: The specified service does not exist as an installed service.
[1462194383] SERVICE ALERT: ADDC22;AD Replication;WARNING;SOFT;2;C:\NagiosMonitoring\NagiosMonitoring_AD_ReplicationCheck.ps1 : The term
[1462194415] SERVICE ALERT: ADDC55;AD Replication;WARNING;HARD;3;C:\NagiosMonitoring\NagiosMonitoring_AD_ReplicationCheck.ps1 : The term
[1462194415] SERVICE ALERT: ADDC11;Service: Dell Backup Agent;UNKNOWN;HARD;3;Failed to open service ErdAgent: 1060: The specified service does not exist as an installed service.
[1462194430] SERVICE ALERT: ADDC44;AD Replication;WARNING;HARD;3;C:\NagiosMonitoring\NagiosMonitoring_AD_ReplicationCheck.ps1 : The term
[1462194444] SERVICE ALERT: ADDC33;AD Replication;WARNING;SOFT;1;C:\NagiosMonitoring\NagiosMonitoring_AD_ReplicationCheck.ps1 : The term
[1462194469] SERVICE ALERT: ADDC22;Service: Dell Backup Agent;UNKNOWN;SOFT;2;Failed to open service ErdAgent: 1060: The specified service does not exist as an installed service.
[1462194502] SERVICE ALERT: ADDC22;AD Replication;WARNING;HARD;3;C:\NagiosMonitoring\NagiosMonitoring_AD_ReplicationCheck.ps1 : The term
[1462194502] SERVICE ALERT: ADDC55;Service: Dell Backup Agent;UNKNOWN;HARD;3;Failed to open service ErdAgent: 1060: The specified service does not exist as an installed service.
[1462194522] SERVICE ALERT: ADDC33;Service: Dell Backup Agent;UNKNOWN;SOFT;1;Failed to open service ErdAgent: 1060: The specified service does not exist as an installed service.
[1462194555] SERVICE ALERT: ADDC33;AD Replication;WARNING;SOFT;2;C:\NagiosMonitoring\NagiosMonitoring_AD_ReplicationCheck.ps1 : The term
[1462194589] SERVICE ALERT: ADDC22;Service: Dell Backup Agent;UNKNOWN;HARD;3;Failed to open service ErdAgent: 1060: The specified service does not exist as an installed service.
[1462194642] SERVICE ALERT: ADDC33;Service: Dell Backup Agent;UNKNOWN;SOFT;2;Failed to open service ErdAgent: 1060: The specified service does not exist as an installed service.
[1462194656] SERVICE DOWNTIME ALERT: NETIDADPROXY;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1462194676] SERVICE ALERT: ADDC33;AD Replication;WARNING;HARD;3;C:\NagiosMonitoring\NagiosMonitoring_AD_ReplicationCheck.ps1 : The term
[1462194762] SERVICE ALERT: ADDC33;Service: Dell Backup Agent;UNKNOWN;HARD;3;Failed to open service ErdAgent: 1060: The specified service does not exist as an installed service.
[1462196130] SERVICE ALERT: ADDC11;AD Replication;UNKNOWN;HARD;3;CHECK_NRPE: Invalid packet type received from server.
[1462196221] SERVICE ALERT: ADDC55;AD Replication;OK;HARD;3;OK: replication is up and running
[1462196236] SERVICE ALERT: ADDC44;AD Replication;OK;HARD;3;OK: replication is up and running
[1462196252] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1462196308] SERVICE ALERT: ADDC22;AD Replication;OK;HARD;3;OK: replication is up and running
[1462196371] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1462196389] SERVICE ALERT: ADDC22;Service: Dell Backup Agent;OK;HARD;3;ErdAgent: Started
[1462196481] SERVICE ALERT: ADDC33;AD Replication;OK;HARD;3;OK: replication is up and running
[1462196491] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1462196514] SERVICE ALERT: ADDC11;Service: Dell Backup Agent;OK;HARD;3;ErdAgent: Started
[1462196562] SERVICE ALERT: ADDC33;Service: Dell Backup Agent;OK;HARD;3;ErdAgent: Started
[1462196729] SERVICE ALERT: ADDC44;Service: Dell Backup Agent;OK;HARD;3;ErdAgent: Started
[1462196902] SERVICE ALERT: ADDC55;Service: Dell Backup Agent;OK;HARD;3;ErdAgent: Started
[1462196923] SERVICE ALERT: ADDC22;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.7.202 and port 12489: Connection refused
[1462196932] SERVICE ALERT: ADDC22;AD Replication;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 30 seconds.
[1462196989] SERVICE ALERT: ADDC22;Service: Dell Backup Agent;CRITICAL;SOFT;1;ErdAgent: Starting
[1462197003] SERVICE ALERT: ADDC33;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1462197031] SERVICE ALERT: ADDC22;AD Replication;OK;SOFT;2;OK: replication is up and running
[1462197038] SERVICE ALERT: ADDC44;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1462197040] SERVICE ALERT: ADDC22;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1462197092] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462197099] SERVICE ALERT: ADDC55;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Starting
[1462197103] SERVICE ALERT: ADDC66;Service: NetLogin;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1462197109] SERVICE ALERT: ADDC22;Service: Dell Backup Agent;OK;SOFT;2;ErdAgent: Started
[1462197123] SERVICE ALERT: ADDC33;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1462197139] SERVICE ALERT: ADDC66;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1462197158] SERVICE ALERT: ADDC44;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1462197216] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1462197219] SERVICE ALERT: ADDC55;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1462197223] SERVICE ALERT: ADDC66;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1462197259] SERVICE ALERT: ADDC66;Service: Windows Time;CRITICAL;SOFT;2;W32Time: Stopped
[1462197336] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462197364] SERVICE ALERT: ADDC5;AD Replication;WARNING;SOFT;1;WARNING: Replication with ADDC22 failed!
[1462197366] SERVICE ALERT: ADDC55;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.7.205 and port 12489: Connection refused
[1462197379] SERVICE ALERT: ADDC44;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.7.204 and port 12489: Connection refused
[1462197379] SERVICE ALERT: ADDC66;Service: Windows Time;OK;SOFT;3;W32Time: Started
[1462197412] SERVICE ALERT: ADDC11;AD Replication;WARNING;HARD;3;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1462197447] SERVICE ALERT: ADDC4;AD Replication;WARNING;SOFT;1;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1462197456] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1462197483] SERVICE ALERT: ADDC22;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462197484] SERVICE ALERT: ADDC5;AD Replication;WARNING;SOFT;2;WARNING: Replication with ADDC11 failed!
[1462197485] SERVICE ALERT: ADDC6;AD Replication;WARNING;SOFT;1;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1462197486] SERVICE ALERT: ADDC55;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1462197490] SERVICE ALERT: ADDC44;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1462197502] SERVICE ALERT: ADDC55;Service: Dell Backup Agent;CRITICAL;SOFT;1;ErdAgent: Starting
[1462197527] SERVICE ALERT: ADDC66;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.7.206 and port 12489: Connection refused
[1462197572] SERVICE ALERT: ADDC4;AD Replication;OK;SOFT;2;OK: replication is up and running
[1462197583] SERVICE ALERT: ADDC22;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1462197604] SERVICE ALERT: ADDC5;AD Replication;WARNING;HARD;3;WARNING: Replication with ADDC22 failed!
[1462197608] SERVICE ALERT: ADDC6;AD Replication;OK;SOFT;2;OK: replication is up and running
[1462197622] SERVICE ALERT: ADDC55;Service: Dell Backup Agent;OK;SOFT;2;ErdAgent: Started
[1462197638] SERVICE ALERT: ADDC66;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1462198056] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 82% (15 min average)
[1462198291] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 85% (15 min average)
[1462198597] SERVICE ALERT: BOEAPP;Memory Usage;WARNING;SOFT;1;Memory usage: total:18429.84 MB - used: 16602.69 MB (90%) - free: 1827.14 MB (10%)
[1462198717] SERVICE ALERT: BOEAPP;Memory Usage;WARNING;SOFT;2;Memory usage: total:18429.84 MB - used: 16606.47 MB (90%) - free: 1823.37 MB (10%)
[1462198837] SERVICE ALERT: BOEAPP;Memory Usage;WARNING;HARD;3;Memory usage: total:18429.84 MB - used: 16625.91 MB (90%) - free: 1803.93 MB (10%)
[1462198891] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1462199011] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1462199131] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462199131] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1462199856] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462199976] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1462200096] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462200806] SERVICE ALERT: ADDC5;AD Replication;OK;HARD;3;OK: replication is up and running
[1462201009] SERVICE ALERT: ADDC11;AD Replication;UNKNOWN;HARD;3;CHECK_NRPE: Invalid packet type received from server.
[1462201009] SERVICE FLAPPING ALERT: ADDC11;AD Replication;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1462201531] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462201837] SERVICE ALERT: BOEAPP;Memory Usage;OK;HARD;3;Memory usage: total:18684.43 MB - used: 16484.51 MB (88%) - free: 2199.92 MB (12%)
[1462202107] SERVICE ALERT: BOEAPP;Service: WMI;WARNING;SOFT;1;No data was received from host!
[1462202211] SERVICE ALERT: BOEAPP;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1462202731] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462203096] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462203696] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462204896] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 34% (15 min average)
[1462204900] SERVICE FLAPPING ALERT: ADDC11;Service: NetLogin;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1462205132] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462205496] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462205616] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462205736] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1462206331] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462207536] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 67% (15 min average)
[1462208129] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 87% (15 min average)
[1462208403] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STARTED; Service appears to have started flapping (43.6% change >= 20.0% threshold)
[1462208403] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (40.7% change >= 20.0% threshold)
[1462208403] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1462208403] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1462208403] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1462208403] SERVICE DOWNTIME ALERT: NETIDADPROXY;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1462208403] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1462208736] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462208856] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1462208976] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462209330] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462209449] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 83% (15 min average)
[1462209576] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 59% (15 min average)
[1462210176] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462210296] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462210416] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462211701] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462211761] SERVICE ALERT: hallefirst2;Device Status;OK;SOFT;2;Warming up
[1462212449] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1462212568] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 88% (15 min average)
[1462213168] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1462213290] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1462213408] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1462214008] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1462214160] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462214220] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462214280] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462214616] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462215208] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1462215216] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462215328] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 77% (15 min average)
[1462216081] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Ready
[1462217016] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1462217402] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1462217402] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (43.8% change >= 20.0% threshold)
[1462217403] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1462217403] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1462217403] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1462217403] SERVICE DOWNTIME ALERT: NETIDADPROXY;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1462217403] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1462217479] SERVICE ALERT: ADDC11;Service: DFS Namespace;CRITICAL;SOFT;1;connect to address 164.76.7.201 and port 12489: Connection refused
[1462217513] SERVICE ALERT: ADDC11;Service: Dell Backup Agent;CRITICAL;SOFT;1;ErdAgent: Starting
[1462217599] SERVICE ALERT: ADDC11;Service: DFS Namespace;OK;SOFT;2;Dfs: Started
[1462217616] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462217617] SERVICE ALERT: ADDC11;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1462217633] SERVICE ALERT: ADDC11;Service: Dell Backup Agent;OK;SOFT;2;ErdAgent: Started
[1462217635] SERVICE ALERT: ADDC6;AD Replication;WARNING;SOFT;1;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1462217736] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1462217737] SERVICE ALERT: ADDC11;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1462217758] SERVICE ALERT: ADDC6;AD Replication;OK;SOFT;2;OK: replication is up and running
[1462217856] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462218328] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462218448] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1462220248] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1462220368] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 91% (15 min average)
[1462221456] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 81% (15 min average)
[1462222056] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462222176] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462222296] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462223968] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1462224088] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1462224096] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462224208] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 74% (15 min average)
[1462224696] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462225296] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 86% (15 min average)
[1462225896] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1462226008] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1462226016] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 62% (15 min average)
[1462226616] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1462226736] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462226856] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462229256] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1462229856] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462229976] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462230096] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462230696] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1462231296] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462231416] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1462231536] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1462233336] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 82% (15 min average)
[1462234536] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1462234656] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1462234776] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1462237756] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;SOFT;1;e:\ - total: 499.87 Gb - used: 407.88 Gb (82%) - free 91.99 Gb (18%)
[1462237876] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;SOFT;2;e:\ - total: 499.87 Gb - used: 407.88 Gb (82%) - free 91.99 Gb (18%)
[1462237996] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;HARD;3;e:\ - total: 499.87 Gb - used: 407.88 Gb (82%) - free 91.99 Gb (18%)
[1462238376] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462238976] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462239796] SERVICE ALERT: AD-RMAD;E:\ Drive Space;OK;HARD;3;e:\ - total: 499.87 Gb - used: 397.98 Gb (80%) - free 101.90 Gb (20%)
[1462240615] SERVICE ALERT: SQL1;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462240715] SERVICE ALERT: SQL1;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1462241043] SERVICE ALERT: SQL1;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462241162] SERVICE ALERT: SQL1;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1462241532] SERVICE ALERT: SQL1;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462241632] SERVICE ALERT: SQL1;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1462241762] SERVICE ALERT: SQL1;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1462241816] SERVICE ALERT: PAPERCUT;Uptime;CRITICAL;SOFT;1;connect to address 164.76.102.17 and port 12489: No route to host
[1462241819] HOST ALERT: PAPERCUT;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.17)
[1462241825] SERVICE ALERT: PAPERCUT;Service: User Profile;CRITICAL;SOFT;1;connect to address 164.76.102.17 and port 12489: No route to host
[1462241872] HOST ALERT: PAPERCUT;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1462241882] SERVICE ALERT: SQL1;CPU Load;OK;SOFT;2;CPU Load 3% (15 min average)
[1462241923] SERVICE ALERT: PAPERCUT;Uptime;CRITICAL;SOFT;2;connect to address 164.76.102.17 and port 12489: Connection refused
[1462241942] SERVICE ALERT: PAPERCUT;Service: User Profile;CRITICAL;SOFT;2;connect to address 164.76.102.17 and port 12489: Connection refused
[1462241953] HOST ALERT: PAPERCUT;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462242013] HOST ALERT: PAPERCUT;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1462242013] HOST FLAPPING ALERT: PAPERCUT;STARTED; Host appears to have started flapping (23.0% change > 20.0% threshold)
[1462242042] SERVICE ALERT: PAPERCUT;Uptime;OK;SOFT;3;System Uptime - 0 day(s) 0 hour(s) 2 minute(s)
[1462242062] SERVICE ALERT: PAPERCUT;Service: User Profile;OK;SOFT;3;ProfSvc: Started
[1462242642] HOST FLAPPING ALERT: PAPERCUT;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1462244608] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462244728] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 90% (15 min average)
[1462245328] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1462245448] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1462245569] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462246168] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 90% (15 min average)
[1462246768] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462246888] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 74% (15 min average)