[1458706241] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 76% (5 min average)
[1458706841] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1458706961] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1458707081] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1458707681] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1458708281] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1458708281] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1458713641] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 84.90 Gb - used: 78.82 Gb (93%) - free 6.08 Gb (7%)
[1458713761] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 84.90 Gb - used: 78.82 Gb (93%) - free 6.08 Gb (7%)
[1458713881] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 84.90 Gb - used: 78.82 Gb (93%) - free 6.08 Gb (7%)
[1458714480] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1458714600] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 75% (5 min average)
[1458715200] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1458715320] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 95% (5 min average)
[1458715440] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 93% (5 min average)
[1458717240] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 79% (5 min average)
[1458719681] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1458720840] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1458720960] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 56% (5 min average)
[1458723281] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 77% (5 min average)
[1458726281] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1458726401] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 70% (5 min average)
[1458727601] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458727721] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1458727841] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1458729374] SERVICE ALERT: BBTS;Service: Group Policy;CRITICAL;SOFT;1;gpsvc: Stopped
[1458729399] SERVICE ALERT: BBTS_DB;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458729494] SERVICE ALERT: BBTS;Service: Group Policy;CRITICAL;SOFT;2;gpsvc: Stopped
[1458729499] SERVICE ALERT: BBTS_DB;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1458729614] SERVICE ALERT: BBTS;Service: Group Policy;CRITICAL;HARD;3;gpsvc: Stopped
[1458731441] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 92% (5 min average)
[1458732041] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 85% (5 min average)
[1458732041] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1458732641] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458732761] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 69% (5 min average)
[1458733027] SERVICE ALERT: BBTS;Service: NetLogin;CRITICAL;SOFT;1;connect to address 164.76.227.3 and port 12489: Connection refused
[1458733144] SERVICE ALERT: BBTS;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1458733214] SERVICE ALERT: BBTS;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1458733342] HOST ALERT: halleground2;UP;HARD;10;PING OK - Packet loss = 50%, RTA = 0.77 ms
[1458733361] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458733481] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 81% (5 min average)
[1458733601] SERVICE ALERT: halleground2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1458733927] SERVICE ALERT: halleground2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1458733937] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1458734081] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1458734082] SERVICE ALERT: marshallpk;Black Cartridge HP CC364X;OK;HARD;3;Black Cartridge HP CC364X is at 100% - OK!
[1458734201] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1458734319] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1458734484] SERVICE DOWNTIME ALERT: BOEAPPTEST;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1458736718] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 18% (5 min average)
[1458737802] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (26.2% change >= 20.0% threshold)
[1458737802] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1458737802] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1458737802] HOST DOWNTIME ALERT: MAGICDBPROD;STARTED; Host has entered a period of scheduled downtime
[1458737802] SERVICE DOWNTIME ALERT: BOEAPPTEST;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1458737802] HOST DOWNTIME ALERT: MAGICAPPTEST;STARTED; Host has entered a period of scheduled downtime
[1458737829] SERVICE ALERT: idmrbstest.emich.edu;Banner Driver;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1458737927] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (25.6% change >= 20.0% threshold)
[1458737927] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1458737927] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1458737927] HOST DOWNTIME ALERT: MAGICDBPROD;STARTED; Host has entered a period of scheduled downtime
[1458737927] SERVICE DOWNTIME ALERT: BOEAPPTEST;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1458737927] HOST DOWNTIME ALERT: MAGICAPPTEST;STARTED; Host has entered a period of scheduled downtime
[1458737999] SERVICE ALERT: idmrbstest.emich.edu;Banner Driver;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 60 seconds.
[1458737999] SERVICE ALERT: idmrbstest.emich.edu;Banner Driver;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 60 seconds.
[1458738069] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1458738130] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1458738190] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1458739118] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1458739238] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1458739358] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1458739958] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 63% (5 min average)
[1458740101] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (36.0% change >= 20.0% threshold)
[1458740101] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1458740101] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1458740101] HOST DOWNTIME ALERT: MAGICDBPROD;STARTED; Host has entered a period of scheduled downtime
[1458740101] SERVICE DOWNTIME ALERT: BOEAPPTEST;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1458740101] HOST DOWNTIME ALERT: MAGICAPPTEST;STARTED; Host has entered a period of scheduled downtime
[1458740279] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;OK;HARD;3;c:\ - total: 84.90 Gb - used: 75.66 Gb (89%) - free 9.24 Gb (11%)
[1458740558] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458740678] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 91% (5 min average)
[1458740798] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1458741497] SERVICE ALERT: idmrbstest.emich.edu;eDirectory Service;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1458741499] HOST ALERT: idmrbstest.emich.edu;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.102)
[1458741563] HOST ALERT: idmrbstest.emich.edu;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.70 ms
[1458741587] SERVICE ALERT: idmrbstest.emich.edu;LDAP Stats;UNKNOWN;SOFT;1;Difference is negative. LDAP recently rebooted.
[1458741607] SERVICE ALERT: idmrbstest.emich.edu;eDirectory Service;OK;SOFT;2;OK: ndsd status is good.
[1458741707] SERVICE ALERT: idmrbstest.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1458741711] SERVICE ALERT: EMUFILES;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1458741831] SERVICE ALERT: NETENG-PA;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1458741831] SERVICE ALERT: EMUFILES;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1458741869] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (36.8% change >= 20.0% threshold)
[1458741869] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1458741869] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1458741869] HOST DOWNTIME ALERT: MAGICDBPROD;STARTED; Host has entered a period of scheduled downtime
[1458741869] SERVICE DOWNTIME ALERT: BOEAPPTEST;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1458741869] HOST DOWNTIME ALERT: MAGICAPPTEST;STARTED; Host has entered a period of scheduled downtime
[1458741916] SERVICE ALERT: idmrbstest.emich.edu;Banner Driver;UNKNOWN;SOFT;1;Incorrect command line arguments supplied
[1458741951] SERVICE ALERT: NETENG-PA;CPU Load;OK;SOFT;2;CPU Load 23% (5 min average)
[1458742036] SERVICE ALERT: idmrbstest.emich.edu;Banner Driver;UNKNOWN;SOFT;2;Incorrect command line arguments supplied
[1458742156] SERVICE ALERT: idmrbstest.emich.edu;Banner Driver;UNKNOWN;HARD;3;Incorrect command line arguments supplied
[1458742226] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (35.9% change >= 20.0% threshold)
[1458742226] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1458742226] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1458742226] HOST DOWNTIME ALERT: MAGICDBPROD;STARTED; Host has entered a period of scheduled downtime
[1458742226] SERVICE DOWNTIME ALERT: BOEAPPTEST;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1458742226] HOST DOWNTIME ALERT: MAGICAPPTEST;STARTED; Host has entered a period of scheduled downtime
[1458742757] SERVICE ALERT: idmrbstest.emich.edu;Banner Driver;OK;HARD;3;OK: cn=Ellucian Banner,cn=driverset1,o=system is running.
[1458744398] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 92% (5 min average)
[1458744789] SERVICE ALERT: crossroadspk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1458745202] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (30.3% change >= 20.0% threshold)
[1458745203] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1458745203] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1458745203] HOST DOWNTIME ALERT: MAGICDBPROD;STARTED; Host has entered a period of scheduled downtime
[1458745203] SERVICE DOWNTIME ALERT: BOEAPPTEST;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1458745203] HOST DOWNTIME ALERT: MAGICAPPTEST;STARTED; Host has entered a period of scheduled downtime
[1458745598] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1458746656] SERVICE ALERT: studentcenterpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1458746716] SERVICE ALERT: studentcenterpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1458746718] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1458746776] SERVICE ALERT: studentcenterpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1458746778] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1458746838] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1458747010] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1458747061] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1458747070] SERVICE ALERT: studentcenterpk;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! 1 of 3 Trays Still Full
[1458747122] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 100% - OK!
[1458747158] SERVICE ALERT: owen101lab3;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1458747218] SERVICE ALERT: owen101lab3;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1458747278] SERVICE ALERT: owen101lab3;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1458747376] SERVICE ALERT: studentcenterpk;Device Status;OK;HARD;3;Clearing paper path
[1458747438] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1458747899] SERVICE ALERT: ADWSUSSERVERS;Service: Windows Time;WARNING;SOFT;1;W32Time: Error
[1458747923] SERVICE ALERT: ADWSUSSERVERS;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458748019] SERVICE ALERT: ADWSUSSERVERS;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1458748023] SERVICE ALERT: ADWSUSSERVERS;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1458749450] SERVICE ALERT: hallegroundcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1458749458] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1458749517] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1458749578] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1458749678] SERVICE ALERT: owen101lab3;Device Status;OK;HARD;3;Warming up
[1458749797] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 82% (5 min average)
[1458749837] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1458749867] HOST ALERT: owen101lab3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1458749897] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1458749927] HOST ALERT: owen101lab3;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1458749957] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1458749987] HOST ALERT: owen101lab3;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1458750077] HOST ALERT: owen101lab3;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1458750094] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1458750124] HOST ALERT: owen101lab3;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1458750154] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1458750178] SERVICE ALERT: eateriespk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1458750184] HOST ALERT: owen101lab3;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1458750214] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1458750244] HOST ALERT: owen101lab3;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1458750284] SERVICE ALERT: owen101lab3;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1458750314] HOST ALERT: owen101lab3;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1458750344] SERVICE ALERT: owen101lab3;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1458750374] HOST ALERT: owen101lab3;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1458750397] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1458750404] SERVICE ALERT: owen101lab3;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1458750434] HOST ALERT: owen101lab3;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1458750517] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1458750637] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1458751638] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1458751699] SERVICE ALERT: hallefirst2;Device Status;OK;SOFT;2;MENUS  RETRIEVE JOB  INFORMATION  PAPER HANDLING
[1458753637] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 93% (5 min average)
[1458754237] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1458754698] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1458754758] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1458754819] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1458755758] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1458755878] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 65% (5 min average)
[1458756618] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;ORDER CARTRIDGE LESS THAN 1700 PAGES 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Processing job from tray 4
[1458756618] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1458756726] SERVICE ALERT: halleground2;Device Status;OK;HARD;3;Ready
[1458757078] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1458757198] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1458757218] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1458757278] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1458757318] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 82% (5 min average)
[1458757338] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1458757803] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (24.3% change >= 20.0% threshold)
[1458757803] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STARTED; Service appears to have started flapping (26.2% change >= 20.0% threshold)
[1458757803] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1458757803] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1458757803] HOST DOWNTIME ALERT: MAGICDBPROD;STARTED; Host has entered a period of scheduled downtime
[1458757803] SERVICE DOWNTIME ALERT: BOEAPPTEST;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1458757803] HOST DOWNTIME ALERT: MAGICAPPTEST;STARTED; Host has entered a period of scheduled downtime
[1458759037] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (5 min average)
[1458759637] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458759757] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1458759877] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1458760318] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1458760438] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 89% (5 min average)
[1458760477] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 88% (5 min average)
[1458761077] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458761197] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1458761317] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1458762138] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Warming up
[1458764007] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 50% - OK!
[1458764011] HOST ALERT: owen101lab3;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.49 ms
[1458764038] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1458764156] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 72% (5 min average)
[1458764234] HOST ALERT: owen101lab3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1458764324] HOST ALERT: owen101lab3;DOWN;SOFT;2;(Host check timed out after 30.00 seconds)
[1458764387] HOST ALERT: owen101lab3;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1458764477] HOST ALERT: owen101lab3;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1458764567] HOST ALERT: owen101lab3;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1458764613] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1458764643] HOST ALERT: owen101lab3;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1458764673] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1458764703] HOST ALERT: owen101lab3;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1458764733] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1458764763] HOST ALERT: owen101lab3;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1458764834] HOST ALERT: owen101lab3;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1458764924] HOST ALERT: owen101lab3;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1458765516] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 94% (5 min average)
[1458765537] SERVICE ALERT: rooseveltpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1458765598] SERVICE ALERT: rooseveltpk;Device Status;OK;SOFT;2;Genuine HP cartridge installed
[1458766116] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1458767113] SERVICE ALERT: ADAUTH-RMAD;E:\ Drive Space;OK;HARD;3;e:\ - total: 199.87 Gb - used: 99.56 Gb (50%) - free 100.31 Gb (50%)
[1458767538] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1458767598] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1458767658] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1458767916] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 65% (5 min average)
[1458768510] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1458768516] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458768518] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1458768570] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1458768577] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 7% - OK!
[1458768630] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1458768636] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 73% (5 min average)
[1458768858] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;ORDER CARTRIDGE LESS THAN 1800 PAGES 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1458769236] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458769356] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1458769476] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1458770756] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 94% (5 min average)
[1458770876] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 66% (5 min average)
[1458771876] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 87% (5 min average)
[1458772476] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458772596] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1458772716] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1458773316] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 80% (5 min average)
[1458773916] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1458774036] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1458774156] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 97% (5 min average)
[1458778369] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1458778489] SERVICE ALERT: BOEAPPTEST;CPU Load;OK;SOFT;2;CPU Load 10% (5 min average)
[1458779658] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1458781755] SERVICE ALERT: IT-2012;CPU Load;WARNING;SOFT;1;CPU Load 91% (5 min average)
[1458781875] SERVICE ALERT: IT-2012;CPU Load;OK;SOFT;2;CPU Load 65% (5 min average)
[1458783289] SERVICE ALERT: BOEAPPTEST;CPU Load;WARNING;SOFT;1;CPU Load 94% (5 min average)
[1458783409] SERVICE ALERT: BOEAPPTEST;CPU Load;OK;SOFT;2;CPU Load 9% (5 min average)
[1458785556] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1458785817] SERVICE ALERT: INTLDB;CPU Load;WARNING;SOFT;1;CPU Load 94% (5 min average)
[1458785937] SERVICE ALERT: INTLDB;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1458786046] SERVICE ALERT: LANSWEEPER;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458786146] SERVICE ALERT: LANSWEEPER;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1458786531] SERVICE ALERT: INTLTESTAPP;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1458786600] HOST DOWNTIME ALERT: PROCARE;STARTED; Host has entered a period of scheduled downtime
[1458786600] HOST DOWNTIME ALERT: ADWSUSCLIENTS;STARTED; Host has entered a period of scheduled downtime
[1458786600] HOST DOWNTIME ALERT: EMUSCRIPT;STARTED; Host has entered a period of scheduled downtime
[1458786600] HOST DOWNTIME ALERT: REFWIN;STARTED; Host has entered a period of scheduled downtime
[1458786600] HOST DOWNTIME ALERT: SANDBOX-DB;STARTED; Host has entered a period of scheduled downtime
[1458786600] HOST DOWNTIME ALERT: EMUSQLEXPRESS;STARTED; Host has entered a period of scheduled downtime
[1458786600] HOST DOWNTIME ALERT: TABLEAU;STARTED; Host has entered a period of scheduled downtime
[1458786600] HOST DOWNTIME ALERT: TI22;STARTED; Host has entered a period of scheduled downtime
[1458786600] HOST DOWNTIME ALERT: ADWU1;STARTED; Host has entered a period of scheduled downtime
[1458786600] HOST DOWNTIME ALERT: Echo-01;STARTED; Host has entered a period of scheduled downtime
[1458786600] HOST DOWNTIME ALERT: TI22-PSYCLINIC;STARTED; Host has entered a period of scheduled downtime
[1458786600] HOST DOWNTIME ALERT: UHSPHARM-2012;STARTED; Host has entered a period of scheduled downtime
[1458786600] HOST DOWNTIME ALERT: FTK;STARTED; Host has entered a period of scheduled downtime
[1458786600] HOST DOWNTIME ALERT: WINKMS;STARTED; Host has entered a period of scheduled downtime
[1458786600] HOST DOWNTIME ALERT: XDM;STARTED; Host has entered a period of scheduled downtime
[1458786600] HOST DOWNTIME ALERT: CBORD;STARTED; Host has entered a period of scheduled downtime
[1458786600] HOST DOWNTIME ALERT: HDBACKUPS;STARTED; Host has entered a period of scheduled downtime
[1458786600] HOST DOWNTIME ALERT: IMAGES;STARTED; Host has entered a period of scheduled downtime
[1458786600] HOST DOWNTIME ALERT: EMU-VALUE01;STARTED; Host has entered a period of scheduled downtime
[1458786600] HOST DOWNTIME ALERT: LANSWEEPER;STARTED; Host has entered a period of scheduled downtime
[1458786600] HOST DOWNTIME ALERT: CEDARON;STARTED; Host has entered a period of scheduled downtime
[1458786600] HOST DOWNTIME ALERT: MBAM;STARTED; Host has entered a period of scheduled downtime
[1458786600] HOST DOWNTIME ALERT: MBAMDB;STARTED; Host has entered a period of scheduled downtime
[1458786600] HOST DOWNTIME ALERT: AD-RMAD;STARTED; Host has entered a period of scheduled downtime
[1458786600] HOST DOWNTIME ALERT: COE-WEBDAV;STARTED; Host has entered a period of scheduled downtime
[1458786600] HOST DOWNTIME ALERT: NESSUS;STARTED; Host has entered a period of scheduled downtime
[1458786600] HOST DOWNTIME ALERT: PNCDB-2012;STARTED; Host has entered a period of scheduled downtime
[1458786600] HOST DOWNTIME ALERT: NETENG-PA;STARTED; Host has entered a period of scheduled downtime
[1458786600] HOST DOWNTIME ALERT: DIRECTACCESS;STARTED; Host has entered a period of scheduled downtime
[1458786600] HOST DOWNTIME ALERT: OFFICEKMS;STARTED; Host has entered a period of scheduled downtime
[1458786600] HOST DOWNTIME ALERT: OPIE;STARTED; Host has entered a period of scheduled downtime
[1458786600] HOST DOWNTIME ALERT: ADAUTH-RMAD;STARTED; Host has entered a period of scheduled downtime
[1458786600] HOST DOWNTIME ALERT: DPS-SECURE02;STARTED; Host has entered a period of scheduled downtime
[1458786600] HOST DOWNTIME ALERT: PNCAPP-2012;STARTED; Host has entered a period of scheduled downtime
[1458786600] HOST DOWNTIME ALERT: PNCWEB-2012;STARTED; Host has entered a period of scheduled downtime
[1458786651] SERVICE ALERT: INTLTESTAPP;CPU Load;OK;SOFT;2;CPU Load 32% (5 min average)
[1458787751] HOST ALERT: ADWSUSSERVERS;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.133)
[1458787772] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10032;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1458787775] HOST ALERT: ADWSUSSERVERS;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.133)
[1458787795] HOST ALERT: ADWSUSSERVERS;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1458787889] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10032;OK;SOFT;2;Eventlog check ok
[1458788876] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1458788996] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 96% (5 min average)
[1458789116] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 93% (5 min average)
[1458789716] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1458790916] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 93% (5 min average)
[1458791516] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 78% (5 min average)
[1458791516] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1458791628] SERVICE ALERT: EMUFILES;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1458791748] SERVICE ALERT: EMUFILES;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)