[1464149050] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 70% (15 min average)
[1464149650] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1464149770] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464149890] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1464150317] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 207.84 Gb (95%) - free 12.06 Gb (5%)
[1464150490] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464151091] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1464151690] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1464152117] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 209.22 Gb (95%) - free 10.68 Gb (5%)
[1464152890] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1464153490] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464154517] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 207.84 Gb (95%) - free 12.06 Gb (5%)
[1464155290] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 91% (15 min average)
[1464155717] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 210.66 Gb (96%) - free 9.24 Gb (4%)
[1464155890] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1464156010] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464156130] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464156730] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1464157930] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464158117] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 208.00 Gb (95%) - free 11.90 Gb (5%)
[1464159317] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 209.95 Gb (95%) - free 9.95 Gb (5%)
[1464159730] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1464159821] SERVICE ALERT: DB01-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1464159941] SERVICE ALERT: DB01-2012;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1464160061] SERVICE ALERT: DB01-2012;CPU Load;OK;SOFT;3;CPU Load 62% (15 min average)
[1464160330] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1464160450] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464160570] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464161261] SERVICE ALERT: DB01-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1464161381] SERVICE ALERT: DB01-2012;CPU Load;OK;SOFT;2;CPU Load 91% (15 min average)
[1464161717] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 207.97 Gb (95%) - free 11.93 Gb (5%)
[1464163806] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464164117] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 209.25 Gb (95%) - free 10.65 Gb (5%)
[1464164406] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 698 found since last run. oneLevelSearchOps OK. searchOps is too high! 719 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 657 found since last run.
[1464164526] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 268 found since last run. oneLevelSearchOps OK. searchOps is too high! 273 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464164646] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464165246] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 300 found since last run. oneLevelSearchOps OK. searchOps is too high! 362 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464165366] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464165966] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 260 found since last run. oneLevelSearchOps OK. searchOps is too high! 282 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464166086] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464166517] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 207.94 Gb (95%) - free 11.96 Gb (5%)
[1464167770] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1464167886] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 366 found since last run. oneLevelSearchOps OK. searchOps is too high! 388 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464168006] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464168317] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 212.89 Gb (97%) - free 7.01 Gb (3%)
[1464168606] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 436 found since last run. oneLevelSearchOps OK. searchOps is too high! 457 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464168726] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464168970] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1464169090] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 87% (15 min average)
[1464169326] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 243 found since last run. oneLevelSearchOps OK. searchOps is too high! 271 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464169446] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464169893] SERVICE ALERT: ADWU1;Service: Firewall;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1464170013] SERVICE ALERT: ADWU1;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1464170046] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 205 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464170166] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464170204] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;1;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1464170324] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;2;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1464170444] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;HARD;3;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1464170717] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 207.94 Gb (95%) - free 11.96 Gb (5%)
[1464170766] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 349 found since last run. oneLevelSearchOps OK. searchOps is too high! 389 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464170886] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464170890] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1464171010] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1464171130] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464171486] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 268 found since last run. oneLevelSearchOps OK. searchOps is too high! 290 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464171606] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464171917] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 211.17 Gb (96%) - free 8.72 Gb (4%)
[1464172206] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 295 found since last run. oneLevelSearchOps OK. searchOps is too high! 325 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464172326] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464172330] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1464172926] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 291 found since last run. oneLevelSearchOps OK. searchOps is too high! 318 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464172930] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (15 min average)
[1464173046] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464173530] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464173646] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 553 found since last run. oneLevelSearchOps OK. searchOps is too high! 575 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464173650] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464173766] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464173770] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464173981] SERVICE ALERT: DB01-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1464174101] SERVICE ALERT: DB01-2012;CPU Load;OK;SOFT;2;CPU Load 1% (15 min average)
[1464174317] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 207.90 Gb (95%) - free 12.00 Gb (5%)
[1464174367] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 313 found since last run. oneLevelSearchOps OK. searchOps is too high! 386 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 520 found since last run.
[1464174485] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464175085] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 311 found since last run. oneLevelSearchOps OK. searchOps is too high! 335 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464175206] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464176003] SERVICE ALERT: marshallpk;Device Status;OK;HARD;3;Clearing paper path
[1464176015] SERVICE ALERT: marshallpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.137.49. Is the printer off?  Is the network down?
[1464176041] HOST ALERT: marshallpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1464176075] SERVICE ALERT: marshallpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.137.49. Is the printer off?  Is the network down?
[1464176094] HOST ALERT: marshallpk;UP;SOFT;2;PING OK - Packet loss = 75%, RTA = 0.73 ms
[1464176117] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 209.82 Gb (95%) - free 10.08 Gb (5%)
[1464176129] SERVICE ALERT: marshallpk;Paper Tray Status;OK;SOFT;3;Printer Paper Is OK! All Trays Still Full
[1464176170] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1464176770] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1464177004] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 306 found since last run. oneLevelSearchOps OK. searchOps is too high! 385 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464177124] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464177644] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;OK;HARD;3;OK: No entries found
[1464177724] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 748 found since last run. oneLevelSearchOps OK. searchOps is too high! 792 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464177844] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464178083] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1464178143] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1464178203] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1464178444] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 2204 found since last run. oneLevelSearchOps OK. searchOps is too high! 2236 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 1025 found since last run.
[1464178516] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 207.89 Gb (95%) - free 12.01 Gb (5%)
[1464178564] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464178803] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 93% - OK!
[1464179164] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 1356 found since last run. oneLevelSearchOps OK. searchOps is too high! 1609 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 810 found since last run.
[1464179253] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1464179284] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 238 found since last run. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464179314] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1464179373] SERVICE ALERT: mckennypk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1464179404] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 217 found since last run. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464179716] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 210.21 Gb (96%) - free 9.69 Gb (4%)
[1464180881] SERVICE ALERT: ADWSUSSERVERS;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464180881] SERVICE ALERT: ADWSUSSERVERS;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464180981] SERVICE ALERT: ADWSUSSERVERS;CPU Load;OK;SOFT;2;CPU Load 5% (15 min average)
[1464180981] SERVICE ALERT: ADWSUSSERVERS;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1464182116] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 207.90 Gb (95%) - free 12.00 Gb (5%)
[1464183131] SERVICE FLAPPING ALERT: BOEAPP;C:\ Drive Space;STARTED; Service appears to have started flapping (30.1% change >= 20.0% threshold)
[1464183131] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (25.9% change >= 20.0% threshold)
[1464183131] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1464183131] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1464183131] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1464183131] SERVICE DOWNTIME ALERT: NETIDADPROXY;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1464183131] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1464183131] SERVICE DOWNTIME ALERT: idm1.emich.edu;AD Auth Driver;STARTED; Service has entered a period of scheduled downtime
[1464183236] SERVICE FLAPPING ALERT: BOEAPP;C:\ Drive Space;STARTED; Service appears to have started flapping (30.1% change >= 20.0% threshold)
[1464183236] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (25.9% change >= 20.0% threshold)
[1464183236] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1464183236] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1464183236] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1464183236] SERVICE DOWNTIME ALERT: NETIDADPROXY;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1464183236] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1464183236] SERVICE DOWNTIME ALERT: idm1.emich.edu;AD Auth Driver;STARTED; Service has entered a period of scheduled downtime
[1464183272] SERVICE FLAPPING ALERT: BOEAPP;C:\ Drive Space;STARTED; Service appears to have started flapping (30.1% change >= 20.0% threshold)
[1464183272] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (25.9% change >= 20.0% threshold)
[1464183272] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1464183272] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1464183272] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1464183272] SERVICE DOWNTIME ALERT: NETIDADPROXY;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1464183272] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1464183316] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 211.01 Gb (96%) - free 8.89 Gb (4%)
[1464183370] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1464183493] SERVICE ALERT: ADDC6;Remote Loader for IDM;OK;HARD;3;dirxml_remote.exe: Running
[1464183970] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464184090] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1464184210] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 66% (15 min average)
[1464184810] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1464184930] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464185050] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1464185650] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 90% (15 min average)
[1464185716] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 207.99 Gb (95%) - free 11.91 Gb (5%)
[1464186250] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464186370] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464186490] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464188398] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1464188716] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 211.05 Gb (96%) - free 8.85 Gb (4%)
[1464188890] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1464189490] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1464190203] SERVICE FLAPPING ALERT: BOEAPP;C:\ Drive Space;STARTED; Service appears to have started flapping (28.4% change >= 20.0% threshold)
[1464190203] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (33.0% change >= 20.0% threshold)
[1464190203] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1464190203] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1464190203] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1464190203] SERVICE DOWNTIME ALERT: NETIDADPROXY;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1464190203] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1464190513] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1464190573] SERVICE ALERT: hallegroundcolor;Device Status;OK;SOFT;2;Checking printer
[1464191116] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 207.98 Gb (95%) - free 11.92 Gb (5%)
[1464192002] SERVICE FLAPPING ALERT: BOEAPP;C:\ Drive Space;STARTED; Service appears to have started flapping (25.2% change >= 20.0% threshold)
[1464192003] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (30.7% change >= 20.0% threshold)
[1464192003] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1464192003] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1464192003] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1464192003] SERVICE DOWNTIME ALERT: NETIDADPROXY;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1464192003] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1464192316] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 210.50 Gb (96%) - free 9.40 Gb (4%)
[1464193306] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 376234 found since last run. oneLevelSearchOps OK. searchOps is too high! 376235 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464193424] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464194023] SERVICE ALERT: halleground1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1464194023] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 293697 found since last run. oneLevelSearchOps OK. searchOps is too high! 293698 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464194083] SERVICE ALERT: halleground1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1464194143] SERVICE ALERT: halleground1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1464194143] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 85457 found since last run. oneLevelSearchOps OK. searchOps is too high! 85458 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464194263] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464194714] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 207.98 Gb (95%) - free 11.92 Gb (5%)
[1464194884] SERVICE ALERT: mckennypk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1464194944] SERVICE ALERT: mckennypk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1464195003] SERVICE ALERT: mckennypk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1464195489] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 84% (15 min average)
[1464195688] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;HARD;3;e:\ - total: 649.87 Gb - used: 543.99 Gb (84%) - free 105.88 Gb (16%)
[1464195914] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 209.52 Gb (95%) - free 10.38 Gb (5%)
[1464196089] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464196209] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 70% (15 min average)
[1464196640] SERVICE ALERT: idm2.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 4317 found since last run.
[1464196760] SERVICE ALERT: idm2.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464196809] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464196929] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464197049] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464197649] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 64% (15 min average)
[1464198249] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464198369] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 85% (15 min average)
[1464198573] SERVICE ALERT: mckennypk;Device Status;OK;HARD;3;Warming up
[1464198603] SERVICE ALERT: mckennypk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1464198914] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 208.00 Gb (95%) - free 11.90 Gb (5%)
[1464198969] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464199089] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 92% (15 min average)
[1464199173] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1464199233] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1464199293] SERVICE ALERT: mckennypk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1464199689] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1464199809] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464199929] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464200714] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 210.57 Gb (96%) - free 9.33 Gb (4%)
[1464201129] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 57% (15 min average)
[1464201342] SERVICE ALERT: halleground1;Device Status;OK;HARD;3;Clearing paper path
[1464201729] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464201849] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464201969] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1464203114] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 208.03 Gb (95%) - free 11.87 Gb (5%)
[1464204314] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 210.15 Gb (96%) - free 9.75 Gb (4%)
[1464204369] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1464204969] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 80% (15 min average)
[1464205213] SERVICE ALERT: INTLAPP-ISSMAP;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 63.90 Gb - used: 59.30 Gb (93%) - free 4.60 Gb (7%)
[1464205288] SERVICE FLAPPING ALERT: BOEAPP;C:\ Drive Space;STARTED; Service appears to have started flapping (30.5% change >= 20.0% threshold)
[1464205288] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (42.5% change >= 20.0% threshold)
[1464205288] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1464205288] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1464205288] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1464205288] SERVICE DOWNTIME ALERT: NETIDADPROXY;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1464205288] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1464205333] SERVICE ALERT: INTLAPP-ISSMAP;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 63.90 Gb - used: 59.30 Gb (93%) - free 4.60 Gb (7%)
[1464205453] SERVICE ALERT: INTLAPP-ISSMAP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 63.90 Gb - used: 59.30 Gb (93%) - free 4.60 Gb (7%)
[1464205569] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464205689] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1464205809] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 84% (15 min average)
[1464206409] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1464206529] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464206649] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 92% (15 min average)
[1464206714] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 208.02 Gb (95%) - free 11.88 Gb (5%)
[1464207169] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1464207249] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1464207369] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464207489] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 93% (15 min average)
[1464207914] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 210.22 Gb (96%) - free 9.68 Gb (4%)
[1464208089] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1464208209] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464208329] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464208893] SERVICE ALERT: mckennypk;Device Status;OK;HARD;3;Ready
[1464208911] SERVICE ALERT: hallesouthcommons;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1464208970] SERVICE ALERT: hallesouthcommons;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1464209031] SERVICE ALERT: hallesouthcommons;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1464210314] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 208.03 Gb (95%) - free 11.87 Gb (5%)
[1464211929] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1464212529] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1464212649] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 92% (15 min average)
[1464213249] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464213312] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 211.28 Gb (96%) - free 8.62 Gb (4%)
[1464213369] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1464213489] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 56% (15 min average)
[1464213596] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 1 Trays Empty
[1464213657] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 1 Trays Empty
[1464213716] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 1 Trays Empty
[1464214089] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464214209] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 74% (15 min average)
[1464214316] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1464214809] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464214929] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1464215049] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464215631] SERVICE ALERT: hallesouthcommons;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1464215649] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 67% (15 min average)
[1464215712] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 208.03 Gb (95%) - free 11.87 Gb (5%)
[1464216249] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1464216312] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 209.22 Gb (95%) - free 10.68 Gb (5%)
[1464216369] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464216489] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1464217689] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1464218289] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1464218409] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1464218529] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464219312] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 207.95 Gb (95%) - free 11.95 Gb (5%)
[1464219729] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1464219912] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 209.31 Gb (95%) - free 10.59 Gb (5%)
[1464220329] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464222128] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1464222312] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 207.96 Gb (95%) - free 11.94 Gb (5%)
[1464222728] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1464222848] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464222968] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1464224168] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1464224712] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 212.91 Gb (97%) - free 6.99 Gb (3%)
[1464224768] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464224888] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1464225008] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1464225124] SERVICE ALERT: hallefirst1;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1464225183] SERVICE ALERT: hallefirst1;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1464225243] SERVICE ALERT: hallefirst1;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1464226512] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 207.96 Gb (95%) - free 11.94 Gb (5%)
[1464226808] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1464227408] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1464227528] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1464227642] SERVICE ALERT: hallefirst1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1464227648] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464227712] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 209.86 Gb (95%) - free 10.04 Gb (5%)
[1464229800] HOST DOWNTIME ALERT: MBAMDB;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: UHSPHARM-2012;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: SANDBOX-DB;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: MBAM;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: REFWIN;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: IMAGES;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: HDBACKUPS;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: TI22-PSYCLINIC;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: R25-ZETA2012;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: FTK;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: Echo-01;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: R25-WEBVIEWER;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: EMUSQLEXPRESS;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: EMUSCRIPT;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: TI22-COE;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: R25-WEBSERVICES;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: EMU-VALUE01;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: DPS-SECURE02;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: R25-INTERFACE;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: DIRECTACCESS;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: COE-WEBDAV;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: TI22;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: R25-DELTA2012;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: CEDARON;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: CBORD;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: R25-ALPHA2012;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: ADWU1;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: ADWSUSCLIENTS;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: TABLEAU;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: PROCARE;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: ADDC5;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: XDM;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: ADDC4;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: PNCWEB-2012;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: ADDC2;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: ADDC1;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: SP2010DB;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: PNCDB-2012;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: ADCA2;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: ADCA1;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: PNCAPP-2012;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: ADAUTHDC1-OLD;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: ADAUTHDC1;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: SP2010;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: OPIE;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: ADAUTH-RMAD;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: ADARC;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: OFFICEKMS;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: AD-RMAD;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: NETENG-PA;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: WINKMS;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: SCCM-PRIM01;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: NESSUS;STARTED; Host has entered a period of scheduled downtime
[1464229800] HOST DOWNTIME ALERT: LANSWEEPER;STARTED; Host has entered a period of scheduled downtime
[1464230112] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 207.96 Gb (95%) - free 11.94 Gb (5%)
[1464231312] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 209.76 Gb (95%) - free 10.14 Gb (5%)
[1464231616] SERVICE ALERT: MBAM;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1464231630] SERVICE ALERT: OFFICEKMS;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1464231736] SERVICE ALERT: MBAM;CPU Load;OK;SOFT;2;CPU Load 85% (15 min average)
[1464231750] SERVICE ALERT: OFFICEKMS;CPU Load;OK;SOFT;2;CPU Load 1% (15 min average)
[1464231781] SERVICE ALERT: SP2010;Service: Group Policy;CRITICAL;SOFT;1;gpsvc: Stopped
[1464231782] SERVICE ALERT: PROCARE;Service: Windows Time;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1464231798] SERVICE ALERT: REFWIN;Service: NetLogin;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1464231825] SERVICE ALERT: ADCA2;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1464231848] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1464231891] SERVICE ALERT: FTK;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464231901] SERVICE ALERT: SP2010;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1464231902] SERVICE ALERT: PROCARE;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1464231914] HOST ALERT: FTK;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1464231918] SERVICE ALERT: REFWIN;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1464231945] SERVICE ALERT: ADCA2;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1464232004] HOST ALERT: FTK;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1464232011] SERVICE ALERT: FTK;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464232041] HOST ALERT: FTK;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1464232042] SERVICE ALERT: FTK;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464232044] SERVICE ALERT: NETENG-PA;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1464232059] SERVICE ALERT: FTK;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464232072] HOST ALERT: FTK;UP;SOFT;4;PING WARNING - Packet loss = 86%, RTA = 0.38 ms
[1464232111] SERVICE ALERT: FTK;Memory Usage;OK;SOFT;3;Memory usage: total:65524.62 MB - used: 19341.83 MB (30%) - free: 46182.79 MB (70%)
[1464232142] SERVICE ALERT: FTK;C:\ Drive Space;OK;SOFT;2;c:\ - total: 408.28 Gb - used: 138.75 Gb (34%) - free 269.52 Gb (66%)
[1464232159] SERVICE ALERT: FTK;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1464232164] SERVICE ALERT: NETENG-PA;CPU Load;OK;SOFT;2;CPU Load 0% (15 min average)
[1464232224] SERVICE ALERT: PNCWEB-2012;Service: Group Policy;CRITICAL;SOFT;1;gpsvc: Stopped
[1464232232] SERVICE ALERT: OFFICEKMS;Service: NetLogin;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1464232236] SERVICE ALERT: PNCAPP-2012;Service: RPC;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1464232244] SERVICE ALERT: PNCAPP-2012;Service: WMI;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1464232244] SERVICE ALERT: PNCAPP-2012;Service: Group Policy;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1464232247] SERVICE ALERT: PNCDB-2012;Service: Firewall;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1464232258] SERVICE ALERT: TI22-COE;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464232271] SERVICE ALERT: AD-RMAD;VMWare Tools;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1464232278] SERVICE ALERT: ADWSUSCLIENTS;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464232278] SERVICE ALERT: ADWSUSCLIENTS;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464232278] SERVICE ALERT: ADWSUSCLIENTS;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464232299] SERVICE ALERT: ADDC1;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464232304] SERVICE ALERT: ADDC5;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464232304] SERVICE ALERT: ADDC5;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464232305] SERVICE ALERT: ADDC5;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464232309] SERVICE ALERT: SANDBOX-DB;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464232310] SERVICE ALERT: SANDBOX-DB;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464232313] SERVICE ALERT: SANDBOX-DB;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464232318] SERVICE ALERT: HDBACKUPS;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464232324] SERVICE ALERT: SCCM-PRIM01;Service: RPC;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1464232327] SERVICE ALERT: R25-DELTA2012;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464232327] SERVICE ALERT: OPIE;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464232328] SERVICE ALERT: R25-DELTA2012;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464232332] SERVICE ALERT: ADAUTHDC1-OLD;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464232344] SERVICE ALERT: PNCWEB-2012;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1464232348] HOST ALERT: HDBACKUPS;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1464232352] SERVICE ALERT: OFFICEKMS;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1464232356] SERVICE ALERT: PNCAPP-2012;Service: RPC;OK;SOFT;2;RpcSs: Started
[1464232357] SERVICE ALERT: TI22-COE;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1464232364] SERVICE ALERT: PNCAPP-2012;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1464232364] SERVICE ALERT: PNCAPP-2012;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1464232367] SERVICE ALERT: PNCDB-2012;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1464232376] HOST ALERT: R25-ALPHA2012;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1464232378] SERVICE ALERT: ADWSUSCLIENTS;VMWare Tools;OK;SOFT;2;VMTools: Started
[1464232378] SERVICE ALERT: ADWSUSCLIENTS;CPU Load;OK;SOFT;2;CPU Load 7% (15 min average)
[1464232378] SERVICE ALERT: ADWSUSCLIENTS;C:\ Drive Space;OK;SOFT;2;c:\ - total: 79.66 Gb - used: 39.96 Gb (50%) - free 39.69 Gb (50%)
[1464232391] SERVICE ALERT: AD-RMAD;VMWare Tools;OK;SOFT;2;VMTools: Started
[1464232398] HOST ALERT: R25-ALPHA2012;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.36 ms
[1464232399] SERVICE ALERT: ADDC1;Service: Dell Backup Agent;CRITICAL;SOFT;1;ErdAgent: Starting
[1464232399] SERVICE ALERT: ADDC1;VMWare Tools;OK;SOFT;2;VMTools: Started
[1464232402] SERVICE ALERT: HDBACKUPS;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464232404] SERVICE ALERT: ADDC5;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1464232404] SERVICE ALERT: ADDC5;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1464232405] SERVICE ALERT: ADDC5;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1464232409] SERVICE ALERT: SANDBOX-DB;Service: RPC;OK;SOFT;2;RpcSs: Started
[1464232410] SERVICE ALERT: SANDBOX-DB;C:\ Drive Space;OK;SOFT;2;c:\ - total: 39.66 Gb - used: 30.39 Gb (77%) - free 9.26 Gb (23%)
[1464232411] SERVICE ALERT: SCCM-PRIM01;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464232412] SERVICE ALERT: SCCM-PRIM01;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464232413] SERVICE ALERT: SANDBOX-DB;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1464232416] SERVICE ALERT: ADDC3;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1464232427] SERVICE ALERT: R25-DELTA2012;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1464232427] SERVICE ALERT: OPIE;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1464232428] SERVICE ALERT: R25-DELTA2012;CPU Load;OK;SOFT;2;CPU Load 0% (15 min average)
[1464232430] SERVICE ALERT: HDBACKUPS;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1464232432] SERVICE ALERT: ADAUTHDC1-OLD;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1464232432] HOST ALERT: HDBACKUPS;UP;SOFT;2;PING WARNING - Packet loss = 86%, RTA = 0.28 ms
[1464232441] HOST ALERT: SCCM-PRIM01;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1464232463] SERVICE ALERT: SCCM-PRIM01;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464232470] SERVICE ALERT: TABLEAU;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464232473] HOST ALERT: SCCM-PRIM01;UP;SOFT;2;PING OK - Packet loss = 54%, RTA = 0.38 ms
[1464232485] SERVICE ALERT: SCCM-PRIM01;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464232486] SERVICE ALERT: SCCM-PRIM01;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464232502] SERVICE ALERT: HDBACKUPS;Service: RPC;OK;SOFT;2;RpcSs: Started
[1464232512] HOST ALERT: EMU-VALUE01;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1464232512] SERVICE ALERT: SCCM-PRIM01;C:\ Drive Space;OK;SOFT;2;c:\ - total: 194.90 Gb - used: 139.05 Gb (71%) - free 55.85 Gb (29%)
[1464232512] SERVICE ALERT: SCCM-PRIM01;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1464232512] HOST ALERT: DPS-SECURE02;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1464232519] SERVICE ALERT: ADDC1;Service: Dell Backup Agent;OK;SOFT;2;ErdAgent: Started
[1464232536] SERVICE ALERT: ADDC3;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1464232563] SERVICE ALERT: SCCM-PRIM01;Service: RPC;OK;SOFT;3;RpcSs: Started
[1464232565] HOST ALERT: DPS-SECURE02;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1464232568] SERVICE ALERT: EMU-VALUE01;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464232585] SERVICE ALERT: SCCM-PRIM01;Memory Usage;OK;SOFT;2;Memory usage: total:32756.34 MB - used: 4680.88 MB (14%) - free: 28075.46 MB (86%)
[1464232586] SERVICE ALERT: SCCM-PRIM01;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 2 minute(s)
[1464232589] SERVICE ALERT: TABLEAU;CPU Load;OK;SOFT;2;CPU Load 9% (15 min average)
[1464232598] HOST ALERT: EMU-VALUE01;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1464232672] SERVICE ALERT: EMU-VALUE01;Memory Usage;OK;SOFT;2;Memory usage: total:9444.43 MB - used: 1884.45 MB (20%) - free: 7559.98 MB (80%)
[1464232688] HOST ALERT: EMU-VALUE01;UP;SOFT;3;PING WARNING - Packet loss = 90%, RTA = 0.37 ms
[1464232742] SERVICE ALERT: ADAUTHDC1;AD Replication;UNKNOWN;SOFT;1;CHECK_NRPE: Invalid packet type received from server.
[1464232861] SERVICE ALERT: ADAUTHDC1;AD Replication;UNKNOWN;SOFT;2;CHECK_NRPE: Invalid packet type received from server.
[1464232981] SERVICE ALERT: ADAUTHDC1;AD Replication;UNKNOWN;HARD;3;CHECK_NRPE: Invalid packet type received from server.
[1464233048] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (15 min average)
[1464233582] SERVICE ALERT: ADAUTHDC1;AD Replication;WARNING;HARD;3;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1464233648] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1464233768] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 92% (15 min average)
[1464234368] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464234488] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464234608] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)