[1460953540] SERVICE ALERT: ADWU1;Windows Update Services Event 364;CRITICAL;SOFT;1;Windows Server Update Services, Content file download failed. Reason: File cert verification failure.  Source File: /d/msdownload/update/software/crup/2016/04/windows10.0-kb3154879-x86_c7b4d1ac4829121785ae79b2df7550004eca089f.cab  Destination File: C:\WSUSupdates\WsusContent\9F\C7B4D1AC4829121785AE79B2DF7550004ECA089F.cab
[1460953660] SERVICE ALERT: ADWU1;Windows Update Services Event 364;CRITICAL;SOFT;2;Windows Server Update Services, Content file download failed. Reason: File cert verification failure.  Source File: /d/msdownload/update/software/crup/2016/04/windows10.0-kb3154879-x86_c7b4d1ac4829121785ae79b2df7550004eca089f.cab  Destination File: C:\WSUSupdates\WsusContent\9F\C7B4D1AC4829121785AE79B2DF7550004ECA089F.cab
[1460953780] SERVICE ALERT: ADWU1;Windows Update Services Event 364;CRITICAL;HARD;3;Windows Server Update Services, Content file download failed. Reason: File cert verification failure.  Source File: /d/msdownload/update/software/crup/2016/04/windows10.0-kb3154879-x86_c7b4d1ac4829121785ae79b2df7550004eca089f.cab  Destination File: C:\WSUSupdates\WsusContent\9F\C7B4D1AC4829121785AE79B2DF7550004ECA089F.cab
[1460954203] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1460954453] SERVICE ALERT: ADWU1;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460954573] SERVICE ALERT: ADWU1;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1460954803] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460954803] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1460955403] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 93% (5 min average)
[1460955971] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 364;CRITICAL;SOFT;1;Windows Server Update Services(error, 364, success)[Content file download failed.
[1460956003] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460956090] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 364;CRITICAL;SOFT;2;Windows Server Update Services(error, 364, success)[Content file download failed.
[1460956210] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 364;CRITICAL;HARD;3;Windows Server Update Services(error, 364, success)[Content file download failed.
[1460957436] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1460957556] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 84% (5 min average)
[1460959356] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460959476] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 94% (5 min average)
[1460959596] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1460959909] SERVICE FLAPPING ALERT: hallegroundcolor;Device Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1460960796] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 80% (5 min average)
[1460960980] SERVICE ALERT: ADWU1;Windows Update Services Event 364;OK;HARD;3;OK: No entries found
[1460961396] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1460961516] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 73% (5 min average)
[1460962116] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 94% (5 min average)
[1460962236] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 93% (5 min average)
[1460962355] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 73% (5 min average)
[1460962404] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460962524] SERVICE ALERT: BOEAPPTEST;CPU Load;OK;SOFT;2;CPU Load 28% (5 min average)
[1460962956] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460963075] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 96% (5 min average)
[1460963195] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 96% (5 min average)
[1460963374] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;CRITICAL;SOFT;1;Windows Server Update Services, The server is failing to download some updates.
[1460963410] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 364;OK;HARD;3;Eventlog check ok
[1460963494] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;CRITICAL;SOFT;2;Windows Server Update Services, The server is failing to download some updates.
[1460963614] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;CRITICAL;HARD;3;Windows Server Update Services, The server is failing to download some updates.
[1460964344] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;CRITICAL;SOFT;1;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1460964395] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 88% (5 min average)
[1460964395] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1460964464] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;CRITICAL;SOFT;2;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1460964584] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;CRITICAL;HARD;3;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1460965595] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1460965715] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 89% (5 min average)
[1460966315] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1460966435] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 64% (5 min average)
[1460967403] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 83% (5 min average)
[1460969802] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460969922] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 95% (5 min average)
[1460970035] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1460970042] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 97% (5 min average)
[1460970155] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 64% (5 min average)
[1460970214] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;OK;HARD;3;OK: No entries found
[1460970642] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 87% (5 min average)
[1460971785] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;OK;HARD;3;Eventlog check ok
[1460972442] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1460972562] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1460972682] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460973282] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1460973754] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1460973874] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 71% (5 min average)
[1460973882] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460975264] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1460975674] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1460976882] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 60% (5 min average)
[1460977482] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1460977602] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1460977722] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460978074] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460978194] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1460978314] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460978914] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 93% (5 min average)
[1460979514] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 67% (5 min average)
[1460979585] HOST ALERT: hallefirst2;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.45 ms
[1460979704] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 37% - OK!
[1460979804] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Ready
[1460979845] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1460980210] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 364;CRITICAL;SOFT;1;Windows Server Update Services(error, 364, success)[Content file download failed.
[1460980330] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 364;CRITICAL;SOFT;2;Windows Server Update Services(error, 364, success)[Content file download failed.
[1460980451] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 364;CRITICAL;HARD;3;Windows Server Update Services(error, 364, success)[Content file download failed.
[1460980714] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460980834] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 68% (5 min average)
[1460981004] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460981064] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460981124] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460981144] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (33.2% change >= 20.0% threshold)
[1460981144] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1460981144] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1460981144] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1460981144] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1460981322] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 87% (15 min average)
[1460981385] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;CRITICAL;SOFT;1;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1460981504] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;CRITICAL;SOFT;2;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1460981624] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;CRITICAL;HARD;3;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1460981724] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Warming up
[1460981922] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1460982042] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 67% (15 min average)
[1460982325] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460982384] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460982444] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460982444] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1460983164] SERVICE ALERT: EVENTS-2012;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1460983284] SERVICE ALERT: EVENTS-2012;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1460983404] SERVICE ALERT: EVENTS-2012;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1460983842] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1460983962] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 55% (15 min average)
[1460984004] SERVICE ALERT: EVENTS-2012;CPU Load;OK;HARD;3;CPU Load 1% (15 min average)
[1460984434] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1460984554] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 67% (15 min average)
[1460984592] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (26.7% change >= 20.0% threshold)
[1460984592] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1460984592] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1460984592] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1460984592] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1460984592] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1460984773] HOST DOWNTIME ALERT: LANSWEEPER;STARTED; Host has entered a period of scheduled downtime
[1460984965] SERVICE ALERT: porterpk;Black Cartridge HP CC364X;OK;HARD;3;Black Cartridge HP CC364X is at 100% - OK!
[1460984994] SERVICE ALERT: LANSWEEPER;Service: NetLogin;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1460985097] HOST DOWNTIME ALERT: IMAGES;STARTED; Host has entered a period of scheduled downtime
[1460985104] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1460985114] SERVICE ALERT: LANSWEEPER;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1460985169] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1460985174] SERVICE ALERT: IMAGES;Service: PXE Boot;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460985204] HOST ALERT: IMAGES;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460985224] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;SOFT;3;Black Print Cartridge HP C8543X is at 85% - OK!
[1460985224] SERVICE ALERT: IMAGES;Service: MDT Monitor;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460985254] HOST ALERT: IMAGES;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460985294] SERVICE ALERT: IMAGES;Service: PXE Boot;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460985298] HOST ALERT: IMAGES;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1460985395] SERVICE ALERT: IMAGES;Service: PXE Boot;CRITICAL;HARD;3;WDSServer: Stopped
[1460985444] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Warming up
[1460985520] SERVICE ALERT: IMAGES;PXE Boot Service;CRITICAL;SOFT;1;WDSServer: Stopped
[1460985640] SERVICE ALERT: IMAGES;PXE Boot Service;CRITICAL;SOFT;2;WDSServer: Stopped
[1460985760] SERVICE ALERT: IMAGES;PXE Boot Service;CRITICAL;HARD;3;WDSServer: Stopped
[1460985762] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1460985804] SERVICE ALERT: IMAGES;Service: MDT Monitor;UNKNOWN;HARD;3;Failed to open service MDT_Monitor: 1060: The specified service does not exist as an installed service.
[1460985882] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1460985994] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460986034] SERVICE ALERT: idm2.emich.edu;LDAP Stats;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1460986054] SERVICE ALERT: hallefirst1;Device Status;OK;SOFT;2;Move left accessory, if attached.
[1460986060] HOST ALERT: idm4.emich.edu;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460986064] HOST ALERT: idm2.emich.edu;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460986074] SERVICE ALERT: idm1.emich.edu;Loopback Driver;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 60 seconds.
[1460986104] HOST ALERT: idm1.emich.edu;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460986114] SERVICE ALERT: idm2.emich.edu;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460986144] SERVICE ALERT: idm1.emich.edu;NSDS Errors;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1460986144] HOST ALERT: idm2.emich.edu;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460986149] SERVICE ALERT: idm1.emich.edu;LDAP Stats;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1460986150] HOST ALERT: idm4.emich.edu;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460986154] SERVICE ALERT: idm2.emich.edu;LDAP Stats;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1460986162] HOST ALERT: idm2.emich.edu;UP;SOFT;3;PING OK - Packet loss = 44%, RTA = 1.11 ms
[1460986174] HOST ALERT: idm1.emich.edu;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460986180] SERVICE ALERT: idm4.emich.edu;NSDS Errors;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1460986191] HOST ALERT: idm4.emich.edu;UP;SOFT;3;PING OK - Packet loss = 58%, RTA = 0.29 ms
[1460986194] SERVICE ALERT: idm1.emich.edu;Loopback Driver;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 60 seconds.
[1460986203] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STARTED; Service appears to have started flapping (27.0% change >= 20.0% threshold)
[1460986203] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1460986203] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1460986203] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1460986203] HOST DOWNTIME ALERT: LANSWEEPER;STARTED; Host has entered a period of scheduled downtime
[1460986203] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1460986203] HOST DOWNTIME ALERT: IMAGES;STARTED; Host has entered a period of scheduled downtime
[1460986217] HOST ALERT: idm1.emich.edu;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1460986228] SERVICE ALERT: idm2.emich.edu;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1460986254] SERVICE ALERT: idm1.emich.edu;NSDS Errors;OK;SOFT;2;OK: 0 Errors Reported
[1460986255] SERVICE ALERT: idm1.emich.edu;Loopback Driver;CRITICAL;HARD;3;Unknown: Error reported?
[1460986259] SERVICE ALERT: idm1.emich.edu;LDAP Stats;CRITICAL;SOFT;2;CRITICAL: Failed to connect to LDAP!
[1460986264] SERVICE ALERT: idm2.emich.edu;LDAP Stats;CRITICAL;HARD;3;CRITICAL: Failed to connect to LDAP!
[1460986290] SERVICE ALERT: idm4.emich.edu;NSDS Errors;OK;SOFT;2;OK:
[1460986324] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;CRITICAL;SOFT;1;CRITICAL: cn=ADAUTH Driver,cn=driverset1,o=system is stopped!
[1460986360] SERVICE ALERT: IMAGES;PXE Boot Service;OK;HARD;3;WDSServer: Started
[1460986379] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1460986430] SERVICE ALERT: idm3.emich.edu;LDAP Stats;UNKNOWN;SOFT;1;Difference is negative. LDAP recently rebooted.
[1460986445] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;CRITICAL;SOFT;2;CRITICAL: cn=ADAUTH Driver,cn=driverset1,o=system is stopped!
[1460986550] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1460986565] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;OK;SOFT;3;0
[1460986594] SERVICE ALERT: IMAGES;Service: PXE Boot;OK;HARD;3;WDSServer: Started
[1460986856] SERVICE ALERT: idm1.emich.edu;Loopback Driver;OK;HARD;3;0
[1460986864] SERVICE ALERT: idm2.emich.edu;LDAP Stats;OK;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1460987544] SERVICE ALERT: idmrbs.emich.edu;NSDS Errors;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1460987604] SERVICE ALERT: idmrbs.emich.edu;LDAP Stats;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1460987674] SERVICE ALERT: idmrbs.emich.edu;NSDS Errors;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1460987682] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1460987704] HOST ALERT: idmrbs.emich.edu;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460987722] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STARTED; Service appears to have started flapping (25.7% change >= 20.0% threshold)
[1460987722] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1460987722] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1460987722] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1460987722] HOST DOWNTIME ALERT: LANSWEEPER;STARTED; Host has entered a period of scheduled downtime
[1460987722] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1460987722] HOST DOWNTIME ALERT: IMAGES;STARTED; Host has entered a period of scheduled downtime
[1460987734] SERVICE ALERT: idmrbs.emich.edu;LDAP Stats;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1460987764] HOST ALERT: idmrbs.emich.edu;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460987794] SERVICE ALERT: idmrbs.emich.edu;NSDS Errors;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1460987794] SERVICE ALERT: strongpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1460987798] HOST ALERT: strongpk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.55 ms
[1460987802] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1460987824] HOST ALERT: idmrbs.emich.edu;DOWN;SOFT;3;(Host check timed out after 30.00 seconds)
[1460987828] HOST ALERT: eateriespk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1460987844] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460987853] SERVICE ALERT: eateriespk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1460987854] SERVICE ALERT: idmrbs.emich.edu;LDAP Stats;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1460987856] SERVICE ALERT: ADAUTHDC2;Service: Key Distribution Center;CRITICAL;SOFT;1;KDC: Stopped
[1460987884] HOST ALERT: idmrbs.emich.edu;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460987885] SERVICE ALERT: ADAUTHDC3;Service: Dell Backup Agent;CRITICAL;SOFT;1;ErdAgent: Starting
[1460987890] SERVICE ALERT: ADAUTHDC1;Service: Intersite Messaging;CRITICAL;SOFT;1;IsmServ: Stopped
[1460987904] SERVICE ALERT: ADAUTHDC3;Service: AD Domain Service;CRITICAL;SOFT;1;NTDS: Stopping
[1460987904] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460987922] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1460987924] SERVICE ALERT: idmrbs.emich.edu;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460987936] SERVICE ALERT: ADAUTHDC1;Service: Key Distribution Center;CRITICAL;SOFT;1;KDC: Stopped
[1460987954] HOST ALERT: idmrbs.emich.edu;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460987964] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460987996] SERVICE ALERT: ADAUTHDC2;Service: Key Distribution Center;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988004] HOST ALERT: marshallpk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.56 ms
[1460988007] SERVICE ALERT: ADAUTHDC3;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988007] SERVICE ALERT: ADAUTHDC3;Service: DFS Namespace;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988018] HOST ALERT: idmrbs.emich.edu;UP;SOFT;6;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1460988025] SERVICE ALERT: ADAUTHDC3;Service: Dell Backup Agent;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988030] SERVICE ALERT: ADAUTHDC1;Service: Intersite Messaging;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988034] SERVICE ALERT: ADAUTHDC3;AD Replication;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 30 seconds.
[1460988034] SERVICE ALERT: eateriespk;Device Status;OK;HARD;3;Sleep Mode on
[1460988038] SERVICE ALERT: idmrbs.emich.edu;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1460988044] SERVICE ALERT: ADAUTHDC3;Service: AD Domain Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988055] SERVICE ALERT: ADAUTHDC1;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988055] SERVICE ALERT: ADAUTHDC2;Service: AD Web Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988064] SERVICE ALERT: ADAUTHDC2;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988064] SERVICE ALERT: ADAUTHDC3;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988064] SERVICE ALERT: ADAUTHDC3;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988076] SERVICE ALERT: ADAUTHDC1;Service: Key Distribution Center;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988091] SERVICE ALERT: idmrbs.emich.edu;Banner Driver;CRITICAL;SOFT;1;Unknown: Error reported?
[1460988094] SERVICE ALERT: ADAUTHDC1;AD Replication;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 30 seconds.
[1460988104] SERVICE ALERT: ADAUTHDC1;Service: AD Domain Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988104] SERVICE ALERT: ADAUTHDC1;Service: Remote Management;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988113] SERVICE ALERT: strongpk;Device Status;OK;HARD;3;Ready
[1460988116] SERVICE ALERT: ADAUTHDC2;Service: Key Distribution Center;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988124] SERVICE ALERT: ADAUTHDC2;Service: Remote Management;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988127] SERVICE ALERT: ADAUTHDC3;Service: DFS Namespace;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988127] SERVICE ALERT: ADAUTHDC3;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988128] SERVICE ALERT: ADAUTHDC1;Service: DFS Namespace;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988128] SERVICE ALERT: ADAUTHDC1;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988134] SERVICE ALERT: ADAUTHDC2;Service: AD Domain Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988134] SERVICE ALERT: ADAUTHDC1;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988144] SERVICE ALERT: ADAUTHDC1;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988145] SERVICE ALERT: ADAUTHDC3;Service: Dell Backup Agent;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988145] SERVICE ALERT: ADAUTHDC2;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988147] SERVICE ALERT: ADAUTHDC3;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988149] SERVICE ALERT: ADAUTHDC1;Service: AD Web Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988150] SERVICE ALERT: ADAUTHDC1;Service: Intersite Messaging;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988154] SERVICE ALERT: ADAUTHDC3;AD Replication;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 30 seconds.
[1460988164] SERVICE ALERT: ADAUTHDC3;Service: AD Domain Service;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988173] SERVICE ALERT: ADAUTHDC3;Service: Distributed File System Replication;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988173] SERVICE ALERT: ADAUTHDC1;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988175] SERVICE ALERT: ADAUTHDC1;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988175] SERVICE ALERT: ADAUTHDC2;Service: AD Web Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988184] SERVICE ALERT: ADAUTHDC2;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988184] SERVICE ALERT: ADAUTHDC3;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988184] SERVICE ALERT: ADAUTHDC3;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988196] SERVICE ALERT: ADAUTHDC1;Service: Key Distribution Center;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988211] SERVICE ALERT: idmrbs.emich.edu;Banner Driver;CRITICAL;SOFT;2;CRITICAL: cn=Ellucian Banner,cn=driverset1,o=system is stopped!
[1460988214] SERVICE ALERT: ADAUTHDC1;AD Replication;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 30 seconds.
[1460988224] SERVICE ALERT: ADAUTHDC1;Service: Remote Management;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988224] SERVICE ALERT: ADAUTHDC1;Service: AD Domain Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988224] SERVICE ALERT: ADAUTHDC2;Service: DFS Namespace;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988224] SERVICE ALERT: ADAUTHDC2;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988224] SERVICE ALERT: ADAUTHDC3;Service: DNS Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988234] SERVICE ALERT: ADAUTHDC1;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988236] SERVICE ALERT: ADAUTHDC2;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988236] SERVICE ALERT: ADAUTHDC1;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988238] SERVICE ALERT: ADAUTHDC3;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988239] SERVICE ALERT: ADAUTHDC3;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988241] SERVICE ALERT: ADAUTHDC2;Service: DNS Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988244] SERVICE ALERT: ADAUTHDC2;Service: Remote Management;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988244] SERVICE ALERT: marshallpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1460988247] SERVICE ALERT: ADAUTHDC3;Service: DFS Namespace;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988247] SERVICE ALERT: ADAUTHDC3;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988248] SERVICE ALERT: ADAUTHDC1;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988248] SERVICE ALERT: ADAUTHDC1;Service: DFS Namespace;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988254] SERVICE ALERT: ADAUTHDC2;Service: AD Domain Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988254] SERVICE ALERT: ADAUTHDC1;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988264] SERVICE ALERT: ADAUTHDC1;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988265] SERVICE ALERT: ADAUTHDC2;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988265] SERVICE ALERT: ADAUTHDC2;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988267] SERVICE ALERT: ADAUTHDC3;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988269] SERVICE ALERT: ADAUTHDC1;Service: AD Web Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988274] SERVICE ALERT: ADAUTHDC3;AD Replication;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 30 seconds.
[1460988281] SERVICE ALERT: ADAUTHDC1;Service: DNS Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988281] SERVICE ALERT: ADAUTHDC1;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988282] SERVICE ALERT: ADAUTHDC1;Service: RPC;UNKNOWN;SOFT;2;Failed to open service manager: 1115: A system shutdown is in progress.
[1460988282] SERVICE ALERT: ADAUTHDC3;Service: Distributed File System Replication;UNKNOWN;SOFT;2;Failed to open service manager: 1115: A system shutdown is in progress.
[1460988284] SERVICE ALERT: ADAUTHDC1;NSClient++ Version;OK;SOFT;3;NSClient++ 0.4.4.15 2015-11-25
[1460988284] SERVICE ALERT: ADAUTHDC3;Service: Windows Time;UNKNOWN;HARD;3;Failed to open service manager: 1115: A system shutdown is in progress.
[1460988284] SERVICE ALERT: ADAUTHDC3;Service: DNS Client;UNKNOWN;HARD;3;Failed to open service manager: 1115: A system shutdown is in progress.
[1460988284] SERVICE ALERT: ADAUTHDC2;Service: User Profile;UNKNOWN;HARD;3;Failed to open service manager: 1115: A system shutdown is in progress.
[1460988284] SERVICE ALERT: ADAUTHDC2;Service: AD Web Service;UNKNOWN;HARD;3;Failed to open service manager: 1115: A system shutdown is in progress.
[1460988294] SERVICE ALERT: marshallpk;Device Status;OK;HARD;3;Ready
[1460988322] SERVICE ALERT: ADAUTHDC1;Service: Dell Backup Agent;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988322] SERVICE ALERT: ADAUTHDC1;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988322] SERVICE ALERT: ADAUTHDC1;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988323] SERVICE ALERT: ADAUTHDC2;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988324] SERVICE ALERT: ADAUTHDC2;Service: Distributed File System Replication;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988324] SERVICE ALERT: ADAUTHDC3;Service: Intersite Messaging;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988328] SERVICE ALERT: ADAUTHDC2;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988328] SERVICE ALERT: ADAUTHDC3;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988331] SERVICE ALERT: idmrbs.emich.edu;Banner Driver;OK;SOFT;3;0
[1460988334] SERVICE ALERT: ADAUTHDC1;AD Replication;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 30 seconds.
[1460988344] SERVICE ALERT: ADAUTHDC1;Service: AD Domain Service;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988344] SERVICE ALERT: ADAUTHDC1;Service: Remote Management;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988344] SERVICE ALERT: ADAUTHDC2;Service: DFS Namespace;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988344] SERVICE ALERT: ADAUTHDC3;Service: DNS Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988344] SERVICE ALERT: ADAUTHDC2;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988354] SERVICE ALERT: ADAUTHDC1;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988354] SERVICE ALERT: ADAUTHDC2;Service: Dell Backup Agent;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988356] SERVICE ALERT: ADAUTHDC2;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988356] SERVICE ALERT: ADAUTHDC1;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988358] SERVICE ALERT: ADAUTHDC3;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988359] SERVICE ALERT: ADAUTHDC3;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988361] SERVICE ALERT: ADAUTHDC2;Service: DNS Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988364] SERVICE ALERT: ADAUTHDC2;Service: Remote Management;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988368] SERVICE ALERT: ADAUTHDC1;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988368] SERVICE ALERT: ADAUTHDC1;Service: DFS Namespace;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988372] SERVICE ALERT: ADAUTHDC1;Service: Distributed File System Replication;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988372] SERVICE ALERT: ADAUTHDC2;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988374] SERVICE ALERT: ADAUTHDC2;Service: AD Domain Service;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988374] SERVICE ALERT: ADAUTHDC1;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988374] SERVICE ALERT: ADAUTHDC3;Service: Key Distribution Center;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988374] SERVICE ALERT: ADAUTHDC3;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988383] SERVICE ALERT: ADAUTHDC2;AD Replication;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 30 seconds.
[1460988384] SERVICE ALERT: ADAUTHDC1;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988384] SERVICE ALERT: idmrbs.emich.edu;NSDS Errors;OK;HARD;3;OK:
[1460988385] SERVICE ALERT: ADAUTHDC2;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988385] SERVICE ALERT: ADAUTHDC2;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988387] SERVICE ALERT: ADAUTHDC3;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988389] SERVICE ALERT: ADAUTHDC1;Service: AD Web Service;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988401] SERVICE ALERT: ADAUTHDC1;Service: DNS Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988401] SERVICE ALERT: ADAUTHDC1;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988413] SERVICE ALERT: ADAUTHDC1;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988413] SERVICE ALERT: ADAUTHDC3;Service: Distributed File System Replication;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988424] SERVICE ALERT: ADAUTHDC2;Service: Intersite Messaging;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988424] SERVICE ALERT: ADAUTHDC3;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988424] SERVICE ALERT: ADAUTHDC2;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988425] SERVICE ALERT: ADAUTHDC1;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988426] SERVICE ALERT: ADAUTHDC3;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988426] SERVICE ALERT: ADAUTHDC3;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988442] SERVICE ALERT: ADAUTHDC1;Service: Dell Backup Agent;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988442] SERVICE ALERT: ADAUTHDC1;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988442] SERVICE ALERT: ADAUTHDC1;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988443] SERVICE ALERT: ADAUTHDC2;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988444] SERVICE ALERT: ADAUTHDC3;Service: Intersite Messaging;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988444] SERVICE ALERT: ADAUTHDC2;Service: Distributed File System Replication;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988444] SERVICE ALERT: idmrbs.emich.edu;LDAP Stats;UNKNOWN;HARD;3;Difference is negative. LDAP recently rebooted.
[1460988448] SERVICE ALERT: ADAUTHDC2;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988448] SERVICE ALERT: ADAUTHDC3;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988464] SERVICE ALERT: ADAUTHDC2;Service: DFS Namespace;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988464] SERVICE ALERT: ADAUTHDC3;Service: DNS Service;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988464] SERVICE ALERT: ADAUTHDC2;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988472] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;OK;HARD;3;Eventlog check ok
[1460988473] SERVICE ALERT: ADAUTHDC3;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988474] SERVICE ALERT: ADAUTHDC1;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988474] SERVICE ALERT: ADAUTHDC2;Service: Dell Backup Agent;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988475] SERVICE ALERT: ADAUTHDC1;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988475] SERVICE ALERT: ADAUTHDC1;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988476] SERVICE ALERT: ADAUTHDC2;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988476] SERVICE ALERT: ADAUTHDC1;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988476] SERVICE ALERT: ADAUTHDC2;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988478] SERVICE ALERT: ADAUTHDC3;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988478] SERVICE ALERT: ADAUTHDC3;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988479] SERVICE ALERT: ADAUTHDC3;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988481] SERVICE ALERT: ADAUTHDC2;Service: DNS Service;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988492] SERVICE ALERT: ADAUTHDC1;Service: Distributed File System Replication;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988492] SERVICE ALERT: ADAUTHDC2;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988494] SERVICE ALERT: ADAUTHDC3;Service: Key Distribution Center;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988494] SERVICE ALERT: ADAUTHDC3;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988503] SERVICE ALERT: ADAUTHDC2;AD Replication;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 30 seconds.
[1460988505] SERVICE ALERT: ADAUTHDC2;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988511] SERVICE ALERT: ADAUTHDC1;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988521] SERVICE ALERT: ADAUTHDC1;Service: DNS Service;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988521] SERVICE ALERT: ADAUTHDC1;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988524] SERVICE ALERT: ADAUTHDC3;Service: Remote Management;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988524] SERVICE ALERT: ADAUTHDC2;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988524] SERVICE ALERT: ADAUTHDC2;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988544] SERVICE ALERT: ADAUTHDC2;Service: Intersite Messaging;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988544] SERVICE ALERT: ADAUTHDC3;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988544] SERVICE ALERT: ADAUTHDC2;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988545] SERVICE ALERT: ADAUTHDC1;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988546] SERVICE ALERT: ADAUTHDC3;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988546] SERVICE ALERT: ADAUTHDC3;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988556] SERVICE ALERT: ADAUTHDC3;Service: AD Web Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988560] SERVICE ALERT: ADAUTHDC2;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988562] SERVICE ALERT: ADAUTHDC1;Service: Dell Backup Agent;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988562] SERVICE ALERT: ADAUTHDC1;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988562] SERVICE ALERT: ADAUTHDC1;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988563] SERVICE ALERT: ADAUTHDC2;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988563] SERVICE ALERT: ADAUTHDC2;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988563] SERVICE ALERT: ADAUTHDC3;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988564] SERVICE ALERT: ADAUTHDC2;Service: Distributed File System Replication;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988564] SERVICE ALERT: ADAUTHDC3;Service: Intersite Messaging;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988565] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Ready
[1460988568] SERVICE ALERT: ADAUTHDC2;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988568] SERVICE ALERT: ADAUTHDC3;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988593] SERVICE ALERT: ADAUTHDC3;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988594] SERVICE ALERT: ADAUTHDC2;Service: Dell Backup Agent;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988595] SERVICE ALERT: ADAUTHDC1;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988595] SERVICE ALERT: ADAUTHDC1;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988596] SERVICE ALERT: ADAUTHDC2;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988598] SERVICE ALERT: ADAUTHDC3;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988612] SERVICE ALERT: ADAUTHDC1;Service: Distributed File System Replication;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988612] SERVICE ALERT: ADAUTHDC2;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988614] SERVICE ALERT: ADAUTHDC3;Service: Key Distribution Center;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988614] SERVICE ALERT: ADAUTHDC3;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988623] SERVICE ALERT: ADAUTHDC2;AD Replication;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 30 seconds.
[1460988631] SERVICE ALERT: ADAUTHDC1;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988644] SERVICE ALERT: ADAUTHDC3;Service: Remote Management;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988644] SERVICE ALERT: ADAUTHDC2;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988644] SERVICE ALERT: ADAUTHDC2;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988664] SERVICE ALERT: ADAUTHDC2;Service: Intersite Messaging;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988664] SERVICE ALERT: ADAUTHDC2;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988664] SERVICE ALERT: ADAUTHDC3;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988665] SERVICE ALERT: ADAUTHDC1;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988666] SERVICE ALERT: ADAUTHDC3;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988666] SERVICE ALERT: ADAUTHDC3;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988676] SERVICE ALERT: ADAUTHDC3;Service: AD Web Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988680] SERVICE ALERT: ADAUTHDC2;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988683] SERVICE ALERT: ADAUTHDC2;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988683] SERVICE ALERT: ADAUTHDC3;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460988713] SERVICE ALERT: ADAUTHDC3;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988715] SERVICE ALERT: ADAUTHDC1;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988715] SERVICE ALERT: ADAUTHDC1;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988716] SERVICE ALERT: ADAUTHDC2;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988718] SERVICE ALERT: ADAUTHDC3;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988751] SERVICE ALERT: ADAUTHDC1;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988764] SERVICE ALERT: ADAUTHDC3;Service: Remote Management;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988764] SERVICE ALERT: ADAUTHDC2;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988764] SERVICE ALERT: ADAUTHDC2;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988776] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 364;OK;HARD;3;Eventlog check ok
[1460988783] SERVICE ALERT: ADAUTHDC2;Service: RPC;CRITICAL;HARD;3;connect to address 164.76.102.172 and port 12489: Connection refused
[1460988783] SERVICE ALERT: ADAUTHDC2;NSClient++ Version;CRITICAL;HARD;3;connect to address 164.76.102.172 and port 12489: Connection refused
[1460988796] SERVICE ALERT: ADAUTHDC3;Service: AD Web Service;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988803] SERVICE ALERT: ADAUTHDC3;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988895] SERVICE ALERT: ADAUTHDC1;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460988895] SERVICE ALERT: ADAUTHDC2;Service: AD Web Service;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988904] SERVICE ALERT: ADAUTHDC3;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988904] SERVICE ALERT: ADAUTHDC2;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460988904] SERVICE ALERT: ADAUTHDC3;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460989015] SERVICE ALERT: ADAUTHDC1;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460989044] SERVICE ALERT: idmrbs.emich.edu;LDAP Stats;OK;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1460989044] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;CRITICAL;SOFT;1;Black Cartridge HP CE260X is at CRITICAL level!
[1460989105] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;CRITICAL;SOFT;2;Black Cartridge HP CE260X is at CRITICAL level!
[1460989122] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1460989122] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (24.5% change >= 20.0% threshold)
[1460989135] SERVICE ALERT: ADAUTHDC1;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460989142] SERVICE ALERT: ADAUTHDC1;Service: WMI;OK;HARD;3;Winmgmt: Started
[1460989142] SERVICE ALERT: ADAUTHDC1;VMWare Tools;OK;HARD;3;VMTools: Started
[1460989142] SERVICE ALERT: ADAUTHDC1;Service: Dell Backup Agent;OK;HARD;3;ErdAgent: Started
[1460989143] SERVICE ALERT: ADAUTHDC2;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1460989144] SERVICE ALERT: ADAUTHDC2;Service: Distributed File System Replication;OK;HARD;3;DFSR: Started
[1460989144] SERVICE ALERT: ADAUTHDC3;Service: Intersite Messaging;OK;HARD;3;IsmServ: Started
[1460989148] SERVICE ALERT: ADAUTHDC2;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 6 minute(s)
[1460989148] SERVICE ALERT: ADAUTHDC3;C:\ Drive Space;OK;HARD;3;c:\ - total: 79.66 Gb - used: 16.55 Gb (21%) - free 63.10 Gb (79%)
[1460989164] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;CRITICAL;HARD;3;Black Cartridge HP CE260X is at CRITICAL level!
[1460989174] SERVICE ALERT: ADAUTHDC2;Service: Dell Backup Agent;OK;HARD;3;ErdAgent: Started
[1460989192] SERVICE ALERT: ADAUTHDC1;Service: Distributed File System Replication;OK;HARD;3;DFSR: Started
[1460989192] SERVICE ALERT: ADAUTHDC2;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1460989194] SERVICE ALERT: ADAUTHDC3;Service: Key Distribution Center;OK;HARD;3;KDC: Started
[1460989194] SERVICE ALERT: ADAUTHDC3;CPU Load;OK;HARD;3;CPU Load 2% (15 min average)
[1460989195] SERVICE ALERT: ADAUTHDC2;AD Replication;WARNING;HARD;3;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1460989244] SERVICE ALERT: ADAUTHDC2;Service: Intersite Messaging;OK;HARD;3;IsmServ: Started
[1460989244] SERVICE ALERT: ADAUTHDC3;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1460989244] SERVICE ALERT: ADAUTHDC2;C:\ Drive Space;OK;HARD;3;c:\ - total: 79.66 Gb - used: 15.07 Gb (19%) - free 64.59 Gb (81%)
[1460989245] SERVICE ALERT: ADAUTHDC1;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1460989246] SERVICE ALERT: ADAUTHDC3;Memory Usage;OK;HARD;3;Memory usage: total:10943.64 MB - used: 777.30 MB (7%) - free: 10166.35 MB (93%)
[1460989246] SERVICE ALERT: ADAUTHDC3;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1460989293] SERVICE ALERT: ADAUTHDC3;Service: RPC;OK;HARD;3;RpcSs: Started
[1460989295] SERVICE ALERT: ADAUTHDC1;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1460989295] SERVICE ALERT: ADAUTHDC1;C:\ Drive Space;OK;HARD;3;c:\ - total: 79.66 Gb - used: 15.93 Gb (20%) - free 63.73 Gb (80%)
[1460989296] SERVICE ALERT: ADAUTHDC2;CPU Load;OK;HARD;3;CPU Load 0% (15 min average)
[1460989296] SERVICE ALERT: ADAUTHDC2;Service: Key Distribution Center;OK;HARD;3;KDC: Started
[1460989298] SERVICE ALERT: ADAUTHDC3;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1460989325] SERVICE ALERT: ADAUTHDC3;Service: Dell Backup Agent;OK;HARD;3;ErdAgent: Started
[1460989330] SERVICE ALERT: ADAUTHDC1;Service: Intersite Messaging;OK;HARD;3;IsmServ: Started
[1460989331] SERVICE ALERT: ADAUTHDC1;CPU Load;OK;HARD;3;CPU Load 0% (15 min average)
[1460989344] SERVICE ALERT: ADAUTHDC3;Service: Remote Management;OK;HARD;3;WinRM: Started
[1460989344] SERVICE ALERT: ADAUTHDC2;Memory Usage;OK;HARD;3;Memory usage: total:10239.64 MB - used: 1038.24 MB (10%) - free: 9201.40 MB (90%)
[1460989344] SERVICE ALERT: ADAUTHDC3;Service: AD Domain Service;OK;HARD;3;NTDS: Started
[1460989344] SERVICE ALERT: ADAUTHDC2;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1460989376] SERVICE ALERT: ADAUTHDC3;Service: AD Web Service;OK;HARD;3;ADWS: Started
[1460989376] SERVICE ALERT: ADAUTHDC1;Service: Key Distribution Center;OK;HARD;3;KDC: Started
[1460989380] SERVICE ALERT: ADAUTHDC2;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1460989383] SERVICE ALERT: ADAUTHDC2;Service: RPC;OK;HARD;3;RpcSs: Started
[1460989383] SERVICE ALERT: ADAUTHDC3;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1460989427] SERVICE ALERT: ADAUTHDC3;Service: DFS Namespace;OK;HARD;3;Dfs: Started
[1460989427] SERVICE ALERT: ADAUTHDC3;Service: WMI;OK;HARD;3;Winmgmt: Started
[1460989447] SERVICE ALERT: ADAUTHDC3;AD Replication;WARNING;HARD;3;Send-MailMessage : The remote name could not be resolved: 'meritsmtp.emich.edu'
[1460989475] SERVICE ALERT: ADAUTHDC2;Service: AD Web Service;OK;HARD;3;ADWS: Started
[1460989484] SERVICE ALERT: ADAUTHDC3;Service: Windows Time;OK;HARD;3;W32Time: Started
[1460989484] SERVICE ALERT: ADAUTHDC3;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1460989484] SERVICE ALERT: ADAUTHDC2;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1460989505] SERVICE ALERT: ADAUTHDC1;AD Replication;WARNING;HARD;3;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1460989524] SERVICE ALERT: ADAUTHDC1;Service: Remote Management;OK;HARD;3;WinRM: Started
[1460989524] SERVICE ALERT: ADAUTHDC1;Service: AD Domain Service;OK;HARD;3;NTDS: Started
[1460989544] SERVICE ALERT: ADAUTHDC2;Service: Remote Management;OK;HARD;3;WinRM: Started
[1460989548] SERVICE ALERT: ADAUTHDC1;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1460989548] SERVICE ALERT: ADAUTHDC1;Service: DFS Namespace;OK;HARD;3;Dfs: Started
[1460989554] SERVICE ALERT: ADAUTHDC2;Service: AD Domain Service;OK;HARD;3;NTDS: Started
[1460989554] SERVICE ALERT: ADAUTHDC1;Memory Usage;OK;HARD;3;Memory usage: total:10943.64 MB - used: 1120.04 MB (10%) - free: 9823.60 MB (90%)
[1460989564] SERVICE ALERT: ADAUTHDC1;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1460989565] SERVICE ALERT: ADAUTHDC2;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1460989567] SERVICE ALERT: ADAUTHDC3;VMWare Tools;OK;HARD;3;VMTools: Started
[1460989569] SERVICE ALERT: ADAUTHDC1;Service: AD Web Service;OK;HARD;3;ADWS: Started
[1460989593] SERVICE ALERT: ADAUTHDC1;Service: RPC;OK;HARD;3;RpcSs: Started
[1460989593] SERVICE ALERT: ADAUTHDC3;Service: Distributed File System Replication;OK;HARD;3;DFSR: Started
[1460989644] SERVICE ALERT: ADAUTHDC2;Service: DFS Namespace;OK;HARD;3;Dfs: Started
[1460989644] SERVICE ALERT: ADAUTHDC3;Service: DNS Service;OK;HARD;3;DNS: Started
[1460989644] SERVICE ALERT: ADAUTHDC2;Service: WMI;OK;HARD;3;Winmgmt: Started
[1460989654] SERVICE ALERT: ADAUTHDC1;Service: Windows Time;OK;HARD;3;W32Time: Started
[1460989655] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460989656] SERVICE ALERT: ADAUTHDC1;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1460989656] SERVICE ALERT: ADAUTHDC2;VMWare Tools;OK;HARD;3;VMTools: Started
[1460989658] SERVICE ALERT: ADAUTHDC3;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 13 minute(s)
[1460989659] SERVICE ALERT: ADAUTHDC3;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1460989661] SERVICE ALERT: ADAUTHDC2;Service: DNS Service;OK;HARD;3;DNS: Started
[1460989685] SERVICE ALERT: ADAUTHDC2;Service: Windows Time;OK;HARD;3;W32Time: Started
[1460989701] SERVICE ALERT: ADAUTHDC1;Service: DNS Service;OK;HARD;3;DNS: Started
[1460989701] SERVICE ALERT: ADAUTHDC1;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 14 minute(s)
[1460989714] SERVICE ALERT: hallefirst1;Device Status;OK;SOFT;2;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1460989715] SERVICE ALERT: ADAUTHDC1;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1460989722] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1460989818] HOST ALERT: hallefirst2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460989842] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1460989882] HOST ALERT: hallefirst2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.58 ms
[1460989962] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1460990173] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (26.5% change >= 20.0% threshold)
[1460990173] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STARTED; Service appears to have started flapping (34.9% change >= 20.0% threshold)
[1460990173] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1460990173] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1460990173] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1460990173] HOST DOWNTIME ALERT: LANSWEEPER;STARTED; Host has entered a period of scheduled downtime
[1460990173] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1460990173] HOST DOWNTIME ALERT: IMAGES;STARTED; Host has entered a period of scheduled downtime
[1460990183] HOST ALERT: hallethirdsouth;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.36 ms
[1460990394] SERVICE ALERT: hallethirdsouth;Device Status;OK;HARD;3;Ready
[1460990434] SERVICE ALERT: hallethirdsouth;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1460990855] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1460990914] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460990914] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1460990974] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460990975] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1460991033] SERVICE ALERT: eateriespk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460991034] SERVICE ALERT: hallefirst1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460991093] SERVICE ALERT: eateriespk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460991097] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1460991097] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STARTED; Service appears to have started flapping (33.1% change >= 20.0% threshold)
[1460991097] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1460991097] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1460991097] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1460991097] HOST DOWNTIME ALERT: LANSWEEPER;STARTED; Host has entered a period of scheduled downtime
[1460991097] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1460991097] HOST DOWNTIME ALERT: IMAGES;STARTED; Host has entered a period of scheduled downtime
[1460991153] SERVICE ALERT: eateriespk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460991598] SERVICE ALERT: ADAUTHDC2;AD Replication;OK;HARD;3;OK: replication is up and running
[1460991634] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;Ready
[1460991850] SERVICE ALERT: ADAUTHDC3;AD Replication;OK;HARD;3;OK: replication is up and running
[1460991909] SERVICE ALERT: ADAUTHDC1;AD Replication;OK;HARD;3;OK: replication is up and running
[1460991966] HOST DOWNTIME ALERT: LANSWEEPER;STOPPED; Host has exited from a period of scheduled downtime
[1460992174] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1460992234] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460992290] HOST DOWNTIME ALERT: IMAGES;STOPPED; Host has exited from a period of scheduled downtime
[1460992294] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460992354] SERVICE ALERT: hallefirst1;Device Status;OK;SOFT;3;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1460992360] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1460992622] SERVICE ALERT: owen101lab3;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460992682] SERVICE ALERT: owen101lab3;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460992728] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (26.9% change >= 20.0% threshold)
[1460992728] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STARTED; Service appears to have started flapping (27.6% change >= 20.0% threshold)
[1460992728] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1460992728] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1460992728] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1460992728] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1460992742] SERVICE ALERT: owen101lab3;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460992960] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1460993560] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 90% (15 min average)
[1460993853] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;CRITICAL;SOFT;1;Cyan Cartridge HP CE261A is at CRITICAL level!
[1460993914] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;CRITICAL;SOFT;2;Cyan Cartridge HP CE261A is at CRITICAL level!
[1460993973] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;CRITICAL;HARD;3;Cyan Cartridge HP CE261A is at CRITICAL level!
[1460994384] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1460994444] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1460994505] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;SOFT;3;Black Print Cartridge HP C8543X is at 55% - OK!
[1460994760] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1460994880] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1460995000] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1460995165] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;OK;HARD;3;
[1460995174] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;OK;HARD;3;
[1460995353] SERVICE ALERT: eateriespk;Device Status;OK;HARD;3;Checking printer
[1460995600] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1460995742] SERVICE ALERT: owen101lab3;Device Status;OK;HARD;3;Clearing paper path
[1460996184] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1460996200] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1460996244] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 32% - OK!
[1460996689] SERVICE ALERT: hallegroundcolor;Magenta Cartridge HP CE263A;OK;HARD;3;
[1460996948] SERVICE ALERT: owen101lab3;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1460996978] HOST ALERT: owen101lab3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460997008] SERVICE ALERT: owen101lab3;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1460997038] HOST ALERT: owen101lab3;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460997068] SERVICE ALERT: owen101lab3;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1460997098] HOST ALERT: owen101lab3;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460997188] HOST ALERT: owen101lab3;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460997278] HOST ALERT: owen101lab3;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460997299] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1460997329] HOST ALERT: owen101lab3;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1460997359] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1460997389] HOST ALERT: owen101lab3;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1460997419] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1460997449] HOST ALERT: owen101lab3;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1460997450] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1460997480] HOST ALERT: owen101lab3;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1460997484] SERVICE ALERT: hallegroundcolor;Black Cartridge HP CE260X;OK;HARD;3;Black Cartridge HP CE260X is at 100% - OK!
[1460997510] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1460997540] HOST ALERT: owen101lab3;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1460997570] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1460998664] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 1 Trays Empty
[1460998725] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 1 Trays Empty
[1460998784] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 1 Trays Empty
[1460999085] SERVICE ALERT: owen101lab1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460999145] SERVICE ALERT: owen101lab1;Device Status;OK;SOFT;2;Clearing paper path
[1460999199] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1460999799] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1460999963] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1461000399] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 81% (15 min average)
[1461000999] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461001119] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1461001184] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1461001239] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461001355] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461001414] SERVICE ALERT: hallefirst1;Device Status;OK;SOFT;2;CLOSE RIGHT DOOR  46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1461001763] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461001823] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461001883] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461002694] SERVICE ALERT: marshallpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461002754] SERVICE ALERT: marshallpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461002814] SERVICE ALERT: marshallpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461003312] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;SOFT;1;WARNING: Replication with ADDC5 failed!
[1461003414] SERVICE ALERT: marshallpk;Device Status;OK;HARD;3;TRAY 3 EMPTY PLAIN LETTER   Ready
[1461003431] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;SOFT;2;WARNING: Replication with ADDC5 failed!
[1461003552] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;HARD;3;WARNING: Replication with ADDC5 failed!
[1461004203] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (47.8% change >= 20.0% threshold)
[1461004203] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1461004203] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1461004203] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1461004203] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1461004239] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461004454] SERVICE ALERT: prayharroldpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1461004515] SERVICE ALERT: prayharroldpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1461004567] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;WARNING;SOFT;1;0
[1461004575] SERVICE ALERT: prayharroldpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1461004687] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;WARNING;SOFT;2;0
[1461004807] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;WARNING;HARD;3;0
[1461004839] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461004858] SERVICE ALERT: idm1.emich.edu;Loopback Driver;WARNING;SOFT;1;0
[1461004977] SERVICE ALERT: idm1.emich.edu;Loopback Driver;WARNING;SOFT;2;0
[1461005096] SERVICE ALERT: idm1.emich.edu;Loopback Driver;WARNING;HARD;3;0
[1461005175] SERVICE ALERT: prayharroldpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1461006005] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;OK;HARD;3;0
[1461006295] SERVICE ALERT: idm1.emich.edu;Loopback Driver;OK;HARD;3;0
[1461006298] SERVICE ALERT: cava2;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.7.192 and port 12489: Connection refused
[1461006328] HOST ALERT: cava2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1461006362] SERVICE ALERT: cava2;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461006364] SERVICE ALERT: cava2;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461006389] HOST ALERT: cava2;UP;SOFT;2;PING WARNING - Packet loss = 81%, RTA = 0.40 ms
[1461006418] SERVICE ALERT: cava2;Service: RPC;OK;SOFT;2;RpcSs: Started
[1461006463] SERVICE ALERT: cava2;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1461006465] SERVICE ALERT: cava2;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1461006584] SERVICE ALERT: cava2;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1461007169] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1461007184] SERVICE ALERT: cava2;CPU Load;OK;HARD;3;CPU Load 2% (15 min average)
[1461007230] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1461007289] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1461007883] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Processing job from tray 4
[1461007895] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1461007920] SERVICE ALERT: markjeffersonpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1461007974] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461008034] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461008436] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461008489] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 42% - OK!
[1461008489] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1461008519] HOST ALERT: hallefirst2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1461008549] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1461008575] HOST ALERT: hallefirst2;UP;SOFT;2;PING WARNING - Packet loss = 81%, RTA = 0.55 ms
[1461008609] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1461008620] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1461008629] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1461008635] SERVICE ALERT: markjeffersonpk;Device Status;OK;HARD;3;Checking printer
[1461008649] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1461008680] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1461008689] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1461008709] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1461008710] HOST ALERT: hallefirst1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1461008734] SERVICE ALERT: hallefirst1;Device Status;OK;SOFT;3;Warming up
[1461008738] HOST ALERT: hallefirst1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1461008749] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1461008769] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1461008818] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1461008878] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1461008907] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1461008938] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1461008967] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1461009027] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1461009069] SERVICE ALERT: hallefirst3;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1461009089] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1461009129] SERVICE ALERT: hallefirst3;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1461009149] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1461009159] HOST ALERT: hallefirst3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1461009189] SERVICE ALERT: hallefirst3;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1461009209] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1461009209] SERVICE FLAPPING ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1461009219] HOST ALERT: hallefirst3;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1461009248] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1461009274] SERVICE ALERT: hallefirst3;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1461009278] HOST ALERT: hallefirst3;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1461009308] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1461009334] SERVICE ALERT: hallefirst3;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1461009338] HOST ALERT: hallefirst3;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1461009340] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1461009368] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1461009394] SERVICE ALERT: hallefirst3;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1461009398] HOST ALERT: hallefirst3;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1461009400] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1461009460] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1461009488] HOST ALERT: hallefirst3;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1461009578] HOST ALERT: hallefirst3;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1461009602] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (34.0% change >= 20.0% threshold)
[1461009602] SERVICE FLAPPING ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1461009603] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1461009603] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1461009603] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1461009603] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1461009668] HOST ALERT: hallefirst3;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1461009758] HOST ALERT: hallefirst3;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1461009804] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Warming up
[1461009804] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1461009820] HOST ALERT: hallefirst3;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1461009909] HOST ALERT: hallefirst3;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1461009944] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 35% - OK!
[1461009964] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1461009982] HOST ALERT: eateriespk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1461009999] HOST ALERT: hallefirst3;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1461010046] HOST ALERT: eateriespk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.62 ms
[1461010055] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;Ready
[1461010062] HOST ALERT: hallefirst3;UP;SOFT;7;PING OK - Packet loss = 0%, RTA = 0.61 ms
[1461010133] SERVICE ALERT: hallefirst1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1461010221] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 49% - OK!
[1461010234] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461010383] SERVICE ALERT: hallefirst3;Device Status;OK;HARD;3;PERFORM PRINTER MAINTENANCE 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Processing job from tray 4
[1461010562] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 9% - OK!
[1461010588] SERVICE ALERT: hallefirst3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1461010651] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1461010711] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1461010771] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1461011003] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461011063] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461011123] SERVICE ALERT: hallefirst2;Device Status;OK;SOFT;3;Processing job from tray 4
[1461012323] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461012383] SERVICE ALERT: hallefirst2;Device Status;OK;SOFT;2;There is no help available for this item. 54 6F 20 65 78 69 74 20 70 72 65 73 73 20 1F 2E
[1461013234] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 91% (15 min average)
[1461013834] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1461013954] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1461014074] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461015209] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;CRITICAL;HARD;3;Black Print Cartridge HP C8543X is at 1% - CRITICAL!
[1461015458] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1461015518] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 27% - OK!
[1461015874] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 81% (15 min average)
[1461015983] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461016043] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461016103] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461016118] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1461016178] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 27% - OK!
[1461016474] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461016594] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461016703] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Ready
[1461016714] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461016951] SERVICE ALERT: eateriespk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461017011] SERVICE ALERT: eateriespk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461017071] SERVICE ALERT: eateriespk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461017383] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 1 Trays Empty
[1461017442] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! All Trays Still Full
[1461017671] SERVICE ALERT: eateriespk;Device Status;OK;HARD;3;TRAY 4 EMPTY PLAIN LETTER   Processing job
[1461017969] SERVICE ALERT: eateriespk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1461019703] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461019763] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461019823] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461021024] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Ready
[1461021209] SERVICE FLAPPING ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1461021514] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 72% (15 min average)
[1461022101] SERVICE ALERT: owen101lab1;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1461022114] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461022161] SERVICE ALERT: owen101lab1;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1461022221] SERVICE ALERT: owen101lab1;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1461022234] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1461022354] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1461024755] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 90% (15 min average)
[1461024928] SERVICE ALERT: ADRODC-LIV;AD Replication;OK;HARD;3;OK: replication is up and running
[1461025952] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461026072] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461026192] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461027392] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1461027661] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;SOFT;1;e:\ - total: 499.87 Gb - used: 404.66 Gb (81%) - free 95.21 Gb (19%)
[1461027781] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;SOFT;2;e:\ - total: 499.87 Gb - used: 404.66 Gb (81%) - free 95.21 Gb (19%)
[1461027901] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;HARD;3;e:\ - total: 499.87 Gb - used: 404.66 Gb (81%) - free 95.21 Gb (19%)
[1461027992] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461028112] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461028232] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461028823] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461028883] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461028943] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461029253] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461029313] SERVICE ALERT: hallefirst1;Device Status;OK;SOFT;2;Open front door.
[1461029432] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461029543] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1461030032] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461030301] SERVICE ALERT: AD-RMAD;E:\ Drive Space;OK;HARD;3;e:\ - total: 499.87 Gb - used: 398.80 Gb (80%) - free 101.07 Gb (20%)
[1461030632] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461031232] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1461032432] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461032552] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461032635] SERVICE ALERT: CBORD;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461032672] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1461032734] SERVICE ALERT: CBORD;CPU Load;OK;SOFT;2;CPU Load 2% (15 min average)
[1461032867] SERVICE ALERT: DB01-2012;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461032987] SERVICE ALERT: DB01-2012;CPU Load;OK;SOFT;2;CPU Load 9% (15 min average)
[1461033872] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461034472] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461036568] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1461036628] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1461036688] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty