[1454318913] SERVICE ALERT: hallefirst1;Paper Tray Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.104.232! Make sure host is up and SNMP is configured properly.
[1454318919] HOST ALERT: hallefirst1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1454318973] SERVICE ALERT: hallefirst1;Paper Tray Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.104.232! Make sure host is up and SNMP is configured properly.
[1454319003] HOST ALERT: hallefirst1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1454319033] SERVICE ALERT: hallefirst1;Paper Tray Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.104.232! Make sure host is up and SNMP is configured properly.
[1454319063] HOST ALERT: hallefirst1;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1454319153] HOST ALERT: hallefirst1;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1454319243] HOST ALERT: hallefirst1;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1454319333] HOST ALERT: hallefirst1;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1454319423] HOST ALERT: hallefirst1;DOWN;SOFT;7;(Host check timed out after 30.00 seconds)
[1454319513] HOST ALERT: hallefirst1;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1454319603] HOST ALERT: hallefirst1;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1454319663] HOST ALERT: hallefirst1;DOWN;HARD;10;(Host check timed out after 30.00 seconds)
[1454331317] SERVICE ALERT: LANSWEEPER;CPU Load;WARNING;SOFT;1;CPU Load 88% (5 min average)
[1454331437] SERVICE ALERT: LANSWEEPER;CPU Load;OK;SOFT;2;CPU Load 72% (5 min average)
[1454332440] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1454332580] SERVICE ALERT: owen101lab1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1454333040] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at 5% - WARNING!
[1454333101] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at 5% - WARNING!
[1454333160] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at 5% - WARNING!
[1454333579] HOST ALERT: sillpk;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1454333643] HOST ALERT: sillpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.45 ms
[1454334143] SERVICE ALERT: sillpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1454334250] SERVICE ALERT: sillpk;Black Cartridge HP CC364X;OK;HARD;3;Black Cartridge HP CC364X is at 100% - OK!
[1454334960] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1454335561] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1454335561] SERVICE FLAPPING ALERT: studentcenterpk;Black Print Cartridge HP C8543X;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1454343384] SERVICE ALERT: owen101lab1;Paper Tray Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.203.101! Make sure host is up and SNMP is configured properly.
[1454343414] HOST ALERT: owen101lab1;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1454343444] SERVICE ALERT: owen101lab1;Paper Tray Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.203.101! Make sure host is up and SNMP is configured properly.
[1454343474] HOST ALERT: owen101lab1;DOWN;SOFT;2;(Host check timed out after 30.00 seconds)
[1454343504] SERVICE ALERT: owen101lab1;Paper Tray Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.203.101! Make sure host is up and SNMP is configured properly.
[1454343534] HOST ALERT: owen101lab1;DOWN;SOFT;3;(Host check timed out after 30.00 seconds)
[1454343624] HOST ALERT: owen101lab1;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1454343714] HOST ALERT: owen101lab1;DOWN;SOFT;5;(Host check timed out after 30.00 seconds)
[1454343750] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.203.101! Make sure host is up and SNMP is configured properly.
[1454343780] HOST ALERT: owen101lab1;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1454343810] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.203.101! Make sure host is up and SNMP is configured properly.
[1454343840] HOST ALERT: owen101lab1;DOWN;SOFT;7;(Host check timed out after 30.00 seconds)
[1454343870] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;WARNING;HARD;3;WARNING: No SNMP response from 164.76.203.101! Make sure host is up and SNMP is configured properly.
[1454343900] HOST ALERT: owen101lab1;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1454343984] SERVICE ALERT: owen101lab3;Paper Tray Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1454343990] HOST ALERT: owen101lab1;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1454343991] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1454344014] HOST ALERT: owen101lab3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1454344044] SERVICE ALERT: owen101lab3;Paper Tray Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1454344051] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1454344054] HOST ALERT: owen101lab1;UP;SOFT;10;PING OK - Packet loss = 0%, RTA = 0.73 ms
[1454344074] HOST ALERT: owen101lab3;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1454344098] SERVICE ALERT: owen101lab1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1454344099] SERVICE ALERT: owen101lab3;Paper Tray Status;OK;SOFT;3;Printer Paper Is OK! All Trays Still Full
[1454344103] HOST ALERT: owen101lab3;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.54 ms
[1454344105] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;OK;SOFT;3;Black Print Cartridge HP C8543X is at 36% - OK!
[1454344465] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 37% - OK!
[1454346437] SERVICE ALERT: LANSWEEPER;CPU Load;WARNING;SOFT;1;CPU Load 82% (5 min average)
[1454346557] SERVICE ALERT: LANSWEEPER;CPU Load;OK;SOFT;2;CPU Load 15% (5 min average)
[1454346960] SERVICE FLAPPING ALERT: studentcenterpk;Black Print Cartridge HP C8543X;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1454355555] SERVICE ALERT: LANSWEEPER;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1454355675] SERVICE ALERT: LANSWEEPER;CPU Load;CRITICAL;SOFT;2;CPU Load 97% (5 min average)
[1454355795] SERVICE ALERT: LANSWEEPER;CPU Load;OK;SOFT;3;CPU Load 46% (5 min average)
[1454358327] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1454358387] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1454358447] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1454358790] SERVICE ALERT: ADWSUSSERVERS;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.102.133 and port 12489: No route to host
[1454358793] HOST ALERT: ADWSUSSERVERS;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.133)
[1454358843] HOST ALERT: ADWSUSSERVERS;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1454358898] SERVICE ALERT: ADWSUSSERVERS;NSClient++ Version;OK;SOFT;2;NSClient++ 0,4,1,73 2012-12-17
[1454371046] SERVICE ALERT: studentcenterpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full