[1478495183] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1478495783] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478495903] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1478496023] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1478497223] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1478497823] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1478498423] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1478498709] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 102 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478498828] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478498946] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1478499023] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1478499066] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;2;Warning:  idm1  are more than 120 seconds behind!
[1478499186] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1478499787] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1478500200] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1478500388] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1478501423] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 90% (15 min average)
[1478501586] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1478501586] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1478501675] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 364;CRITICAL;SOFT;1;Windows Server Update Services(error, 364, success)[Content file download failed.
[1478501794] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 364;CRITICAL;SOFT;2;Windows Server Update Services(error, 364, success)[Content file download failed.
[1478501914] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 364;CRITICAL;HARD;3;Windows Server Update Services(error, 364, success)[Content file download failed.
[1478502023] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1478502143] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 73% (15 min average)
[1478502743] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1478502786] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1478502863] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1478502983] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1478504586] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1478504783] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 65% (15 min average)
[1478505383] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478505503] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1478505624] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1478505786] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1478505888] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 107 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478506008] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478506116] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10032;CRITICAL;SOFT;1;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1478506235] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10032;CRITICAL;SOFT;2;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1478506355] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10032;CRITICAL;HARD;3;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1478506386] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1478506506] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;2;Warning:  idm1  are more than 120 seconds behind!
[1478506626] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1478507226] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1478508347] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1478508426] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1478508467] SERVICE ALERT: epprogsql2;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1478508515] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 364;OK;HARD;3;Eventlog check ok
[1478508587] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1478508623] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 71% (15 min average)
[1478509026] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1478509223] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478509343] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1478509463] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 54% (15 min average)
[1478510385] SERVICE ALERT: epprogsql2;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1478510808] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 105 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478510928] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478510985] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1478511003] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (32.8% change >= 20.0% threshold)
[1478511003] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (53.0% change >= 20.0% threshold)
[1478511003] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1478511003] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1478511003] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1478511003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1478511003] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1478511003] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1478511003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1478511003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1478511003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1478511003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1478511003] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1478511003] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1478511003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1478511003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1478511003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1478511003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1478511003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1478512185] SERVICE ALERT: epprogsql2;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1478512185] SERVICE FLAPPING ALERT: epprogsql2;CPU Load;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1478512623] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 121 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478512743] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478512785] SERVICE ALERT: epprogsql2;CPU Load;OK;HARD;3;CPU Load 22% (15 min average)
[1478512802] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1478512802] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (38.2% change >= 20.0% threshold)
[1478512802] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1478512802] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1478512802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1478512802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1478512802] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1478512802] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1478512802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1478512802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1478512802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1478512802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1478512802] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1478512802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1478512802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1478512802] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1478512802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1478512802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1478512802] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1478512956] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10032;OK;HARD;3;Eventlog check ok
[1478514863] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1478514983] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 92% (15 min average)
[1478515583] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1478515703] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1478515743] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 104 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478515863] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478516328] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 102 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478516448] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478516903] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1478517023] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1478517048] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 138 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478517143] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1478517169] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478517663] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 118 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478517768] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 109 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478517783] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478517887] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478518200] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service has exited from a period of scheduled downtime
[1478518343] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 83% (15 min average)
[1478518943] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478519063] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1478519183] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 92% (15 min average)
[1478519583] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 111 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478519703] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478519783] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478519903] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1478520023] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1478520287] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1048 found since last run.
[1478520303] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1004 found since last run.
[1478520407] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478520423] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478520426] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1478521007] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1563 found since last run.
[1478521023] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1579 found since last run.
[1478521127] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478521143] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478521727] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1596 found since last run.
[1478521743] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1588 found since last run.
[1478521847] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478521863] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478522447] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1671 found since last run.
[1478522463] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1677 found since last run.
[1478522567] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478522583] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478523167] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1688 found since last run.
[1478523183] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1698 found since last run.
[1478523287] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478523303] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478523887] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1196 found since last run.
[1478523903] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1197 found since last run.
[1478523961] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1478523995] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1478524007] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478524021] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1478524023] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478524032] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1478524055] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1478524079] SERVICE ALERT: porterpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1478524081] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1478524092] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1478524099] SERVICE ALERT: porterpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1478524109] HOST ALERT: porterpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1478524115] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1478524133] SERVICE ALERT: porterpk;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 9% - OK!
[1478524137] HOST ALERT: porterpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.53 ms
[1478524145] HOST ALERT: prayharroldpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1478524152] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1478524153] SERVICE ALERT: porterpk;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! All Trays Still Full
[1478524182] HOST ALERT: prayharroldpk;UP;SOFT;2;PING WARNING - Packet loss = 96%, RTA = 1.06 ms
[1478524223] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1478524675] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 29% - OK!
[1478524709] SERVICE ALERT: prayharroldpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1478524746] SERVICE ALERT: prayharroldpk;Device Status;OK;HARD;3;Ready
[1478524823] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1478525207] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6809 found since last run. oneLevelSearchOps OK. searchOps is too high! 7020 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478525327] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478525425] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1478525485] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1478525545] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1478525915] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 364;CRITICAL;SOFT;1;Windows Server Update Services(error, 364, success)[Content file download failed.
[1478525927] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7849 found since last run. oneLevelSearchOps OK. searchOps is too high! 8092 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478526023] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478526034] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 364;CRITICAL;SOFT;2;Windows Server Update Services(error, 364, success)[Content file download failed.
[1478526047] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478526143] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 92% (15 min average)
[1478526153] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 364;CRITICAL;HARD;3;Windows Server Update Services(error, 364, success)[Content file download failed.
[1478526647] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5494 found since last run. oneLevelSearchOps OK. searchOps is too high! 5660 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478526745] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;Clear jam then press OK.
[1478526767] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478527023] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7884 found since last run. oneLevelSearchOps OK. searchOps is too high! 8436 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478527143] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478527343] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1478527345] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1478527367] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6046 found since last run. oneLevelSearchOps OK. searchOps is too high! 6248 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478527406] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1478527463] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 55% (15 min average)
[1478527465] SERVICE ALERT: hallegroundcolor;Device Status;OK;SOFT;3;Clear jam then press OK.
[1478527487] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478527743] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5139 found since last run. oneLevelSearchOps OK. searchOps is too high! 5631 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478527863] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478528063] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1478528065] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1478528087] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8737 found since last run. oneLevelSearchOps OK. searchOps is too high! 8948 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478528125] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1478528183] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1478528185] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1478528207] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478528303] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1478528346] SERVICE ALERT: prayharroldpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1478528406] SERVICE ALERT: prayharroldpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1478528463] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6553 found since last run. oneLevelSearchOps OK. searchOps is too high! 7134 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478528466] SERVICE ALERT: prayharroldpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1478528584] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478528785] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;Clear jam then press OK.
[1478528785] SERVICE FLAPPING ALERT: hallegroundcolor;Device Status;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1478528807] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10976 found since last run. oneLevelSearchOps OK. searchOps is too high! 11176 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478528927] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478529002] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (29.3% change >= 20.0% threshold)
[1478529002] SERVICE FLAPPING ALERT: hallegroundcolor;Device Status;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1478529002] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1478529002] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1478529002] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1478529002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1478529002] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1478529002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1478529002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1478529002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1478529002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1478529002] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1478529002] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1478529002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1478529002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1478529002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1478529002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1478529002] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1478529026] SERVICE ALERT: ADFS;Uptime;OK;HARD;3;System Uptime - 2 day(s) 23 hour(s) 39 minute(s)
[1478529026] SERVICE ALERT: ADFS;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1478529030] HOST ALERT: ADFS;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1478529066] SERVICE ALERT: prayharroldpk;Device Status;OK;HARD;3;TRAY 2 OPEN  46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Processing job
[1478529066] SERVICE FLAPPING ALERT: prayharroldpk;Device Status;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1478529183] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5734 found since last run. oneLevelSearchOps OK. searchOps is too high! 6101 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478529193] SERVICE ALERT: ADFS;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1478529195] SERVICE ALERT: ADFS;C:\ Drive Space;OK;HARD;3;c:\ - total: 79.51 Gb - used: 18.62 Gb (23%) - free 60.89 Gb (77%)
[1478529282] SERVICE ALERT: ADFS;Service: RPC;OK;HARD;3;RpcSs: Started
[1478529282] SERVICE ALERT: ADFS;CPU Load;OK;HARD;3;CPU Load 0% (15 min average)
[1478529303] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478529370] SERVICE ALERT: ADFS;Memory Usage;OK;HARD;3;Memory usage: total:10239.64 MB - used: 1675.32 MB (16%) - free: 8564.32 MB (84%)
[1478529370] SERVICE ALERT: ADFS;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1478529385] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1478529446] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1478529453] SERVICE ALERT: ADFS;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1478529453] SERVICE ALERT: ADFS;Service: WMI;OK;HARD;3;Winmgmt: Started
[1478529503] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 25% (15 min average)
[1478529505] SERVICE ALERT: hallegroundcolor;Device Status;OK;SOFT;3;Clear jam then press OK.
[1478529527] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7942 found since last run. oneLevelSearchOps OK. searchOps is too high! 8121 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478529528] SERVICE ALERT: ADFS;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1478529543] SERVICE ALERT: ADFS;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1478529543] SERVICE ALERT: ADFS;VMWare Tools;OK;HARD;3;VMTools: Started
[1478529543] SERVICE ALERT: ADFS;Service: Windows Time;OK;HARD;3;W32Time: Started
[1478529647] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478529903] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6136 found since last run. oneLevelSearchOps OK. searchOps is too high! 6577 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478530023] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478530237] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1478530247] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6361 found since last run. oneLevelSearchOps OK. searchOps is too high! 6537 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478530297] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1478530356] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1478530367] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478530623] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5215 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478530743] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478530967] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8938 found since last run. oneLevelSearchOps OK. searchOps is too high! 9143 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478531087] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478531301] SERVICE ALERT: owen101lab3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1478531303] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478531336] SERVICE ALERT: owen101lab3;Device Status;OK;HARD;3;Processing job from tray 4
[1478531343] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6646 found since last run. oneLevelSearchOps OK. searchOps is too high! 7163 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478531423] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1478531463] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478531543] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1478531584] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 36% - OK!
[1478531687] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9326 found since last run. oneLevelSearchOps OK. searchOps is too high! 9505 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478531807] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478532063] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6212 found since last run. oneLevelSearchOps OK. searchOps is too high! 6882 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478532183] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478532407] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9337 found since last run. oneLevelSearchOps OK. searchOps is too high! 9510 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478532527] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478532538] SERVICE ALERT: porterpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1478532559] SERVICE ALERT: porterpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1478532598] SERVICE ALERT: porterpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1478532619] SERVICE ALERT: porterpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1478532658] SERVICE ALERT: porterpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1478532668] SERVICE ALERT: porterpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1478532679] SERVICE ALERT: porterpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1478532728] SERVICE ALERT: porterpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1478532743] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1478532753] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 364;OK;HARD;3;Eventlog check ok
[1478532788] SERVICE ALERT: porterpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1478533127] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 13932 found since last run. oneLevelSearchOps OK. searchOps is too high! 14154 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 3035 found since last run.
[1478533247] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478533343] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1478533383] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 12924 found since last run. oneLevelSearchOps OK. searchOps is too high! 13549 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 3769 found since last run.
[1478533504] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478533674] HOST ALERT: ds1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1478533692] HOST ALERT: gmt;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1478533764] HOST ALERT: ds1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1478533782] HOST ALERT: gmt;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1478533836] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1478533848] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 14248 found since last run. oneLevelSearchOps OK. searchOps is too high! 14442 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478533853] SERVICE ALERT: porterpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 9% - OK!
[1478533854] HOST ALERT: ds1;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1478533872] HOST ALERT: gmt;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1478533874] SERVICE ALERT: porterpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1478533896] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1478533944] HOST ALERT: ds1;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1478533956] SERVICE ALERT: marshallpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1478533957] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1478533962] HOST ALERT: gmt;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1478533967] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478533982] SERVICE ALERT: porterpk;Device Status;OK;HARD;3;Ready
[1478534034] HOST ALERT: ds1;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1478534052] HOST ALERT: gmt;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1478534103] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 21109 found since last run. oneLevelSearchOps OK. searchOps is too high! 21528 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478534124] HOST ALERT: ds1;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1478534142] HOST ALERT: gmt;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1478534214] HOST ALERT: ds1;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1478534222] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478534232] HOST ALERT: gmt;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1478534304] HOST ALERT: ds1;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1478534322] HOST ALERT: gmt;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1478534394] HOST ALERT: ds1;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1478534412] HOST ALERT: gmt;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1478534484] HOST ALERT: ds1;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1478534502] HOST ALERT: gmt;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1478534543] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 79% (15 min average)
[1478534557] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1478534566] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 12133 found since last run. oneLevelSearchOps OK. searchOps is too high! 12349 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478534617] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1478534676] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1478534686] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478534822] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 12140 found since last run. oneLevelSearchOps OK. searchOps is too high! 12636 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478534942] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478535143] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1478535263] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1478535286] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10258 found since last run. oneLevelSearchOps OK. searchOps is too high! 10583 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478535383] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1478535406] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478535542] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7307 found since last run. oneLevelSearchOps OK. searchOps is too high! 7859 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478535662] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478535876] SERVICE ALERT: marshallpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1478535876] SERVICE FLAPPING ALERT: marshallpk;Paper Tray Status;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1478536006] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8542 found since last run. oneLevelSearchOps OK. searchOps is too high! 8858 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478536126] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478536262] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8314 found since last run. oneLevelSearchOps OK. searchOps is too high! 8625 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478536356] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;CRITICAL;HARD;3;Black Print Cartridge HP C8543X is at 1% - CRITICAL!
[1478536382] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478536726] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9657 found since last run. oneLevelSearchOps OK. searchOps is too high! 9833 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478536846] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478536982] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5448 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478537102] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478537446] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7181 found since last run. oneLevelSearchOps OK. searchOps is too high! 7268 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478537556] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1478537566] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478537677] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1478537702] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9383 found since last run. oneLevelSearchOps OK. searchOps is too high! 9862 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478537737] SERVICE ALERT: marshallpk;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! 1 of 2 Trays Still Full
[1478537822] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478538166] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10542 found since last run. oneLevelSearchOps OK. searchOps is too high! 10625 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478538286] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478538383] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1478538422] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5529 found since last run. oneLevelSearchOps OK. searchOps is too high! 5957 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478538542] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478538613] SERVICE ALERT: eateriespk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1478538886] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 11450 found since last run. oneLevelSearchOps OK. searchOps is too high! 11513 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478538983] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1478539006] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478539142] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 12001 found since last run. oneLevelSearchOps OK. searchOps is too high! 12595 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478539262] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478539583] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1478539606] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 13378 found since last run. oneLevelSearchOps OK. searchOps is too high! 13426 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478539726] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478539803] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (51.2% change >= 20.0% threshold)
[1478539803] SERVICE FLAPPING ALERT: marshallpk;Paper Tray Status;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1478539803] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1478539803] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1478539803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1478539803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1478539803] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1478539803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1478539803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1478539803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1478539803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1478539803] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1478539803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1478539803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1478539803] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1478539803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1478539803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1478539803] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1478539829] SERVICE ALERT: ADFS;Uptime;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478539829] SERVICE ALERT: ADFS;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478539832] HOST ALERT: ADFS;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.57)
[1478539862] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 12155 found since last run. oneLevelSearchOps OK. searchOps is too high! 13065 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478539895] HOST ALERT: ADFS;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.4.57)
[1478539949] SERVICE ALERT: ADFS;Uptime;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478539949] SERVICE ALERT: ADFS;Service: Firewall;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478539952] HOST ALERT: ADFS;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.4.57)
[1478539982] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478539996] SERVICE ALERT: ADFS;Service: NetLogin;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478539996] SERVICE ALERT: ADFS;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478539999] HOST ALERT: ADFS;DOWN;SOFT;4;CRITICAL - Host Unreachable (164.76.4.57)
[1478540048] HOST ALERT: prayharroldpk;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1478540062] HOST ALERT: ADFS;DOWN;SOFT;5;CRITICAL - Host Unreachable (164.76.4.57)
[1478540069] SERVICE ALERT: ADFS;Uptime;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478540069] SERVICE ALERT: ADFS;Service: Firewall;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478540085] SERVICE ALERT: ADFS;CPU Load;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478540085] SERVICE ALERT: ADFS;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478540088] HOST ALERT: ADFS;DOWN;SOFT;6;CRITICAL - Host Unreachable (164.76.4.57)
[1478540116] SERVICE ALERT: ADFS;C:\ Drive Space;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478540116] SERVICE ALERT: ADFS;Service: NetLogin;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478540119] HOST ALERT: ADFS;DOWN;SOFT;7;CRITICAL - Host Unreachable (164.76.4.57)
[1478540138] HOST ALERT: prayharroldpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1478540173] SERVICE ALERT: ADFS;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478540173] SERVICE ALERT: ADFS;Service: User Profile;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478540176] HOST ALERT: ADFS;DOWN;SOFT;8;CRITICAL - Host Unreachable (164.76.4.57)
[1478540183] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1478540205] SERVICE ALERT: ADFS;CPU Load;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478540205] SERVICE ALERT: ADFS;Service: RPC;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478540208] HOST ALERT: ADFS;DOWN;SOFT;9;CRITICAL - Host Unreachable (164.76.4.57)
[1478540228] HOST ALERT: prayharroldpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1478540236] SERVICE ALERT: ADFS;C:\ Drive Space;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478540236] SERVICE ALERT: ADFS;Service: NetLogin;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478540239] HOST ALERT: ADFS;DOWN;HARD;10;CRITICAL - Host Unreachable (164.76.4.57)
[1478540256] SERVICE ALERT: ADFS;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478540256] SERVICE ALERT: ADFS;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478540281] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1478540293] SERVICE ALERT: ADFS;Service: User Profile;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478540293] SERVICE ALERT: ADFS;Memory Usage;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478540311] HOST ALERT: prayharroldpk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1478540313] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1478540325] SERVICE ALERT: ADFS;CPU Load;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478540325] SERVICE ALERT: ADFS;Service: RPC;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478540326] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6343 found since last run. oneLevelSearchOps OK. searchOps is too high! 6389 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478540331] SERVICE ALERT: ADFS;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478540341] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1478540343] HOST ALERT: prayharroldpk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1478540346] SERVICE ALERT: ADFS;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478540346] SERVICE ALERT: ADFS;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478540346] SERVICE ALERT: ADFS;VMWare Tools;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478540373] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1478540376] SERVICE ALERT: ADFS;Service: WMI;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478540376] SERVICE ALERT: ADFS;NSClient++ Version;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478540401] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1478540403] HOST ALERT: prayharroldpk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1478540413] SERVICE ALERT: ADFS;Memory Usage;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478540413] SERVICE ALERT: ADFS;Service: User Profile;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478540433] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1478540446] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478540451] SERVICE ALERT: ADFS;Service: Group Policy;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478540463] HOST ALERT: prayharroldpk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1478540466] SERVICE ALERT: ADFS;Service: Windows Time;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478540466] SERVICE ALERT: ADFS;Service: DNS Client;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478540466] SERVICE ALERT: ADFS;VMWare Tools;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478540472] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1478540496] SERVICE ALERT: ADFS;Service: WMI;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478540496] SERVICE ALERT: ADFS;NSClient++ Version;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478540502] HOST ALERT: prayharroldpk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1478540532] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1478540562] HOST ALERT: prayharroldpk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1478540571] SERVICE ALERT: ADFS;Service: Group Policy;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478540582] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8684 found since last run. oneLevelSearchOps OK. searchOps is too high! 9424 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478540586] SERVICE ALERT: ADFS;VMWare Tools;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478540586] SERVICE ALERT: ADFS;Service: DNS Client;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478540592] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1478540595] SERVICE ALERT: ADFS;Service: Windows Time;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478540622] HOST ALERT: prayharroldpk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1478540702] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478540747] SERVICE ALERT: studentcenterpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1478540807] SERVICE ALERT: studentcenterpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1478540867] SERVICE ALERT: studentcenterpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1478540995] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 26% - OK!
[1478540999] HOST ALERT: prayharroldpk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.54 ms
[1478541027] SERVICE ALERT: prayharroldpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1478541046] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6977 found since last run. oneLevelSearchOps OK. searchOps is too high! 7027 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478541166] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478541186] SERVICE ALERT: prayharroldpk;Device Status;OK;HARD;3;Ready
[1478541302] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5315 found since last run. oneLevelSearchOps OK. searchOps is too high! 5920 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478541422] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478541654] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1478541709] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1478541713] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1478541762] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1478541766] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 11145 found since last run. oneLevelSearchOps OK. searchOps is too high! 11192 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478541769] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1478541773] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1478541822] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1478541829] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1478541878] SERVICE ALERT: hallefirst3;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1478541882] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1478541882] SERVICE FLAPPING ALERT: hallefirst2;Black Print Cartridge HP C8543X;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1478541886] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478541928] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1478541938] SERVICE ALERT: hallefirst3;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1478541983] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1478541988] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1478541998] SERVICE ALERT: hallefirst3;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1478542028] SERVICE ALERT: hallefirst3;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1478542048] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1478542088] SERVICE ALERT: hallefirst3;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1478542148] SERVICE ALERT: hallefirst3;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1478542277] HOST ALERT: hallefirstcolor;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1478542280] SERVICE ALERT: hallefirstcolor;Toner Collection Unit HP CE265A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1478542310] HOST ALERT: hallefirstcolor;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1478542335] SERVICE ALERT: hallefirstcolor;Toner Collection Unit HP CE265A;OK;SOFT;2;
[1478542339] HOST ALERT: hallefirstcolor;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1478542486] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9555 found since last run. oneLevelSearchOps OK. searchOps is too high! 9600 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478542583] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1478542593] SERVICE ALERT: hallefirst3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1478542606] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478542622] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5602 found since last run. oneLevelSearchOps OK. searchOps is too high! 6132 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478542643] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 24% - OK!
[1478542667] SERVICE ALERT: studentcenterpk;Device Status;OK;HARD;3;TRAY 2 OPEN  46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1478542742] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478542742] SERVICE ALERT: hallefirst3;Device Status;OK;HARD;3;Ready
[1478542967] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1478543023] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Processing job from tray 4
[1478543077] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 82% - OK!
[1478543206] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9930 found since last run. oneLevelSearchOps OK. searchOps is too high! 9976 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478543326] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478543584] SERVICE ALERT: halleground3;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1478543644] SERVICE ALERT: halleground3;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1478543704] SERVICE ALERT: halleground3;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1478543926] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8346 found since last run. oneLevelSearchOps OK. searchOps is too high! 8400 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478543942] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6741 found since last run. oneLevelSearchOps OK. searchOps is too high! 7242 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478544046] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478544062] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478544223] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1478544283] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1478544343] SERVICE ALERT: hallefirst2;Device Status;OK;SOFT;3;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1478544646] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10680 found since last run. oneLevelSearchOps OK. searchOps is too high! 10731 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478544662] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 20064 found since last run. oneLevelSearchOps OK. searchOps is too high! 20623 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478544766] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478544782] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478544983] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 55% (15 min average)
[1478545123] SERVICE ALERT: rooseveltpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.239.219. Is the printer off?  Is the network down?
[1478545153] HOST ALERT: rooseveltpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1478545178] SERVICE ALERT: rooseveltpk;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 36% - OK!
[1478545182] HOST ALERT: rooseveltpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1478545366] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9180 found since last run. oneLevelSearchOps OK. searchOps is too high! 9232 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478545369] SERVICE ALERT: rooseveltpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.239.219. Is the printer off?  Is the network down?
[1478545382] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 17642 found since last run. oneLevelSearchOps OK. searchOps is too high! 18382 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478545429] SERVICE ALERT: rooseveltpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.239.219. Is the printer off?  Is the network down?
[1478545486] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478545489] SERVICE ALERT: rooseveltpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.239.219. Is the printer off?  Is the network down?
[1478545502] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478545579] SERVICE ALERT: rooseveltpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.239.219. Is the printer off?  Is the network down?
[1478545639] SERVICE ALERT: rooseveltpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.239.219. Is the printer off?  Is the network down?
[1478545699] SERVICE ALERT: rooseveltpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.239.219. Is the printer off?  Is the network down?
[1478545783] SERVICE ALERT: rooseveltpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.239.219. Is the printer off?  Is the network down?
[1478545843] SERVICE ALERT: rooseveltpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.239.219. Is the printer off?  Is the network down?
[1478545863] HOST ALERT: reflin;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1478545903] SERVICE ALERT: rooseveltpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.239.219. Is the printer off?  Is the network down?
[1478545953] HOST ALERT: reflin;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1478546043] HOST ALERT: reflin;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1478546086] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6497 found since last run. oneLevelSearchOps OK. searchOps is too high! 6557 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478546102] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 11680 found since last run. oneLevelSearchOps OK. searchOps is too high! 12217 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478546133] HOST ALERT: reflin;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1478546183] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478546206] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478546222] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478546223] HOST ALERT: reflin;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1478546289] HOST ALERT: reflin;UP;SOFT;6;PING OK - Packet loss = 28%, RTA = 0.26 ms
[1478546303] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1478546423] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 92% (15 min average)
[1478546448] SERVICE ALERT: marshallpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.137.49. Is the printer off?  Is the network down?
[1478546497] SERVICE ALERT: rooseveltpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 36% - OK!
[1478546508] SERVICE ALERT: marshallpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.137.49. Is the printer off?  Is the network down?
[1478546568] SERVICE ALERT: marshallpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.137.49. Is the printer off?  Is the network down?
[1478546684] SERVICE ALERT: rooseveltpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1478546740] SERVICE ALERT: marshallpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.137.49. Is the printer off?  Is the network down?
[1478546800] SERVICE ALERT: marshallpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.137.49. Is the printer off?  Is the network down?
[1478546806] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9248 found since last run. oneLevelSearchOps OK. searchOps is too high! 9336 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478546822] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7357 found since last run. oneLevelSearchOps OK. searchOps is too high! 7890 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478546830] HOST ALERT: marshallpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1478546860] SERVICE ALERT: marshallpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.137.49. Is the printer off?  Is the network down?
[1478546890] HOST ALERT: marshallpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1478546893] SERVICE ALERT: rooseveltpk;Device Status;OK;HARD;3;Ready
[1478546926] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478546942] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478546954] HOST ALERT: marshallpk;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.54 ms
[1478547023] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1478547143] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1478547162] SERVICE ALERT: marshallpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 8% - OK!
[1478547263] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1478547303] SERVICE ALERT: halleground3;Device Status;OK;HARD;3;RESUME  CANCEL CURRENT  JOB
[1478547455] SERVICE ALERT: marshallpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1478547499] SERVICE ALERT: rooseveltpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.239.219. Is the printer off?  Is the network down?
[1478547526] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 11688 found since last run. oneLevelSearchOps OK. searchOps is too high! 11769 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478547529] HOST ALERT: rooseveltpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1478547542] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5164 found since last run. oneLevelSearchOps OK. searchOps is too high! 5546 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478547559] SERVICE ALERT: rooseveltpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.239.219. Is the printer off?  Is the network down?
[1478547563] HOST ALERT: rooseveltpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.46 ms
[1478547619] SERVICE ALERT: rooseveltpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.239.219. Is the printer off?  Is the network down?
[1478547646] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478547662] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478547703] SERVICE ALERT: rooseveltpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.239.219. Is the printer off?  Is the network down?
[1478547763] SERVICE ALERT: rooseveltpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.239.219. Is the printer off?  Is the network down?
[1478547823] SERVICE ALERT: rooseveltpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.239.219. Is the printer off?  Is the network down?
[1478547863] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 64% (15 min average)
[1478547889] SERVICE ALERT: rooseveltpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.239.219. Is the printer off?  Is the network down?
[1478547949] SERVICE ALERT: rooseveltpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.239.219. Is the printer off?  Is the network down?
[1478548009] SERVICE ALERT: rooseveltpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.239.219. Is the printer off?  Is the network down?
[1478548246] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8886 found since last run. oneLevelSearchOps OK. searchOps is too high! 9067 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478548262] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9491 found since last run. oneLevelSearchOps OK. searchOps is too high! 9742 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478548366] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478548372] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1478548382] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478548432] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1478548482] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1478548492] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1478548543] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1478548549] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1478548602] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1478548609] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1478548669] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1478548966] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6525 found since last run. oneLevelSearchOps OK. searchOps is too high! 6797 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478548982] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6241 found since last run. oneLevelSearchOps OK. searchOps is too high! 6554 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478549086] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478549102] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478549603] HOST ALERT: hallefirst2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1478549663] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478549686] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 11745 found since last run. oneLevelSearchOps OK. searchOps is too high! 11940 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478549693] HOST ALERT: hallefirst2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1478549783] HOST ALERT: hallefirst2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1478549783] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1478549806] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478549831] HOST ALERT: hallefirst2;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1478549899] HOST ALERT: hallefirst2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1478549903] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 68% (15 min average)
[1478549989] HOST ALERT: hallefirst2;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1478550079] HOST ALERT: hallefirst2;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1478550169] HOST ALERT: hallefirst2;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1478550259] HOST ALERT: hallefirst2;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1478550302] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5842 found since last run. oneLevelSearchOps OK. searchOps is too high! 6275 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478550322] HOST ALERT: hallefirst2;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1478550406] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 16460 found since last run. oneLevelSearchOps OK. searchOps is too high! 16673 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478550422] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478550526] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478550774] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1478550835] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1478550894] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1478551022] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6531 found since last run. oneLevelSearchOps OK. searchOps is too high! 6902 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478551103] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478551126] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10015 found since last run. oneLevelSearchOps OK. searchOps is too high! 10242 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478551142] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478551223] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1478551246] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478551343] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1478551742] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5973 found since last run. oneLevelSearchOps OK. searchOps is too high! 6306 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478551846] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7674 found since last run. oneLevelSearchOps OK. searchOps is too high! 7954 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478551862] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478551966] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478552462] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7103 found since last run. oneLevelSearchOps OK. searchOps is too high! 7685 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478552582] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478553143] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 86% (15 min average)
[1478553182] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 11684 found since last run. oneLevelSearchOps OK. searchOps is too high! 12113 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478553302] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478553743] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1478553863] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 92% (15 min average)
[1478553902] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7899 found since last run. oneLevelSearchOps OK. searchOps is too high! 8259 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478554022] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478554366] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5125 found since last run. oneLevelSearchOps OK. searchOps is too high! 5188 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478554463] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478554486] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478554583] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 83% (15 min average)
[1478554622] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8923 found since last run. oneLevelSearchOps OK. searchOps is too high! 9312 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478554742] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478555086] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8777 found since last run. oneLevelSearchOps OK. searchOps is too high! 8837 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478555183] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478555201] SERVICE FLAPPING ALERT: hallefirst2;Black Print Cartridge HP C8543X;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1478555206] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478555303] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 74% (15 min average)
[1478555342] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9397 found since last run. oneLevelSearchOps OK. searchOps is too high! 9766 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478555428] HOST ALERT: sillpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1478555430] SERVICE ALERT: sillpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1478555460] HOST ALERT: sillpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1478555462] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478555490] SERVICE ALERT: sillpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1478555520] HOST ALERT: sillpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1478555550] SERVICE ALERT: sillpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1478555580] HOST ALERT: sillpk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1478555670] HOST ALERT: sillpk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1478555700] SERVICE ALERT: sillpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1478555730] HOST ALERT: sillpk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1478555806] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7761 found since last run. oneLevelSearchOps OK. searchOps is too high! 7827 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478555820] HOST ALERT: sillpk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1478555903] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478555910] HOST ALERT: sillpk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1478555926] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478555956] SERVICE ALERT: sillpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1478555986] HOST ALERT: sillpk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1478556016] SERVICE ALERT: sillpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1478556023] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1478556046] HOST ALERT: sillpk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1478556062] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7450 found since last run. oneLevelSearchOps OK. searchOps is too high! 7894 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478556076] SERVICE ALERT: sillpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1478556143] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1478556182] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478556526] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6984 found since last run. oneLevelSearchOps OK. searchOps is too high! 7048 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478556647] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478556782] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7513 found since last run. oneLevelSearchOps OK. searchOps is too high! 8122 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478556902] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478557054] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1478557115] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1478557175] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1478557245] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9585 found since last run. oneLevelSearchOps OK. searchOps is too high! 9646 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478557343] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 83% (15 min average)
[1478557365] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478557502] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7415 found since last run. oneLevelSearchOps OK. searchOps is too high! 7888 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478557622] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478557834] HOST ALERT: sillpk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.54 ms
[1478557870] SERVICE ALERT: sillpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 11% - OK!
[1478557943] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1478557945] SERVICE ALERT: sillpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1478557965] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10224 found since last run. oneLevelSearchOps OK. searchOps is too high! 10292 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478558063] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1478558084] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478558094] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Ready
[1478558183] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1478558222] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5374 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478558342] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478558683] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8037 found since last run. oneLevelSearchOps OK. searchOps is too high! 8107 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478558761] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1478558803] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478558821] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1478558881] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1478558942] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5422 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478559062] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478559383] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1478559403] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6661 found since last run. oneLevelSearchOps OK. searchOps is too high! 6721 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478559468] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1478559473] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1478559476] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 54% - OK!
[1478559522] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;OK;SOFT;2;Yellow Cartridge HP CE262A is at 67% - OK!
[1478559523] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478559527] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;OK;SOFT;2;Magenta Cartridge HP CE263A is at 97% - OK!
[1478559662] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6753 found since last run. oneLevelSearchOps OK. searchOps is too high! 7222 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478559776] SERVICE ALERT: cava2;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1478559783] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478559896] SERVICE ALERT: cava2;CPU Load;OK;SOFT;2;CPU Load 92% (15 min average)
[1478559983] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1478560382] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5920 found since last run. oneLevelSearchOps OK. searchOps is too high! 6339 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478560501] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478560723] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8612 found since last run. oneLevelSearchOps OK. searchOps is too high! 8669 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478560843] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478561101] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5102 found since last run. oneLevelSearchOps OK. searchOps is too high! 5516 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478561221] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478561821] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5601 found since last run. oneLevelSearchOps OK. searchOps is too high! 6055 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478561941] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478562541] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5607 found since last run. oneLevelSearchOps OK. searchOps is too high! 6029 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478562643] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7431 found since last run. oneLevelSearchOps OK. searchOps is too high! 7558 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478562661] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478562763] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478563175] SERVICE FLAPPING ALERT: marshallpk;Paper Tray Status;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1478563261] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8234 found since last run. oneLevelSearchOps OK. searchOps is too high! 8622 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478563363] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5474 found since last run. oneLevelSearchOps OK. searchOps is too high! 5664 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478563381] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478563483] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478563496] SERVICE ALERT: cava2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1478563604] SERVICE ALERT: rooseveltpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1478563616] SERVICE ALERT: cava2;CPU Load;OK;SOFT;2;CPU Load 1% (15 min average)
[1478563811] SERVICE ALERT: rooseveltpk;Device Status;OK;HARD;3;Ready
[1478563981] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5327 found since last run. oneLevelSearchOps OK. searchOps is too high! 5595 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478564017] SERVICE ALERT: rooseveltpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 36% - OK!
[1478564083] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7460 found since last run. oneLevelSearchOps OK. searchOps is too high! 7610 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478564101] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478564183] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 78% (15 min average)
[1478564203] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478564701] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5747 found since last run. oneLevelSearchOps OK. searchOps is too high! 6140 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478564783] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478564803] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6421 found since last run. oneLevelSearchOps OK. searchOps is too high! 6563 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478564821] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478564903] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1478564923] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478565023] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1478565421] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7287 found since last run. oneLevelSearchOps OK. searchOps is too high! 7609 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478565523] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5203 found since last run. oneLevelSearchOps OK. searchOps is too high! 5367 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478565541] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478565643] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478566141] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7822 found since last run. oneLevelSearchOps OK. searchOps is too high! 8101 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478566212] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1478566223] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 87% (15 min average)
[1478566234] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1478566242] HOST ALERT: hallefirstcolor;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1478566243] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5020 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478566261] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478566272] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1478566279] HOST ALERT: hallefirstcolor;UP;SOFT;2;PING OK - Packet loss = 37%, RTA = 0.40 ms
[1478566288] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Black Cartridge HP CE260X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1478566327] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;OK;SOFT;3;
[1478566353] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1478566363] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478566383] HOST ALERT: hallefirstcolor;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1478566447] HOST ALERT: hallefirstcolor;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1478566447] HOST FLAPPING ALERT: hallefirstcolor;STARTED; Host appears to have started flapping (22.8% change > 20.0% threshold)
[1478566555] SERVICE ALERT: hallefirstcolor;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1478566573] SERVICE ALERT: hallefirstcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1478566585] HOST ALERT: hallefirstcolor;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1478566589] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1478566615] SERVICE ALERT: hallefirstcolor;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1478566619] HOST ALERT: hallefirstcolor;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1478566633] SERVICE ALERT: hallefirstcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1478566649] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1478566663] HOST ALERT: hallefirstcolor;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1478566675] SERVICE ALERT: hallefirstcolor;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1478566693] SERVICE ALERT: hallefirstcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1478566705] HOST ALERT: hallefirstcolor;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1478566709] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1478566727] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1478566732] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1478566739] HOST ALERT: hallefirstcolor;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1478566771] SERVICE ALERT: hallefirstcolor;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1478566787] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1478566792] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1478566801] HOST ALERT: hallefirstcolor;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1478566801] HOST FLAPPING ALERT: hallefirstcolor;STOPPED; Host appears to have stopped flapping (4.0% change < 5.0% threshold)
[1478566823] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478566831] SERVICE ALERT: hallefirstcolor;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1478566847] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1478566852] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1478566861] HOST ALERT: hallefirstcolor;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1478566891] SERVICE ALERT: hallefirstcolor;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1478566921] HOST ALERT: hallefirstcolor;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1478566932] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1478566939] SERVICE ALERT: hallefirstcolor;Toner Collection Unit HP CE265A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1478566943] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1478566962] HOST ALERT: hallefirstcolor;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1478566963] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5168 found since last run. oneLevelSearchOps OK. searchOps is too high! 5340 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478566991] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1478566999] SERVICE ALERT: hallefirstcolor;Toner Collection Unit HP CE265A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1478567021] HOST ALERT: hallefirstcolor;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1478567051] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1478567059] SERVICE ALERT: hallefirstcolor;Toner Collection Unit HP CE265A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1478567063] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1478567083] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478567663] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1478567683] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5240 found since last run. oneLevelSearchOps OK. searchOps is too high! 5402 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478567803] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478568403] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6266 found since last run. oneLevelSearchOps OK. searchOps is too high! 6410 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478568523] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478568863] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1478568983] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1478569103] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 94% (15 min average)
[1478569123] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7592 found since last run. oneLevelSearchOps OK. searchOps is too high! 7730 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478569243] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478569843] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6593 found since last run. oneLevelSearchOps OK. searchOps is too high! 6739 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478569963] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478570563] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5380 found since last run. oneLevelSearchOps OK. searchOps is too high! 5510 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478570683] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478571061] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5438 found since last run. oneLevelSearchOps OK. searchOps is too high! 5748 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478571181] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478571283] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5283 found since last run. oneLevelSearchOps OK. searchOps is too high! 5412 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478571403] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478571503] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1478571623] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1478571743] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 58% (15 min average)
[1478572343] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478572463] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1478572583] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1478573581] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5201 found since last run. oneLevelSearchOps OK. searchOps is too high! 5489 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478573701] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478573783] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 70% (15 min average)
[1478574301] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5801 found since last run. oneLevelSearchOps OK. searchOps is too high! 6097 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478574383] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478574421] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478574503] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1478575103] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478575223] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1478575343] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1478577743] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1478578343] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1478580743] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 91% (15 min average)