[1454911659] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1454911719] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1454911746] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1454911779] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1454911807] SERVICE ALERT: eateriespk;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! 1 of 3 Trays Still Full
[1454912378] SERVICE ALERT: eateriespk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1454932541] HOST ALERT: halleground1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1454932564] HOST ALERT: halleground1;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1454932631] HOST ALERT: halleground1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1454932654] HOST ALERT: halleground1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1454932673] HOST ALERT: halleground1;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1454932744] HOST ALERT: halleground1;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1454932751] HOST ALERT: halleground1;UP;SOFT;4;PING OK - Packet loss = 73%, RTA = 0.55 ms
[1454932804] HOST ALERT: halleground1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1454932834] HOST ALERT: halleground1;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1454932837] HOST ALERT: halleground1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1454932924] HOST ALERT: halleground1;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1454932927] HOST ALERT: halleground1;DOWN;SOFT;3;(Host check timed out after 30.00 seconds)
[1454933014] HOST ALERT: halleground1;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1454933017] HOST ALERT: halleground1;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1454933074] HOST ALERT: halleground1;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1454933107] HOST ALERT: halleground1;DOWN;SOFT;5;(Host check timed out after 30.00 seconds)
[1454933164] HOST ALERT: halleground1;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1454933197] HOST ALERT: halleground1;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1454933254] HOST ALERT: halleground1;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1454933273] HOST ALERT: halleground1;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1454933344] HOST ALERT: halleground1;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1454933363] HOST ALERT: halleground1;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1454933404] HOST ALERT: halleground1;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1454933437] HOST ALERT: halleground1;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1454933865] HOST ALERT: halleground1;UP;HARD;10;PING OK - Packet loss = 78%, RTA = 0.56 ms
[1454934004] HOST ALERT: halleground1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1454934037] HOST ALERT: halleground1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1454934127] HOST ALERT: halleground1;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1454934217] HOST ALERT: halleground1;DOWN;SOFT;4;(Host check timed out after 30.00 seconds)
[1454934307] HOST ALERT: halleground1;DOWN;SOFT;5;(Host check timed out after 30.00 seconds)
[1454934397] HOST ALERT: halleground1;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1454934473] HOST ALERT: halleground1;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1454934563] HOST ALERT: halleground1;DOWN;SOFT;8;(Host check timed out after 30.00 seconds)
[1454934604] HOST ALERT: halleground1;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1454934637] HOST ALERT: halleground1;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1454936673] HOST ALERT: halleground1;UP;HARD;10;PING WARNING - Packet loss = 83%, RTA = 0.72 ms
[1454936741] HOST ALERT: halleground1;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.56 ms
[1454937004] HOST ALERT: halleground1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1454937037] HOST ALERT: halleground1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1454937101] HOST ALERT: halleground1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.56 ms
[1454937438] SERVICE ALERT: marshallpk;Black Cartridge HP CC364X;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.137.101! Make sure host is up and SNMP is configured properly.
[1454937468] HOST ALERT: marshallpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1454937492] SERVICE ALERT: marshallpk;Black Cartridge HP CC364X;OK;SOFT;2;Black Cartridge HP CC364X is at 61% - OK!
[1454937496] HOST ALERT: marshallpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.41 ms
[1454938447] SERVICE ALERT: ADWU1;C:\ Drive Space;OK;HARD;3;c:\ - total: 399.66 Gb - used: 315.14 Gb (79%) - free 84.52 Gb (21%)
[1454938554] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 12072;OK;HARD;3;Eventlog check ok
[1454938603] SERVICE ALERT: LANSWEEPER;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1454938633] HOST ALERT: LANSWEEPER;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1454938639] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 27% - OK!
[1454938645] HOST ALERT: LANSWEEPER;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1454938649] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 12072;OK;HARD;3;Eventlog check ok
[1454938663] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;OK;HARD;3;OK: No entries found
[1454938693] SERVICE ALERT: LANSWEEPER;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1454938694] SERVICE ALERT: LANSWEEPER;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1454938723] HOST ALERT: LANSWEEPER;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1454938723] HOST ALERT: LANSWEEPER;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1454938723] SERVICE ALERT: LANSWEEPER;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1454938752] SERVICE ALERT: LANSWEEPER;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1454938753] HOST ALERT: LANSWEEPER;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1454938756] SERVICE ALERT: LANSWEEPER;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1454938768] SERVICE ALERT: halleground1;Device Status;OK;HARD;3;Ready
[1454938780] SERVICE ALERT: LANSWEEPER;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1454938781] SERVICE ALERT: LANSWEEPER;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1454938782] HOST ALERT: LANSWEEPER;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1454938786] HOST ALERT: LANSWEEPER;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1454938808] SERVICE ALERT: halleground1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1454938813] SERVICE ALERT: LANSWEEPER;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1454938814] SERVICE ALERT: LANSWEEPER;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1454938816] SERVICE ALERT: LANSWEEPER;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1454938816] SERVICE ALERT: LANSWEEPER;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1454938843] HOST ALERT: LANSWEEPER;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1454938843] SERVICE ALERT: LANSWEEPER;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1454938846] HOST ALERT: LANSWEEPER;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1454938848] SERVICE ALERT: LANSWEEPER;Service: VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1454938872] SERVICE ALERT: LANSWEEPER;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1454938873] SERVICE ALERT: LANSWEEPER;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1454938874] SERVICE ALERT: LANSWEEPER;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1454938876] SERVICE ALERT: LANSWEEPER;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1454938878] HOST ALERT: LANSWEEPER;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1454938900] SERVICE ALERT: LANSWEEPER;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1454938901] SERVICE ALERT: LANSWEEPER;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1454938904] HOST ALERT: LANSWEEPER;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1454938930] HOST ALERT: LANSWEEPER;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1454938933] SERVICE ALERT: LANSWEEPER;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1454938934] SERVICE ALERT: LANSWEEPER;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1454938936] SERVICE ALERT: LANSWEEPER;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1454938936] SERVICE ALERT: LANSWEEPER;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1454938936] SERVICE ALERT: LANSWEEPER;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1454938944] SERVICE ALERT: LANSWEEPER;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1454938954] SERVICE ALERT: LANSWEEPER;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1454938963] HOST ALERT: LANSWEEPER;DOWN;SOFT;7;(Host check timed out after 30.00 seconds)
[1454938966] HOST ALERT: LANSWEEPER;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1454938967] SERVICE ALERT: LANSWEEPER;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1454938968] SERVICE ALERT: LANSWEEPER;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1454938968] SERVICE ALERT: LANSWEEPER;Service: VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1454938992] SERVICE ALERT: LANSWEEPER;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1454938993] SERVICE ALERT: LANSWEEPER;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1454938994] SERVICE ALERT: LANSWEEPER;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1454938994] SERVICE ALERT: LANSWEEPER;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1454938996] SERVICE ALERT: LANSWEEPER;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1454938997] HOST ALERT: LANSWEEPER;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1454938998] HOST ALERT: LANSWEEPER;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1454939008] SERVICE ALERT: LANSWEEPER;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1454939012] SERVICE ALERT: halleground1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1454939020] SERVICE ALERT: LANSWEEPER;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1454939021] SERVICE ALERT: LANSWEEPER;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1454939038] SERVICE ALERT: halleground1;Device Status;OK;HARD;3;Ready
[1454939038] HOST ALERT: LANSWEEPER;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1454939050] HOST ALERT: LANSWEEPER;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1454939056] SERVICE ALERT: LANSWEEPER;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1454939056] SERVICE ALERT: LANSWEEPER;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1454939056] SERVICE ALERT: LANSWEEPER;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1454939059] SERVICE ALERT: LANSWEEPER;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1454939064] SERVICE ALERT: LANSWEEPER;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1454939070] SERVICE ALERT: LANSWEEPER;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1454939074] SERVICE ALERT: LANSWEEPER;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1454939074] SERVICE ALERT: LANSWEEPER;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1454939086] HOST ALERT: LANSWEEPER;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1454939086] HOST ALERT: LANSWEEPER;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1454939087] SERVICE ALERT: LANSWEEPER;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1454939088] SERVICE ALERT: LANSWEEPER;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1454939088] SERVICE ALERT: LANSWEEPER;Service: VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1454939113] SERVICE ALERT: LANSWEEPER;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1454939114] SERVICE ALERT: LANSWEEPER;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1454939114] SERVICE ALERT: LANSWEEPER;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1454939117] SERVICE ALERT: LANSWEEPER;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1454939128] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 27% - OK!
[1454939128] SERVICE ALERT: LANSWEEPER;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1454939136] SERVICE ALERT: LANSWEEPER;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1454939168] SERVICE ALERT: LANSWEEPER;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1454939169] SERVICE ALERT: LANSWEEPER;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1454939176] SERVICE ALERT: LANSWEEPER;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1454939179] SERVICE ALERT: LANSWEEPER;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1454939179] SERVICE ALERT: LANSWEEPER;Service: User Profile;CRITICAL;SOFT;2;connect to address 164.76.2.33 and port 12489: Connection refused
[1454939180] SERVICE ALERT: LANSWEEPER;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1454939183] HOST ALERT: LANSWEEPER;UP;HARD;10;PING OK - Packet loss = 68%, RTA = 0.69 ms
[1454939183] SERVICE ALERT: LANSWEEPER;NSClient++ Version;CRITICAL;HARD;3;connect to address 164.76.2.33 and port 12489: Connection refused
[1454939183] SERVICE ALERT: LANSWEEPER;C:\ Drive Space;CRITICAL;SOFT;2;connect to address 164.76.2.33 and port 12489: Connection refused
[1454939183] HOST ALERT: LANSWEEPER;UP;HARD;10;PING OK - Packet loss = 66%, RTA = 0.52 ms
[1454939184] SERVICE ALERT: LANSWEEPER;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1454939187] SERVICE ALERT: LANSWEEPER;Service: Windows Time;CRITICAL;HARD;3;W32Time: Stopped
[1454939188] SERVICE ALERT: LANSWEEPER;Uptime;OK;SOFT;3;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1454939214] SERVICE ALERT: LANSWEEPER;Service: Group Policy;OK;SOFT;3;gpsvc: Started
[1454939217] SERVICE ALERT: LANSWEEPER;Service: RPC;OK;SOFT;2;RpcSs: Started
[1454939228] SERVICE ALERT: LANSWEEPER;Service: RPC;OK;SOFT;3;RpcSs: Started
[1454939236] SERVICE ALERT: LANSWEEPER;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1454939268] SERVICE ALERT: LANSWEEPER;Memory Usage;OK;SOFT;2;Memory usage: total:11711.16 MB - used: 1861.88 MB (16%) - free: 9849.27 MB (84%)
[1454939269] SERVICE ALERT: LANSWEEPER;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1454939279] SERVICE ALERT: LANSWEEPER;Service: NetLogin;OK;SOFT;3;Netlogon: Started
[1454939280] SERVICE ALERT: LANSWEEPER;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1454939290] SERVICE ALERT: LANSWEEPER;Service: User Profile;OK;SOFT;3;ProfSvc: Started
[1454939294] SERVICE ALERT: LANSWEEPER;C:\ Drive Space;OK;SOFT;3;c:\ - total: 49.48 Gb - used: 33.28 Gb (67%) - free 16.20 Gb (33%)
[1454939423] SERVICE ALERT: LANSWEEPER;Service: Windows Time;OK;HARD;3;W32Time: Started
[1454939449] HOST ALERT: ADRODC-LIV;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 7.09 ms
[1454939514] SERVICE ALERT: LANSWEEPER;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 5 minute(s)
[1454939572] SERVICE ALERT: LANSWEEPER;Memory Usage;OK;HARD;3;Memory usage: total:11711.16 MB - used: 3335.28 MB (28%) - free: 8375.87 MB (72%)
[1454939576] SERVICE ALERT: LANSWEEPER;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1454939600] SERVICE ALERT: LANSWEEPER;C:\ Drive Space;OK;HARD;3;c:\ - total: 49.48 Gb - used: 33.50 Gb (68%) - free 15.98 Gb (32%)
[1454939636] SERVICE ALERT: LANSWEEPER;NSClient++ Version;OK;HARD;3;NSClient++ 0.5.0.19 2015-12-10
[1454939636] SERVICE ALERT: LANSWEEPER;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1454939668] SERVICE ALERT: LANSWEEPER;Service: VMWare Tools;OK;HARD;3;VMTools: Started
[1454939693] SERVICE ALERT: LANSWEEPER;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1454939694] SERVICE ALERT: LANSWEEPER;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1454939756] SERVICE ALERT: LANSWEEPER;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1454939764] SERVICE ALERT: LANSWEEPER;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1454939774] SERVICE ALERT: LANSWEEPER;NSClient++ Version;OK;HARD;3;NSClient++ 0.5.0.19 2015-12-10
[1454939787] SERVICE ALERT: LANSWEEPER;Service: Windows Time;OK;HARD;3;W32Time: Started
[1454939934] HOST ALERT: ADRODC-DET;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 8.10 ms
[1454940113] SERVICE ALERT: LANSWEEPER;CPU Load;OK;HARD;3;CPU Load 86% (5 min average)
[1454940201] SERVICE ALERT: LANSWEEPER;CPU Load;OK;HARD;3;CPU Load 62% (5 min average)
[1454940463] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;1;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1454940583] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;2;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1454940703] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;HARD;3;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1454941445] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1454941449] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1454941505] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1454941508] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1454941565] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1454941568] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1454942112] SERVICE ALERT: studentcenterpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1454942166] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Genuine HP cartridge installed
[1454942168] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Paused  To return to Ready press STOP
[1454942455] SERVICE ALERT: studentcenterpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1454943717] SERVICE ALERT: owen101lab3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1454943721] HOST ALERT: owen101lab3;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.57 ms
[1454943724] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 30% - OK!
[1454943737] SERVICE ALERT: owen101lab3;Device Status;OK;HARD;3;Open right  door.
[1454943741] HOST ALERT: owen101lab3;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.59 ms
[1454943951] SERVICE ALERT: owen101lab3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1454944051] HOST ALERT: owen101lab3;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1454944071] HOST ALERT: owen101lab3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1454944099] HOST ALERT: owen101lab3;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1454944124] HOST ALERT: owen101lab3;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1454944189] HOST ALERT: owen101lab3;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1454944214] HOST ALERT: owen101lab3;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1454944279] HOST ALERT: owen101lab3;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1454944304] HOST ALERT: owen101lab3;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1454944322] SERVICE ALERT: owen101lab3;Paper Tray Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1454944330] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1454944343] SERVICE ALERT: owen101lab3;Device Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1454944352] HOST ALERT: owen101lab3;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1454944373] HOST ALERT: owen101lab3;DOWN;SOFT;5;(Host check timed out after 30.00 seconds)
[1454944382] SERVICE ALERT: owen101lab3;Paper Tray Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1454944390] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1454944403] SERVICE ALERT: owen101lab3;Device Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1454944412] HOST ALERT: owen101lab3;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1454944433] HOST ALERT: owen101lab3;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1454944442] SERVICE ALERT: owen101lab3;Paper Tray Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1454944450] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;WARNING;HARD;3;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1454944463] SERVICE ALERT: owen101lab3;Device Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1454944472] HOST ALERT: owen101lab3;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1454944493] HOST ALERT: owen101lab3;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1454944556] SERVICE ALERT: owen101lab3;Paper Tray Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1454944562] HOST ALERT: owen101lab3;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1454944583] HOST ALERT: owen101lab3;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1454944616] SERVICE ALERT: owen101lab3;Paper Tray Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1454944646] HOST ALERT: owen101lab3;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1454944652] HOST ALERT: owen101lab3;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1454944676] SERVICE ALERT: owen101lab3;Paper Tray Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1454944699] HOST ALERT: owen101lab3;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1454944706] HOST ALERT: owen101lab3;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1454947498] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;OK;HARD;3;OK: No entries found
[1454948267] HOST ALERT: owen101lab3;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.55 ms
[1454948271] SERVICE ALERT: owen101lab3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1454948275] HOST ALERT: owen101lab3;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.55 ms
[1454948287] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 30% - OK!
[1454948637] SERVICE ALERT: owen101lab3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1454948643] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 30% - OK!
[1454948876] SERVICE ALERT: owen101lab3;Paper Tray Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1454948893] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1454948899] HOST ALERT: owen101lab3;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1454948906] HOST ALERT: owen101lab3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1454948936] SERVICE ALERT: owen101lab3;Paper Tray Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1454948953] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1454948966] HOST ALERT: owen101lab3;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1454948989] HOST ALERT: owen101lab3;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1454948996] SERVICE ALERT: owen101lab3;Paper Tray Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1454948996] SERVICE FLAPPING ALERT: owen101lab3;Paper Tray Status;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1454949013] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;WARNING;HARD;3;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1454949026] HOST ALERT: owen101lab3;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1454949079] HOST ALERT: owen101lab3;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1454949116] HOST ALERT: owen101lab3;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1454949169] HOST ALERT: owen101lab3;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1454949206] HOST ALERT: owen101lab3;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1454949242] SERVICE ALERT: owen101lab3;Paper Tray Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1454949249] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1454949259] HOST ALERT: owen101lab3;DOWN;SOFT;5;(Host check timed out after 30.00 seconds)
[1454949293] HOST ALERT: owen101lab3;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1454949302] SERVICE ALERT: owen101lab3;Paper Tray Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1454949309] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1454949332] HOST ALERT: owen101lab3;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1454949362] SERVICE ALERT: owen101lab3;Paper Tray Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1454949362] SERVICE FLAPPING ALERT: owen101lab3;Paper Tray Status;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1454949369] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;WARNING;HARD;3;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1454949369] SERVICE FLAPPING ALERT: owen101lab3;Black Print Cartridge HP C8543X;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1454949383] HOST ALERT: owen101lab3;DOWN;SOFT;7;(Host check timed out after 30.00 seconds)
[1454949392] HOST ALERT: owen101lab3;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1454949473] HOST ALERT: owen101lab3;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1454949482] HOST ALERT: owen101lab3;DOWN;SOFT;8;(Host check timed out after 30.00 seconds)
[1454949563] HOST ALERT: owen101lab3;DOWN;SOFT;9;(Host check timed out after 30.00 seconds)
[1454949572] HOST ALERT: owen101lab3;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1454949626] HOST ALERT: owen101lab3;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1454949662] HOST ALERT: owen101lab3;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1454950925] SERVICE ALERT: ADDC5;GPO Replication;WARNING;SOFT;1;Warning! 1 policies are not replicating properly!
[1454951329] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1454951342] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1454951389] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1454951402] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1454951450] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1454951462] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1454952049] SERVICE ALERT: hallesouthcommons;Device Status;OK;HARD;3;Ready
[1454952062] SERVICE ALERT: hallesouthcommons;Device Status;OK;HARD;3;Ready
[1454952124] SERVICE ALERT: ADDC5;GPO Replication;WARNING;SOFT;2;Warning! 1 policies are not replicating properly!
[1454953321] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1454953328] SERVICE ALERT: ADDC5;GPO Replication;WARNING;HARD;3;Warning! 1 policies are not replicating properly!
[1454953381] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1454953441] SERVICE ALERT: halleground2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1454953492] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1454953551] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1454953611] SERVICE ALERT: halleground2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1454954211] SERVICE ALERT: halleground2;Device Status;OK;HARD;3;Clearing paper path
[1454954515] SERVICE ALERT: ADDC5;GPO Replication;WARNING;SOFT;1;Warning! 1 policies are not replicating properly!
[1454954642] SERVICE ALERT: halleground2;Device Status;OK;HARD;3;Ready
[1454954811] SERVICE ALERT: ADDC5;GPO Replication;OK;SOFT;2;OK! All GPOs are properly replicated.
[1454958240] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1454958301] SERVICE ALERT: halleground2;Device Status;OK;SOFT;2;Open  tray 4. Remove any  jammed
[1454958309] SERVICE ALERT: halleground3;Device Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1454958339] HOST ALERT: halleground3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1454958350] SERVICE ALERT: halleground3;Device Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1454958367] SERVICE ALERT: halleground3;Paper Tray Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1454958369] SERVICE ALERT: halleground3;Device Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1454958380] HOST ALERT: halleground3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1454958397] HOST ALERT: halleground3;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1454958410] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1454958410] SERVICE ALERT: halleground3;Device Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1454958427] SERVICE ALERT: halleground3;Paper Tray Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1454958429] SERVICE ALERT: halleground3;Device Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1454958440] HOST ALERT: halleground3;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1454958457] HOST ALERT: halleground3;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1454958470] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1454958470] SERVICE ALERT: halleground3;Device Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1454958487] SERVICE ALERT: halleground3;Paper Tray Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1454958500] HOST ALERT: halleground3;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1454958517] HOST ALERT: halleground3;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1454958530] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;WARNING;HARD;3;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1454958537] HOST ALERT: halleground3;UP;SOFT;4;PING OK - Packet loss = 37%, RTA = 0.71 ms
[1454958554] SERVICE ALERT: halleground3;Paper Tray Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1454958584] HOST ALERT: halleground3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1454958607] HOST ALERT: halleground3;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1454958614] SERVICE ALERT: halleground3;Paper Tray Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1454958644] HOST ALERT: halleground3;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1454958674] SERVICE ALERT: halleground3;Paper Tray Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1454958697] HOST ALERT: halleground3;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1454958704] HOST ALERT: halleground3;DOWN;SOFT;3;(Host check timed out after 30.00 seconds)
[1454958787] HOST ALERT: halleground3;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1454958794] HOST ALERT: halleground3;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1454958850] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1454958877] HOST ALERT: halleground3;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1454958884] HOST ALERT: halleground3;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1454958910] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1454958940] HOST ALERT: halleground3;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1454958971] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;WARNING;HARD;3;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1454958974] HOST ALERT: halleground3;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1454959001] HOST ALERT: halleground3;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1454959064] HOST ALERT: halleground3;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1454959091] HOST ALERT: halleground3;UP;SOFT;8;PING OK - Packet loss = 77%, RTA = 16.25 ms
[1454959092] HOST ALERT: halleground3;UP;HARD;10;PING OK - Packet loss = 16%, RTA = 0.52 ms
[1454959160] HOST ALERT: halleground3;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1454959160] HOST FLAPPING ALERT: halleground3;STARTED; Host appears to have started flapping (21.3% change > 20.0% threshold)
[1454959250] HOST ALERT: halleground3;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1454959304] HOST ALERT: halleground3;DOWN;SOFT;3;(Host check timed out after 30.00 seconds)
[1454959368] HOST ALERT: halleground3;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 400.77 ms
[1454959422] HOST ALERT: halleground3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1454959490] SERVICE ALERT: owen101lab1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1454959505] SERVICE ALERT: owen101lab1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1454959512] HOST ALERT: halleground3;DOWN;SOFT;2;(Host check timed out after 30.00 seconds)
[1454959550] SERVICE ALERT: owen101lab1;Device Status;OK;SOFT;2;Clearing paper path
[1454959564] SERVICE ALERT: owen101lab1;Device Status;OK;SOFT;2;Genuine HP cartridge installed
[1454959600] HOST ALERT: halleground3;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1454959659] HOST ALERT: halleground3;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1454959698] HOST ALERT: halleground3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1454959717] HOST ALERT: halleground3;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1454959760] HOST ALERT: halleground3;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1454959807] HOST ALERT: halleground3;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1454959850] HOST ALERT: halleground3;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1454959897] HOST ALERT: halleground3;DOWN;SOFT;7;(Host check timed out after 30.00 seconds)
[1454959904] HOST ALERT: halleground3;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1454959987] HOST ALERT: halleground3;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1454959994] HOST ALERT: halleground3;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1454960077] HOST ALERT: halleground3;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1454960084] HOST ALERT: halleground3;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1454960167] HOST ALERT: halleground3;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1454960174] HOST ALERT: halleground3;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1454960264] HOST ALERT: halleground3;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1454960300] HOST ALERT: halleground3;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1454960360] HOST ALERT: halleground3;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1454960396] SERVICE FLAPPING ALERT: owen101lab3;Paper Tray Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1454960524] SERVICE ALERT: ADDC5;GPO Replication;OK;HARD;3;OK! All GPOs are properly replicated.
[1454960762] SERVICE FLAPPING ALERT: owen101lab3;Paper Tray Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1454960769] SERVICE FLAPPING ALERT: owen101lab3;Black Print Cartridge HP C8543X;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1454960900] HOST FLAPPING ALERT: halleground3;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1454961851] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 1 Trays Empty
[1454961912] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 1 Trays Empty
[1454961971] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;SOFT;3;Printer Paper Is OK! All Trays Still Full
[1454965421] HOST ALERT: halleground3;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.62 ms
[1454965608] HOST ALERT: halleground3;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 1.12 ms
[1454965623] SERVICE ALERT: halleground3;Device Status;OK;HARD;3;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1454965664] SERVICE ALERT: halleground3;Device Status;OK;HARD;3;66.11.00 INPUT DEVICE FAILURE Turn off, check connection, turn on 66.11.00 INPUT DEVICE FAILURE 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1454965682] SERVICE ALERT: halleground3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1454965724] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 21% - OK!
[1454965869] SERVICE ALERT: halleground3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1454966163] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 21% - OK!
[1454969330] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.107.53! Make sure host is up and SNMP is configured properly.
[1454969352] HOST ALERT: hallesecondsouth;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1454969360] HOST ALERT: hallesecondsouth;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1454969390] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.107.53! Make sure host is up and SNMP is configured properly.
[1454969420] HOST ALERT: hallesecondsouth;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1454969442] HOST ALERT: hallesecondsouth;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1454969450] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.107.53! Make sure host is up and SNMP is configured properly.
[1454969480] HOST ALERT: hallesecondsouth;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1454969511] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.107.53! Make sure host is up and SNMP is configured properly.
[1454969532] HOST ALERT: hallesecondsouth;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1454969541] HOST ALERT: hallesecondsouth;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1454969571] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.107.53! Make sure host is up and SNMP is configured properly.
[1454969598] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.107.53! Make sure host is up and SNMP is configured properly.
[1454969601] HOST ALERT: hallesecondsouth;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1454969622] HOST ALERT: hallesecondsouth;DOWN;SOFT;4;(Host check timed out after 30.00 seconds)
[1454969631] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;WARNING;HARD;3;WARNING: No SNMP response from 164.76.107.53! Make sure host is up and SNMP is configured properly.
[1454969656] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.107.53! Make sure host is up and SNMP is configured properly.
[1454969658] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.107.53! Make sure host is up and SNMP is configured properly.
[1454969661] HOST ALERT: hallesecondsouth;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1454969686] HOST ALERT: hallesecondsouth;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1454969711] SERVICE ALERT: hallesecondsouth;Paper Tray Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.107.53! Make sure host is up and SNMP is configured properly.
[1454969712] SERVICE ALERT: hallesecondsouth;Paper Tray Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.107.53! Make sure host is up and SNMP is configured properly.
[1454969716] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.107.53! Make sure host is up and SNMP is configured properly.
[1454969718] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;WARNING;HARD;3;WARNING: No SNMP response from 164.76.107.53! Make sure host is up and SNMP is configured properly.
[1454969741] HOST ALERT: hallesecondsouth;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1454969742] HOST ALERT: hallesecondsouth;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1454969771] SERVICE ALERT: hallesecondsouth;Paper Tray Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.107.53! Make sure host is up and SNMP is configured properly.
[1454969772] SERVICE ALERT: hallesecondsouth;Paper Tray Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.107.53! Make sure host is up and SNMP is configured properly.
[1454969776] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.107.53! Make sure host is up and SNMP is configured properly.
[1454969801] HOST ALERT: hallesecondsouth;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1454969802] HOST ALERT: hallesecondsouth;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1454969831] SERVICE ALERT: hallesecondsouth;Paper Tray Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.107.53! Make sure host is up and SNMP is configured properly.
[1454969832] SERVICE ALERT: hallesecondsouth;Paper Tray Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.107.53! Make sure host is up and SNMP is configured properly.
[1454969861] HOST ALERT: hallesecondsouth;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1454969862] HOST ALERT: hallesecondsouth;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1454969941] HOST ALERT: hallesecondsouth;UP;SOFT;10;PING OK - Packet loss = 75%, RTA = 0.60 ms
[1454969942] HOST ALERT: hallesecondsouth;UP;SOFT;9;PING OK - Packet loss = 77%, RTA = 1.42 ms
[1454970044] SERVICE ALERT: hallesecondsouth;Device Status;OK;HARD;3;Warming up
[1454970225] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 74% - OK!
[1454970230] SERVICE ALERT: ADWSUSSERVERS;Service: WMI;CRITICAL;SOFT;1;Winmgmt: Stopped
[1454970312] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 74% - OK!
[1454970350] SERVICE ALERT: ADWSUSSERVERS;Service: WMI;CRITICAL;SOFT;2;Winmgmt: Stopped
[1454970370] SERVICE ALERT: hallesecondsouth;Device Status;OK;HARD;3;Ready
[1454970426] SERVICE ALERT: hallesecondsouth;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1454970427] SERVICE ALERT: hallesecondsouth;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1454970470] SERVICE ALERT: ADWSUSSERVERS;Service: WMI;CRITICAL;HARD;3;Winmgmt: Stopped
[1454970702] SERVICE ALERT: ADWSUSSERVERS;Service: WMI;CRITICAL;SOFT;1;Winmgmt: Stopped
[1454970822] SERVICE ALERT: ADWSUSSERVERS;Service: WMI;CRITICAL;SOFT;2;Winmgmt: Stopped
[1454970942] SERVICE ALERT: ADWSUSSERVERS;Service: WMI;CRITICAL;HARD;3;Winmgmt: Stopped
[1454971542] SERVICE ALERT: ADWSUSSERVERS;Service: WMI;OK;HARD;3;Winmgmt: Started
[1454971670] SERVICE ALERT: ADWSUSSERVERS;Service: WMI;OK;HARD;3;Winmgmt: Started
[1454973156] HOST ALERT: studentcenterpk;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1454973220] HOST ALERT: studentcenterpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.55 ms
[1454976834] SERVICE ALERT: studentcenterpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1454976893] SERVICE ALERT: studentcenterpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1454976952] SERVICE ALERT: studentcenterpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1454976977] SERVICE ALERT: studentcenterpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1454977037] SERVICE ALERT: studentcenterpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1454977097] SERVICE ALERT: studentcenterpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1454978752] SERVICE ALERT: studentcenterpk;Device Status;OK;HARD;3;Warming up
[1454978897] SERVICE ALERT: studentcenterpk;Device Status;OK;HARD;3;TRAY 3 EMPTY PLAIN LETTER   Ready
[1454980853] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1454980913] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1454980973] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1454981103] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.106.52! Make sure host is up and SNMP is configured properly.
[1454981111] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1454981121] HOST ALERT: halleground1;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1454981133] HOST ALERT: halleground2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1454981134] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.106.51! Make sure host is up and SNMP is configured properly.
[1454981163] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.106.52! Make sure host is up and SNMP is configured properly.
[1454981164] HOST ALERT: halleground1;UP;SOFT;2;PING WARNING - Packet loss = 90%, RTA = 0.73 ms
[1454981171] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1454981176] HOST ALERT: halleground2;UP;SOFT;2;PING OK - Packet loss = 64%, RTA = 0.69 ms
[1454981189] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 24% - OK!
[1454981217] SERVICE ALERT: halleground2;Device Status;OK;SOFT;3;Processing job from tray 4
[1454981231] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1454983356] SERVICE ALERT: SANDBOX-DB;CPU Load;WARNING;SOFT;1;CPU Load 80% (5 min average)
[1454983476] SERVICE ALERT: SANDBOX-DB;CPU Load;OK;SOFT;2;CPU Load 77% (5 min average)
[1454984076] SERVICE ALERT: SANDBOX-DB;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1454984196] SERVICE ALERT: SANDBOX-DB;CPU Load;OK;SOFT;2;CPU Load 55% (5 min average)
[1454984561] SERVICE ALERT: DB01-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1454984681] SERVICE ALERT: DB01-2012;CPU Load;OK;SOFT;2;CPU Load 77% (5 min average)
[1454986562] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1454986565] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1454986622] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1454986625] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1454986683] SERVICE ALERT: hallefirst2;Device Status;OK;SOFT;3;RETRIEVE JOB  NO STORED JOBS
[1454986685] SERVICE ALERT: hallefirst2;Device Status;OK;SOFT;3;RETRIEVE JOB  NO STORED JOBS
[1454987282] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1454987285] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1454987342] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1454987345] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1454987402] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1454987405] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1454987761] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1454987820] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1454987880] SERVICE ALERT: crossroadspk;Device Status;OK;SOFT;3;Press OK to continue.
[1454988558] SERVICE ALERT: VC;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1454988563] SERVICE ALERT: VC;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1454988602] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;MENUS  RETRIEVE JOB  INFORMATION  PAPER HANDLING
[1454988604] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;MENUS  RETRIEVE JOB  INFORMATION  PAPER HANDLING
[1454988658] SERVICE ALERT: VC;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1454988663] SERVICE ALERT: VC;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1454989202] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1454989204] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1454989263] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1454989264] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1454989322] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1454989324] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1454990168] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1454990198] HOST ALERT: halleground3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1454990228] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1454990229] SERVICE ALERT: halleground3;Device Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1454990256] HOST ALERT: halleground3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1454990258] HOST ALERT: halleground3;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1454990270] SERVICE ALERT: halleground3;Device Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1454990287] SERVICE ALERT: halleground3;Paper Tray Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1454990288] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;WARNING;HARD;3;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1454990289] SERVICE ALERT: halleground3;Device Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1454990300] HOST ALERT: halleground3;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1454990317] HOST ALERT: halleground3;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1454990330] SERVICE ALERT: halleground3;Device Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1454990330] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1454990343] SERVICE ALERT: halleground3;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! All Trays Still Full
[1454990343] SERVICE ALERT: halleground3;Device Status;OK;SOFT;3;66.11.00 INPUT DEVICE FAILURE Turn off, check connection, turn on 66.11.00 INPUT DEVICE FAILURE 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1454990346] HOST ALERT: halleground3;UP;SOFT;3;PING OK - Packet loss = 70%, RTA = 0.50 ms
[1454990347] HOST ALERT: halleground3;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1454990390] SERVICE ALERT: halleground3;Device Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1454990391] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1454990420] HOST ALERT: halleground3;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1454990450] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;WARNING;HARD;3;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1454990474] SERVICE ALERT: halleground3;Paper Tray Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1454990480] HOST ALERT: halleground3;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1454990534] SERVICE ALERT: halleground3;Paper Tray Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1454990564] HOST ALERT: halleground3;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1454990594] SERVICE ALERT: halleground3;Paper Tray Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1454990623] HOST ALERT: halleground3;UP;SOFT;4;PING WARNING - Packet loss = 83%, RTA = 0.52 ms
[1454990623] HOST FLAPPING ALERT: halleground3;STARTED; Host appears to have started flapping (20.8% change > 20.0% threshold)
[1454990673] HOST ALERT: halleground3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1454990763] HOST ALERT: halleground3;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1454990853] HOST ALERT: halleground3;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1454990901] HOST ALERT: halleground3;UP;SOFT;4;PING OK - Packet loss = 64%, RTA = 0.52 ms
[1454990901] HOST FLAPPING ALERT: halleground3;STARTED; Host appears to have started flapping (21.8% change > 20.0% threshold)
[1454990947] SERVICE ALERT: halleground3;Paper Tray Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1454990949] SERVICE ALERT: halleground3;Device Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1454990953] HOST ALERT: halleground3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1454990977] HOST ALERT: halleground3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1454991007] SERVICE ALERT: halleground3;Paper Tray Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1454991009] SERVICE ALERT: halleground3;Device Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1454991020] HOST ALERT: halleground3;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1454991037] HOST ALERT: halleground3;DOWN;SOFT;2;(Host check timed out after 30.00 seconds)
[1454991067] SERVICE ALERT: halleground3;Paper Tray Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1454991069] SERVICE ALERT: halleground3;Device Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.106.53! Make sure host is up and SNMP is configured properly.
[1454991080] HOST ALERT: halleground3;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1454991097] HOST ALERT: halleground3;DOWN;SOFT;3;(Host check timed out after 30.00 seconds)
[1454991170] HOST ALERT: halleground3;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1454991187] HOST ALERT: halleground3;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1454991224] HOST ALERT: halleground3;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1454991277] HOST ALERT: halleground3;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1454991314] HOST ALERT: halleground3;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1454991367] HOST ALERT: halleground3;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1454991404] HOST ALERT: halleground3;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1454991457] HOST ALERT: halleground3;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1454991494] HOST ALERT: halleground3;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1454991518] HOST ALERT: halleground3;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1454991584] HOST ALERT: halleground3;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1454991608] HOST ALERT: halleground3;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1454991620] HOST ALERT: halleground3;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1454991697] HOST ALERT: halleground3;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1454991999] HOST FLAPPING ALERT: halleground3;STOPPED; Host appears to have stopped flapping (4.0% change < 5.0% threshold)
[1454992190] HOST FLAPPING ALERT: halleground3;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)