[1473739218] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 537 found since last run.
[1473739338] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473739888] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473740488] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 2510 found since last run. oneLevelSearchOps OK. searchOps is too high! 2539 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 506 found since last run.
[1473740608] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473741138] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 2344 found since last run. oneLevelSearchOps OK. searchOps is too high! 2467 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1473741258] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473741858] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 2051 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1473741978] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473742408] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1533 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1473742528] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473743128] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 2498 found since last run. oneLevelSearchOps OK. searchOps is too high! 2521 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1473743248] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473743848] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 2642 found since last run. oneLevelSearchOps OK. searchOps is too high! 2668 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1473743968] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473743985] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1473744103] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;2;Warning:  idm1  are more than 120 seconds behind!
[1473744223] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1473744568] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1653 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1473744644] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 83% (15 min average)
[1473744688] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473744978] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 505 found since last run.
[1473745098] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473745244] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1473745364] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1473745423] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1473745484] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 88% (15 min average)
[1473746023] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1473746084] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1473746143] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;2;Warning:  idm1  are more than 120 seconds behind!
[1473746204] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1473746263] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1473746324] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1473746488] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1621 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1473746608] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473746863] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1473746863] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1473746924] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 43% (15 min average)
[1473747524] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1473747644] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1473747764] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 65% (15 min average)
[1473748063] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1473748364] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1473748484] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1473748604] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1473748663] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1473749008] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1967 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1473749128] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473749263] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1473749298] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 690 found since last run.
[1473749418] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473750404] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1473750463] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1473750535] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1473750583] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;2;Warning:  idm1  are more than 120 seconds behind!
[1473750655] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1473750703] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1473750775] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1473751004] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1473751124] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 92% (15 min average)
[1473751724] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1473751844] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1473751903] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1473751964] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1473752503] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1473752564] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 72% (15 min average)
[1473752623] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1473753164] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1473753284] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1473753404] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1473753823] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1473753943] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1473754063] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1473754604] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 27% (15 min average)
[1473754816] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1534 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1473754936] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473755263] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1473755863] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1473756463] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1473757374] SERVICE ALERT: epprogsql2;CPU Load;OK;HARD;3;CPU Load 54% (15 min average)
[1473757663] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1473757783] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1473757903] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1473758204] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1473758324] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1473758924] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1473759044] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1473759164] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 24% (15 min average)
[1473759703] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1473760303] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1473760423] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1473760543] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1473761143] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1473761564] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1473761684] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1473761804] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1473762328] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1564 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1473762343] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1473762448] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473762463] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1473762583] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1473762736] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 2148 found since last run. oneLevelSearchOps OK. searchOps is too high! 2268 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 537 found since last run.
[1473762856] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473763604] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 27% (15 min average)
[1473763783] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1473764204] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1473764324] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1473764444] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1473764656] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1003 found since last run.
[1473764776] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473764848] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1392 found since last run.
[1473764968] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473765376] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1513 found since last run.
[1473765496] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473765568] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1570 found since last run.
[1473765583] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1473765688] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473765703] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1473766096] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1569 found since last run.
[1473766216] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473766288] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1558 found since last run.
[1473766408] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473766816] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1631 found since last run.
[1473766936] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473767008] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1596 found since last run.
[1473767128] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473767536] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1510 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 1773 found since last run.
[1473767656] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 587 found since last run.
[1473767728] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1538 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 2270 found since last run.
[1473767776] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1224 found since last run.
[1473767848] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 690 found since last run.
[1473767968] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 814 found since last run.
[1473768415] SERVICE ALERT: cava2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1473768535] SERVICE ALERT: cava2;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1473768655] SERVICE ALERT: cava2;CPU Load;OK;SOFT;3;CPU Load 14% (15 min average)
[1473769643] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473769703] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473769763] SERVICE ALERT: hallegroundcolor;Device Status;OK;SOFT;3;Open right door.
[1473770364] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473770423] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473770444] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 77% (15 min average)
[1473770483] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473771083] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;Open right door.
[1473771644] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1473771764] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1473772055] SERVICE ALERT: sillpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1473772115] SERVICE ALERT: sillpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1473772175] SERVICE ALERT: sillpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1473772283] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473772343] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473772403] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473772417] SERVICE ALERT: alexanderpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473772483] SERVICE ALERT: alexanderpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.142.94. Is the printer off?  Is the network down?
[1473772513] HOST ALERT: alexanderpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1473772543] SERVICE ALERT: alexanderpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.142.94. Is the printer off?  Is the network down?
[1473772567] HOST ALERT: alexanderpk;UP;SOFT;2;PING WARNING - Packet loss = 80%, RTA = 0.55 ms
[1473772664] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1473772723] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1473772776] SERVICE ALERT: sillpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 44% - OK!
[1473772784] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1473772964] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1473773003] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;Open right door.
[1473773003] SERVICE FLAPPING ALERT: hallegroundcolor;Device Status;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1473773084] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 77% (15 min average)
[1473773084] SERVICE ALERT: owen101lab2;Device Status;OK;HARD;3;ORDER CARTRIDGE LESS THAN 1200 PAGES 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1473773137] SERVICE ALERT: alexanderpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473773604] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473773663] SERVICE ALERT: hallegroundcolor;Device Status;OK;SOFT;2;Open right door.
[1473773983] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1473774263] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473774284] SERVICE ALERT: owen101lab2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473774323] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473774338] SERVICE ALERT: alexanderpk;Device Status;OK;HARD;3;Ready
[1473774345] SERVICE ALERT: owen101lab2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473774383] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473774404] SERVICE ALERT: owen101lab2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473774614] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1473774659] SERVICE ALERT: owen101lab1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1473774689] HOST ALERT: owen101lab1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1473774703] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1473774712] SERVICE ALERT: owen101lab1;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1473774719] SERVICE ALERT: owen101lab1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1473774719] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1473774742] HOST ALERT: owen101lab1;UP;SOFT;2;PING WARNING - Packet loss = 90%, RTA = 0.95 ms
[1473774767] SERVICE ALERT: owen101lab1;Device Status;OK;SOFT;2;Processing job
[1473774773] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 25% - OK!
[1473774773] SERVICE ALERT: owen101lab1;Paper Tray Status;OK;SOFT;3;Printer Paper Is OK! All Trays Still Full
[1473774983] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;Open right door.
[1473775004] SERVICE ALERT: owen101lab2;Device Status;OK;HARD;3;ORDER CARTRIDGE LESS THAN 900 PAGES 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1473775484] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1473775604] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1473775724] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1473776183] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473776243] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473776303] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473776904] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;Open right door.
[1473778157] HOST ALERT: hallegroundcolor;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1473778221] HOST ALERT: hallegroundcolor;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1473778595] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473778655] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473778703] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473778716] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473778724] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 91% (15 min average)
[1473778763] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473778823] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473779063] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473779129] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1473779159] HOST ALERT: hallefirst1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1473779183] SERVICE ALERT: hallefirst1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473779187] HOST ALERT: hallefirst1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.52 ms
[1473779315] SERVICE ALERT: hallesouthcommons;Device Status;OK;HARD;3;Open top cover.
[1473779324] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1473779423] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;Open right door.
[1473779444] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1473779564] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1473779783] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;ORDER CARTRIDGE LESS THAN 2000 PAGES 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1473779915] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473779975] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473780035] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473780635] SERVICE ALERT: hallesouthcommons;Device Status;OK;HARD;3;Open top cover.
[1473780635] SERVICE FLAPPING ALERT: hallesouthcommons;Device Status;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1473780649] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1473780679] HOST ALERT: halleground2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1473780709] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1473780739] HOST ALERT: halleground2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1473780769] SERVICE ALERT: halleground2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1473780769] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1473780799] HOST ALERT: halleground2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1473780829] SERVICE ALERT: halleground2;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1473780849] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1473780859] HOST ALERT: halleground2;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1473780889] SERVICE ALERT: halleground2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1473780909] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1473780919] HOST ALERT: halleground2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1473780969] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1473780999] HOST ALERT: halleground2;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1473781089] HOST ALERT: halleground2;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1473781179] HOST ALERT: halleground2;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1473781269] HOST ALERT: halleground2;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1473781359] HOST ALERT: halleground2;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1473781836] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473781895] SERVICE ALERT: hallesouthcommons;Device Status;OK;SOFT;2;Open top cover.
[1473782423] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473782483] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473782543] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473783143] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;Open right door.
[1473783576] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1473783637] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1473783696] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1473784203] SERVICE FLAPPING ALERT: hallegroundcolor;Device Status;STARTED; Service appears to have started flapping (48.7% change >= 20.0% threshold)
[1473784203] SERVICE FLAPPING ALERT: hallesouthcommons;Device Status;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1473784203] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1473784203] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1473784203] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1473784203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1473784203] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1473784203] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1473784203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1473784203] HOST DOWNTIME ALERT: NETIDADPROXY;STARTED; Host has entered a period of scheduled downtime
[1473784203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1473784203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1473784203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1473784203] HOST DOWNTIME ALERT: VUM;STARTED; Host has entered a period of scheduled downtime
[1473784203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1473784203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1473784203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1473784203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1473784203] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1473784203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1473784343] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473784403] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473784463] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473784903] HOST ALERT: halleground2;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.52 ms
[1473784963] SERVICE ALERT: halleground2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1473785063] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;Open right door.
[1473785083] SERVICE ALERT: halleground2;Device Status;OK;HARD;3;Ready 164.76.106.52
[1473785163] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 77% - OK!
[1473786003] SERVICE FLAPPING ALERT: hallegroundcolor;Device Status;STARTED; Service appears to have started flapping (49.7% change >= 20.0% threshold)
[1473786003] SERVICE FLAPPING ALERT: hallesouthcommons;Device Status;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1473786003] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1473786003] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1473786003] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1473786003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1473786003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1473786003] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1473786003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1473786003] HOST DOWNTIME ALERT: NETIDADPROXY;STARTED; Host has entered a period of scheduled downtime
[1473786003] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1473786003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1473786003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1473786003] HOST DOWNTIME ALERT: VUM;STARTED; Host has entered a period of scheduled downtime
[1473786003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1473786003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1473786003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1473786003] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1473786003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1473786003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1473786263] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473786323] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473786383] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473786696] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;CRITICAL;HARD;3;Black Print Cartridge HP C8543X is at 1% - CRITICAL!
[1473786983] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;Open right door.
[1473787004] SERVICE ALERT: owen101lab2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473787064] SERVICE ALERT: owen101lab2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473787125] SERVICE ALERT: owen101lab2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473788184] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473788184] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473788243] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473788243] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473788303] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473788304] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473788851] SERVICE ALERT: hallesecondsouth;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1473788903] SERVICE ALERT: hallesecondsouth;Device Status;OK;HARD;3;Clearing paper path
[1473788910] SERVICE ALERT: hallesecondsouth;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! 2 of 3 Trays Still Full
[1473789503] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;Open right door.
[1473789649] SERVICE ALERT: owen101lab2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1473789679] HOST ALERT: owen101lab2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1473789699] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1473789703] HOST ALERT: owen101lab2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.61 ms
[1473789703] SERVICE ALERT: owen101lab2;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! All Trays Still Full
[1473790103] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473790124] SERVICE ALERT: owen101lab2;Device Status;OK;HARD;3;REPLACE BLACK CARTRIDGE 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Processing job from tray 4
[1473790163] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473790223] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473790296] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;CRITICAL;HARD;3;Black Print Cartridge HP C8543X is at 1% - CRITICAL!
[1473790296] SERVICE FLAPPING ALERT: owen101lab2;Black Print Cartridge HP C8543X;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1473790684] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1473790724] SERVICE ALERT: owen101lab2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473790743] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1473790785] SERVICE ALERT: owen101lab2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473790803] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1473790844] SERVICE ALERT: owen101lab2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473791152] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1473791182] HOST ALERT: hallesecondsouth;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1473791212] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1473791242] HOST ALERT: hallesecondsouth;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1473791272] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1473791302] HOST ALERT: hallesecondsouth;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1473791309] SERVICE ALERT: hallesecondsouth;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1473791316] SERVICE ALERT: hallesecondsouth;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1473791339] HOST ALERT: hallesecondsouth;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1473791369] SERVICE ALERT: hallesecondsouth;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1473791376] SERVICE ALERT: hallesecondsouth;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1473791399] HOST ALERT: hallesecondsouth;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1473791402] SERVICE FLAPPING ALERT: hallegroundcolor;Device Status;STARTED; Service appears to have started flapping (55.5% change >= 20.0% threshold)
[1473791402] SERVICE FLAPPING ALERT: owen101lab2;Black Print Cartridge HP C8543X;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1473791402] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1473791402] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1473791402] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1473791402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1473791402] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1473791402] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1473791402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1473791402] HOST DOWNTIME ALERT: NETIDADPROXY;STARTED; Host has entered a period of scheduled downtime
[1473791402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1473791402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1473791402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1473791402] HOST DOWNTIME ALERT: VUM;STARTED; Host has entered a period of scheduled downtime
[1473791402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1473791402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1473791402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1473791402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1473791402] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1473791402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1473791424] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;Open right door.
[1473791429] SERVICE ALERT: hallesecondsouth;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1473791436] SERVICE ALERT: hallesecondsouth;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1473791459] HOST ALERT: hallesecondsouth;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1473791549] HOST ALERT: hallesecondsouth;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1473791639] HOST ALERT: hallesecondsouth;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1473791729] HOST ALERT: hallesecondsouth;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1473791819] HOST ALERT: hallesecondsouth;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1473791902] HOST ALERT: hallesecondsouth;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1473791992] HOST ALERT: hallesecondsouth;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1473792023] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473792059] HOST ALERT: hallesecondsouth;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1473792083] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473792143] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473792149] HOST ALERT: hallesecondsouth;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1473792743] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;Remove all paper found.
[1473792971] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1473793025] SERVICE ALERT: owen101lab3;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! All Trays Still Full
[1473793244] SERVICE ALERT: owen101lab2;Device Status;OK;HARD;3;REPLACE BLACK CARTRIDGE 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1473793244] SERVICE FLAPPING ALERT: owen101lab2;Device Status;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1473793343] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473793364] SERVICE ALERT: prayharroldpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1473793364] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1473793403] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473793424] SERVICE ALERT: prayharroldpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1473793463] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473793484] SERVICE ALERT: prayharroldpk;Paper Tray Status;OK;SOFT;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1473793641] SERVICE ALERT: R25-WEBSERVICES;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.4.83 and port 12489: Connection refused
[1473793711] SERVICE ALERT: R25-WEBVIEWER;Service: NetLogin;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1473793716] SERVICE ALERT: R25-WEBVIEWER;Service: Windows Time;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1473793761] SERVICE ALERT: R25-WEBSERVICES;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1473793817] SERVICE ALERT: R25-INTERFACE;Service: NetLogin;CRITICAL;SOFT;1;connect to address 164.76.4.82 and port 12489: Connection refused
[1473793823] SERVICE ALERT: R25-INTERFACE;VMWare Tools;CRITICAL;SOFT;1;connect to address 164.76.4.82 and port 12489: Connection refused
[1473793831] SERVICE ALERT: R25-WEBVIEWER;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1473793836] SERVICE ALERT: R25-WEBVIEWER;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1473793937] SERVICE ALERT: R25-INTERFACE;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1473793943] SERVICE ALERT: R25-INTERFACE;VMWare Tools;OK;SOFT;2;VMTools: Started
[1473793964] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1473794063] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;Open right door.
[1473794744] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 1 Trays Empty
[1473794803] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 1 Trays Empty
[1473794862] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 1 Trays Empty
[1473795263] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473795323] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473795383] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473796203] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;CRITICAL;HARD;3;Black Print Cartridge HP C8543X is at 1% - CRITICAL!
[1473796583] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;Open right door.
[1473797183] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473797243] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473797303] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473797903] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;Open right door.
[1473798382] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473798442] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473798502] SERVICE ALERT: hallefirst1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473798603] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1473798940] HOST ALERT: hallesecondsouth;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.45 ms
[1473799066] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 32% - OK!
[1473799102] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;ORDER CARTRIDGE LESS THAN 1900 PAGES 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1473799223] SERVICE ALERT: hallesecondsouth;Device Status;OK;HARD;3;Ready
[1473799231] SERVICE ALERT: hallesecondsouth;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1473799702] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473799762] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473799822] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473800403] SERVICE FLAPPING ALERT: hallegroundcolor;Device Status;STARTED; Service appears to have started flapping (65.1% change >= 20.0% threshold)
[1473800403] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1473800403] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1473800403] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1473800403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1473800403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1473800403] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1473800403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1473800403] HOST DOWNTIME ALERT: NETIDADPROXY;STARTED; Host has entered a period of scheduled downtime
[1473800403] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1473800403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1473800403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1473800403] HOST DOWNTIME ALERT: VUM;STARTED; Host has entered a period of scheduled downtime
[1473800403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1473800403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1473800403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1473800403] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1473800403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1473800403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1473800564] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 87% (15 min average)
[1473801022] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;Open right door.
[1473801164] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1473801284] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 80% (15 min average)
[1473801622] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473801683] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473801742] SERVICE ALERT: hallegroundcolor;Device Status;OK;SOFT;3;Open right door.
[1473801884] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1473802004] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 90% (15 min average)
[1473802341] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473802401] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473802461] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473802604] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1473802724] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1473802823] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473802844] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1473802844] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1473802883] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473802943] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473803661] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;Open right door.
[1473804165] SERVICE ALERT: owen101lab1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473804225] SERVICE ALERT: owen101lab1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473804261] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473804285] SERVICE ALERT: owen101lab1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473804321] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473804382] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473804462] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1473804743] SERVICE ALERT: hallesecondsouth;Device Status;OK;HARD;3;INPUT PAPER PATH OPEN 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1473804885] SERVICE ALERT: owen101lab1;Device Status;OK;HARD;3;TRAY 4 EMPTY PLAIN LETTER   Processing job from tray 3
[1473805244] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 86% (15 min average)
[1473805581] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;Open right door.
[1473805844] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1473805964] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1473806084] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 63% (15 min average)
[1473806182] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473806241] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473806301] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473807031] SERVICE ALERT: hallesecondsouth;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1473807096] SERVICE ALERT: hallesecondsouth;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1473807126] HOST ALERT: hallesecondsouth;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1473807147] HOST ALERT: hallesecondsouth;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.49 ms
[1473807151] SERVICE ALERT: hallesecondsouth;Paper Tray Status;OK;SOFT;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1473807501] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;Open right door.
[1473807743] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473807751] SERVICE ALERT: hallesecondsouth;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1473807803] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473807811] SERVICE ALERT: hallesecondsouth;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1473807863] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473807863] SERVICE FLAPPING ALERT: hallesecondsouth;Device Status;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1473807871] SERVICE ALERT: hallesecondsouth;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1473807884] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1473808004] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1473808101] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473808124] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 74% (15 min average)
[1473808162] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473808221] SERVICE ALERT: hallegroundcolor;Device Status;OK;SOFT;3;Open right door.
[1473808464] SERVICE ALERT: hallesecondsouth;Device Status;OK;HARD;3;Clearing paper path
[1473808471] SERVICE ALERT: hallesecondsouth;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1473808821] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473808881] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473808941] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473809324] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1473809444] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 84% (15 min average)
[1473809541] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;Open right door.
[1473810044] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1473810141] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473810164] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1473810201] SERVICE ALERT: hallegroundcolor;Device Status;OK;SOFT;2;Open right door.
[1473810284] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1473810801] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473810861] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473810884] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 72% (15 min average)
[1473810922] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473811521] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;Open right door.
[1473812121] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473812181] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473812241] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473812841] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;Open right door.
[1473813283] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1473813403] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1473813523] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1473814041] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473814101] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473814162] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473814723] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1473814824] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;SOFT;1;e:\ - total: 699.87 Gb - used: 569.25 Gb (81%) - free 130.62 Gb (19%)
[1473814944] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;SOFT;2;e:\ - total: 699.87 Gb - used: 569.25 Gb (81%) - free 130.62 Gb (19%)
[1473815064] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;HARD;3;e:\ - total: 699.87 Gb - used: 569.25 Gb (81%) - free 130.62 Gb (19%)
[1473815083] SERVICE ALERT: prayharroldpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1473815144] SERVICE ALERT: prayharroldpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1473815204] SERVICE ALERT: prayharroldpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1473815323] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1473815361] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;Open right door.
[1473815443] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1473815563] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1473815664] SERVICE ALERT: AD-RMAD;E:\ Drive Space;OK;HARD;3;e:\ - total: 699.87 Gb - used: 547.60 Gb (78%) - free 152.28 Gb (22%)
[1473815961] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473815982] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1473816021] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473816042] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1473816081] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473816103] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1473816264] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;SOFT;1;e:\ - total: 699.87 Gb - used: 569.25 Gb (81%) - free 130.62 Gb (19%)
[1473816384] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;SOFT;2;e:\ - total: 699.87 Gb - used: 569.25 Gb (81%) - free 130.62 Gb (19%)
[1473816504] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;HARD;3;e:\ - total: 699.87 Gb - used: 569.25 Gb (81%) - free 130.62 Gb (19%)
[1473816763] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1473817281] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;Open right door.
[1473817363] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1473817704] SERVICE ALERT: AD-RMAD;E:\ Drive Space;OK;HARD;3;e:\ - total: 699.87 Gb - used: 547.60 Gb (78%) - free 152.28 Gb (22%)
[1473817704] SERVICE FLAPPING ALERT: AD-RMAD;E:\ Drive Space;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1473817881] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473817941] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473817963] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1473818002] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473818083] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1473818203] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1473819201] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;Open right door.
[1473819801] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473819861] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473819863] SERVICE FLAPPING ALERT: hallesecondsouth;Device Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1473819921] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473820003] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1473820603] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1473821121] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;Open right door.
[1473821721] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473821782] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473821841] SERVICE ALERT: hallegroundcolor;Device Status;OK;SOFT;3;Open right door.
[1473822441] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473822501] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473822561] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473823161] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;Open right door.
[1473823603] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1473823761] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473823821] SERVICE ALERT: hallegroundcolor;Device Status;OK;SOFT;2;Open right door.
[1473824203] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1473824421] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473824482] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473824541] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473824803] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1473825369] SERVICE ALERT: eateriespk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.59.141. Is the printer off?  Is the network down?
[1473825399] HOST ALERT: eateriespk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1473825403] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1473825429] SERVICE ALERT: eateriespk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.59.141. Is the printer off?  Is the network down?
[1473825459] HOST ALERT: eateriespk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1473825489] SERVICE ALERT: eateriespk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.59.141. Is the printer off?  Is the network down?
[1473825519] HOST ALERT: eateriespk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1473825523] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 81% (15 min average)