[1460865886] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460866739] SERVICE ALERT: SNAP;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1460866739] SERVICE ALERT: SNAP;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1460866742] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.40)
[1460866806] HOST ALERT: SNAP;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.46 ms
[1460866847] SERVICE ALERT: SNAP;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1460866848] SERVICE ALERT: SNAP;C:\ Drive Space;OK;SOFT;2;c:\ - total: 1116.65 Gb - used: 35.74 Gb (3%) - free 1080.91 Gb (97%)
[1460866905] SERVICE ALERT: ADWU1;Windows Update Services Event 364;CRITICAL;SOFT;1;Windows Server Update Services, Content file download failed. Reason: File cert verification failure.  Source File: /d/msdownload/update/software/crup/2016/04/windows10.0-kb3154879-x86_c7b4d1ac4829121785ae79b2df7550004eca089f.cab  Destination File: C:\WSUSupdates\WsusContent\9F\C7B4D1AC4829121785AE79B2DF7550004ECA089F.cab
[1460867025] SERVICE ALERT: ADWU1;Windows Update Services Event 364;CRITICAL;SOFT;2;Windows Server Update Services, Content file download failed. Reason: File cert verification failure.  Source File: /d/msdownload/update/software/crup/2016/04/windows10.0-kb3154879-x86_c7b4d1ac4829121785ae79b2df7550004eca089f.cab  Destination File: C:\WSUSupdates\WsusContent\9F\C7B4D1AC4829121785AE79B2DF7550004ECA089F.cab
[1460867145] SERVICE ALERT: ADWU1;Windows Update Services Event 364;CRITICAL;HARD;3;Windows Server Update Services, Content file download failed. Reason: File cert verification failure.  Source File: /d/msdownload/update/software/crup/2016/04/windows10.0-kb3154879-x86_c7b4d1ac4829121785ae79b2df7550004eca089f.cab  Destination File: C:\WSUSupdates\WsusContent\9F\C7B4D1AC4829121785AE79B2DF7550004ECA089F.cab
[1460869334] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 364;CRITICAL;SOFT;1;Windows Server Update Services(error, 364, success)[Content file download failed.
[1460869455] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 364;CRITICAL;SOFT;2;Windows Server Update Services(error, 364, success)[Content file download failed.
[1460869574] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 364;CRITICAL;HARD;3;Windows Server Update Services(error, 364, success)[Content file download failed.
[1460872885] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460873005] SERVICE ALERT: BOEAPPTEST;CPU Load;OK;SOFT;2;CPU Load 26% (5 min average)
[1460873680] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1460873800] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 84% (5 min average)
[1460874345] SERVICE ALERT: ADWU1;Windows Update Services Event 364;OK;HARD;3;OK: No entries found
[1460874400] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1460874520] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1460874640] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 95% (5 min average)
[1460875240] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1460875840] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1460876440] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 87% (5 min average)
[1460876440] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1460876617] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;CRITICAL;SOFT;1;Windows Server Update Services, The server is failing to download some updates.
[1460876686] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 93% (5 min average)
[1460876737] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;CRITICAL;SOFT;2;Windows Server Update Services, The server is failing to download some updates.
[1460876774] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 364;OK;HARD;3;Eventlog check ok
[1460876857] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;CRITICAL;HARD;3;Windows Server Update Services, The server is failing to download some updates.
[1460877286] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1460877640] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 91% (5 min average)
[1460877760] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 91% (5 min average)
[1460877880] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 95% (5 min average)
[1460878186] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;CRITICAL;SOFT;1;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1460878306] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;CRITICAL;SOFT;2;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1460878426] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;CRITICAL;HARD;3;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1460879080] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 79% (5 min average)
[1460880280] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1460880400] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 92% (5 min average)
[1460880520] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 97% (5 min average)
[1460880886] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 86% (5 min average)
[1460881512] SERVICE ALERT: NETIDADPROXY;Service: Windows Time;OK;HARD;3;W32Time: Started
[1460881720] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 68% (5 min average)
[1460882086] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460882206] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460882326] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 86% (5 min average)
[1460882926] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1460883046] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460883166] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1460884057] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;OK;HARD;3;OK: No entries found
[1460884366] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 63% (5 min average)
[1460884720] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1460884839] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 72% (5 min average)
[1460884966] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1460885026] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;OK;HARD;3;Eventlog check ok
[1460885086] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 84% (5 min average)
[1460885439] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1460885558] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 95% (5 min average)
[1460885678] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 75% (5 min average)
[1460886278] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1460886398] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 73% (5 min average)
[1460888086] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460888206] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1460888326] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 67% (5 min average)
[1460888926] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460889046] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1460889068] SERVICE ALERT: INTLAPP;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1460889166] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460889187] SERVICE ALERT: INTLAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1460889307] SERVICE ALERT: INTLAPP;CPU Load;OK;SOFT;3;CPU Load 0% (5 min average)
[1460890016] SERVICE ALERT: ADWU1;CPU Load;WARNING;SOFT;1;CPU Load 94% (5 min average)
[1460890136] SERVICE ALERT: ADWU1;CPU Load;WARNING;SOFT;2;CPU Load 93% (5 min average)
[1460890256] SERVICE ALERT: ADWU1;CPU Load;CRITICAL;HARD;3;CPU Load 95% (5 min average)
[1460890856] SERVICE ALERT: ADWU1;CPU Load;OK;HARD;3;CPU Load 0% (5 min average)
[1460892166] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 94% (5 min average)
[1460892766] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460892998] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1460893366] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 94% (5 min average)
[1460893966] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 76% (5 min average)
[1460894566] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1460894685] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 94% (5 min average)
[1460894805] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460895405] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 82% (5 min average)
[1460896605] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460896725] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 82% (5 min average)
[1460897251] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1460897312] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1460897325] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1460897371] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1460897445] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 91% (5 min average)
[1460897565] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1460898457] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;CRITICAL;SOFT;1;Windows Server Update Services, The server is failing to download some updates.
[1460898577] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;CRITICAL;SOFT;2;Windows Server Update Services, The server is failing to download some updates.
[1460898697] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;CRITICAL;HARD;3;Windows Server Update Services, The server is failing to download some updates.
[1460899365] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 86% (5 min average)
[1460899425] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;CRITICAL;SOFT;1;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1460899544] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;CRITICAL;SOFT;2;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1460899664] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;CRITICAL;HARD;3;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1460899965] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460900085] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 90% (5 min average)
[1460900205] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 70% (5 min average)
[1460901405] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1460901525] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460901645] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460905295] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;OK;HARD;3;OK: No entries found
[1460905845] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 64% (5 min average)
[1460906445] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1460906565] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460906644] SERVICE ALERT: studentcenterpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460906685] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 97% (5 min average)
[1460906704] SERVICE ALERT: studentcenterpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460906764] SERVICE ALERT: studentcenterpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460906865] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;OK;HARD;3;Eventlog check ok
[1460908485] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 92% (5 min average)
[1460909085] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1460909165] SERVICE ALERT: studentcenterpk;Device Status;OK;HARD;3;TRAY 4 EMPTY PLAIN LETTER   Ready
[1460909371] SERVICE ALERT: crossroadspk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1460910285] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 88% (5 min average)
[1460910885] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460911005] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 96% (5 min average)
[1460911125] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460911725] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 56% (5 min average)
[1460912325] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460912445] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1460912565] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460916165] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 67% (5 min average)
[1460916444] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460916504] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460916564] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460916765] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460916885] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 50% (5 min average)
[1460917485] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1460917544] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1460917604] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1460917605] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460917664] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1460917725] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1460917770] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1460918044] SERVICE ALERT: halleground3;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460918104] SERVICE ALERT: halleground3;Device Status;OK;SOFT;2;CLOSE RIGHT DOOR  46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1460918264] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 39% - OK!
[1460918364] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460919525] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 85% (5 min average)
[1460919684] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1460919695] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;CRITICAL;SOFT;1;Windows Server Update Services, The server is failing to download some updates.
[1460919744] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1460919804] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1460919815] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;CRITICAL;SOFT;2;Windows Server Update Services, The server is failing to download some updates.
[1460919935] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;CRITICAL;HARD;3;Windows Server Update Services, The server is failing to download some updates.
[1460920125] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1460920245] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460920365] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 66% (5 min average)
[1460920965] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1460921085] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 97% (5 min average)
[1460921205] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460921264] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;CRITICAL;SOFT;1;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1460921384] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;CRITICAL;SOFT;2;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1460921505] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;CRITICAL;HARD;3;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1460924264] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1460924324] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1460924364] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;RESUME  CANCEL CURRENT  JOB
[1460924364] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1460924384] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;SOFT;3;Black Print Cartridge HP C8543X is at 39% - OK!
[1460924805] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1460924964] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460925024] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460925085] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460925205] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;CRITICAL;HARD;3;Black Print Cartridge HP C8543X is at 1% - CRITICAL!
[1460925405] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1460927135] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;OK;HARD;3;OK: No entries found
[1460927190] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1460927200] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1460927200] SERVICE ALERT: hallefirstcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1460927220] HOST ALERT: hallefirstcolor;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460927234] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1460927242] HOST ALERT: hallefirstcolor;UP;SOFT;2;PING OK - Packet loss = 44%, RTA = 0.40 ms
[1460927245] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;OK;SOFT;2;Black Cartridge HP CE260X is at 12% - OK!
[1460927255] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;OK;SOFT;2;Cyan Cartridge HP CE261A is at 31% - OK!
[1460927255] SERVICE ALERT: hallefirstcolor;Image Fuser Kit HP 110V-CE246A;OK;SOFT;2;Image Fuser Kit HP 110V-CE246A, 220V-CE247A is at 100% - OK!
[1460927288] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;OK;SOFT;2;Image Transfer Kit HP CE249A is at 8% - OK!
[1460928104] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;OK;HARD;3;Eventlog check ok
[1460928684] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;TRAY 1 SIZE= *ANY SIZE Setting saved
[1460929026] HOST ALERT: hallefirst2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460929116] HOST ALERT: hallefirst2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460929190] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1460929206] HOST ALERT: hallefirst2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460929250] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1460929280] HOST ALERT: hallefirst2;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460929290] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1460929310] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1460929320] HOST ALERT: hallefirst2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460929336] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1460929350] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1460929366] HOST ALERT: hallefirst2;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1460929396] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1460929405] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1460929410] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1460929426] HOST ALERT: hallefirst2;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1460929456] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1460929486] HOST ALERT: hallefirst2;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1460929576] HOST ALERT: hallefirst2;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1460929666] HOST ALERT: hallefirst2;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1460931814] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460931874] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460931934] SERVICE ALERT: hallefirst1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460933010] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1460933040] HOST ALERT: hallefirst1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460933065] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 86% - OK!
[1460933069] HOST ALERT: hallefirst1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.52 ms
[1460933134] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;Ready
[1460936804] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1460937334] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460937394] SERVICE ALERT: hallefirst1;Device Status;OK;SOFT;2;Clearing paper path
[1460939064] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1460939124] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1460939184] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;SOFT;3;Black Print Cartridge HP C8543X is at 69% - OK!
[1460940810] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1460941534] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;CRITICAL;SOFT;1;Windows Server Update Services, The server is failing to download some updates.
[1460941594] SERVICE ALERT: TARGETX-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1460941626] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;SOFT;1;e:\ - total: 499.87 Gb - used: 400.83 Gb (80%) - free 99.04 Gb (20%)
[1460941654] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;CRITICAL;SOFT;2;Windows Server Update Services, The server is failing to download some updates.
[1460941714] SERVICE ALERT: TARGETX-2012;CPU Load;OK;SOFT;2;CPU Load 71% (5 min average)
[1460941746] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;SOFT;2;e:\ - total: 499.87 Gb - used: 400.83 Gb (80%) - free 99.04 Gb (20%)
[1460941774] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;CRITICAL;HARD;3;Windows Server Update Services, The server is failing to download some updates.
[1460941866] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;HARD;3;e:\ - total: 499.87 Gb - used: 400.83 Gb (80%) - free 99.04 Gb (20%)
[1460942504] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;CRITICAL;SOFT;1;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1460942624] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;CRITICAL;SOFT;2;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1460942745] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;CRITICAL;HARD;3;Windows Server Update Services(error, 10032, success)[The server is failing to download some updates., ], eventlog: 1 > critical
[1460943066] SERVICE ALERT: AD-RMAD;E:\ Drive Space;OK;HARD;3;e:\ - total: 499.87 Gb - used: 394.97 Gb (79%) - free 104.90 Gb (21%)
[1460943890] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460943949] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460944009] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460944609] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;Clear jam then press OK.
[1460945809] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460945869] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460945929] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460946529] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;Clear jam then press OK.
[1460946529] SERVICE FLAPPING ALERT: hallegroundcolor;Device Status;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1460947129] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460947189] SERVICE ALERT: hallegroundcolor;Device Status;OK;SOFT;2;Clear jam then press OK.
[1460947664] SERVICE ALERT: VC;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460947764] SERVICE ALERT: VC;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1460947789] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460947850] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460947909] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460948203] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 81% (5 min average)
[1460948509] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;SUPPLIES VERY LOW  46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1460948803] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1460948923] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 54% (5 min average)
[1460948974] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;OK;HARD;3;OK: No entries found
[1460949523] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460949643] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 82% (5 min average)
[1460949944] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;OK;HARD;3;Eventlog check ok
[1460949996] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1460950116] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 94% (5 min average)
[1460950236] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 86% (5 min average)
[1460950243] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1460950363] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 83% (5 min average)
[1460950963] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1460951083] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 94% (5 min average)
[1460951145] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 1 Trays Empty
[1460951203] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1460951205] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 1 Trays Empty
[1460951266] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 1 Trays Empty