[1461645002] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1461645602] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1461645722] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1461645842] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461648432] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461648493] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461648552] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461648842] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1461649442] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461649562] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1461649682] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461650282] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1461650882] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461651002] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1461651122] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461651252] SERVICE FLAPPING ALERT: hallefirst1;Device Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1461651722] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1461651722] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461651845] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1461651962] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461652397] SERVICE ALERT: INTLTESTAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461652517] SERVICE ALERT: INTLTESTAPP;CPU Load;OK;SOFT;2;CPU Load 0% (15 min average)
[1461652562] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 88% (15 min average)
[1461653162] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1461653282] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1461654482] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1461654602] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1461655202] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461655322] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1461658322] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 90% (15 min average)
[1461659522] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461659642] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1461659762] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461660962] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1461661562] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1461661682] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461661802] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461664202] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 84% (15 min average)
[1461664922] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461665042] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 77% (15 min average)
[1461665402] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1461665522] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461665642] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1461666417] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1461666476] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1461666500] SERVICE ALERT: eateriespk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461666537] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1461666559] SERVICE ALERT: eateriespk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461666619] SERVICE ALERT: eateriespk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461666842] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 88% (15 min average)
[1461667442] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1461667562] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1461667682] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 84% (15 min average)
[1461668282] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1461668353] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;REPLACE BLACK CARTRIDGE 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1461668362] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 21% - OK!
[1461668366] HOST ALERT: halleground1;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.54 ms
[1461668402] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461668522] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461668652] SERVICE ALERT: halleground1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1461668659] SERVICE ALERT: halleground1;Device Status;OK;HARD;3;TRAY 4 EMPTY PLAIN LETTER   Ready
[1461670502] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;WARNING;HARD;3;Yellow Cartridge HP CE262A is at WARNING level!
[1461670702] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1461670878] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1461670908] HOST ALERT: hallesecondsouth;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1461670998] HOST ALERT: hallesecondsouth;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1461671088] HOST ALERT: hallesecondsouth;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1461671103] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;OK;HARD;3;Yellow Cartridge HP CE262A is at 100% - OK!
[1461671178] HOST ALERT: hallesecondsouth;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1461671197] HOST ALERT: hallesecondsouth;UP;SOFT;5;PING OK - Packet loss = 0%, RTA = 0.45 ms
[1461671508] HOST ALERT: hallesecondsouth;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1461671598] HOST ALERT: hallesecondsouth;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1461671688] HOST ALERT: hallesecondsouth;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1461671778] HOST ALERT: hallesecondsouth;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1461671798] SERVICE ALERT: hallesecondsouth;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1461671828] HOST ALERT: hallesecondsouth;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1461671848] SERVICE ALERT: hallesecondsouth;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1461671858] SERVICE ALERT: hallesecondsouth;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1461671878] HOST ALERT: hallesecondsouth;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1461671908] SERVICE ALERT: hallesecondsouth;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1461671918] SERVICE ALERT: hallesecondsouth;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1461671938] HOST ALERT: hallesecondsouth;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1461671968] SERVICE ALERT: hallesecondsouth;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1461671998] HOST ALERT: hallesecondsouth;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1461672088] HOST ALERT: hallesecondsouth;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1461672177] HOST ALERT: hallesecondsouth;UP;SOFT;10;PING WARNING - Packet loss = 83%, RTA = 0.54 ms
[1461672512] SERVICE ALERT: hallesecondsouth;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1461672562] SERVICE ALERT: hallesecondsouth;Device Status;OK;HARD;3;Ready
[1461672672] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1461672722] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461673322] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461673802] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (39.2% change >= 20.0% threshold)
[1461673802] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1461673802] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1461673802] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1461673802] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1461674302] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1461674352] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461674362] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 100% - OK!
[1461674412] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461674472] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1461674478] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1461674508] HOST ALERT: hallefirst2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1461674522] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 67% (15 min average)
[1461674572] HOST ALERT: hallefirst2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.44 ms
[1461675073] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Ready
[1461675122] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1461675242] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1461675362] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461675583] SERVICE ALERT: hallefirst1;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1461675643] SERVICE ALERT: hallefirst1;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1461675702] SERVICE ALERT: hallefirst1;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1461675858] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1461675913] SERVICE ALERT: hallefirst1;Device Status;OK;SOFT;2;Clearing paper path
[1461676303] SERVICE ALERT: hallefirst1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1461676442] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;CRITICAL;HARD;3;Black Print Cartridge HP C8543X is at 1% - CRITICAL!
[1461676512] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461676562] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1461676572] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461676632] SERVICE ALERT: hallefirst1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461676872] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461676932] SERVICE ALERT: hallefirst2;Device Status;OK;SOFT;2;Clearing paper path
[1461677162] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461677232] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;Clearing paper path
[1461677838] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1461677892] SERVICE ALERT: hallefirst1;Device Status;OK;SOFT;2;Processing job from tray 4
[1461678362] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461678492] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461678552] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461678612] SERVICE ALERT: hallefirst1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461678793] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1461678853] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1461678912] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1461680057] SERVICE ALERT: halleground1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461680118] SERVICE ALERT: halleground1;Device Status;OK;SOFT;2;Clearing paper path
[1461680412] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;Ready
[1461680412] SERVICE FLAPPING ALERT: hallefirst1;Device Status;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1461680642] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1461680712] SERVICE ALERT: studentcenterpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1461681462] HOST ALERT: hallefirst2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1461681526] HOST ALERT: hallefirst2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.44 ms
[1461682447] HOST ALERT: sillpk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.44 ms
[1461682535] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;REPLACE MAINTENANCE KIT 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1461682565] SERVICE ALERT: sillpk;Black Cartridge HP CC364X;OK;HARD;3;Black Cartridge HP CC364X is at 12% - OK!
[1461682708] SERVICE ALERT: sillpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1461682951] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (33.0% change >= 20.0% threshold)
[1461682951] SERVICE FLAPPING ALERT: hallefirst1;Device Status;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1461682951] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1461682951] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1461682951] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1461682951] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1461683423] SERVICE ALERT: ADCA1;Service: WMI;CRITICAL;SOFT;1;(Return code of 13 is out of bounds)
[1461683423] HOST ALERT: ADCA1;DOWN;SOFT;1;(Return code of 13 is out of bounds)
[1461683423] SERVICE ALERT: idmrbs.emich.edu;eDirectory Service;CRITICAL;SOFT;1;(Return code of 13 is out of bounds)
[1461683423] HOST ALERT: idmrbs.emich.edu;DOWN;SOFT;1;(Return code of 13 is out of bounds)
[1461683423] SERVICE ALERT: WINKMS;Service: RPC;CRITICAL;SOFT;1;(Return code of 13 is out of bounds)
[1461683423] HOST ALERT: WINKMS;DOWN;SOFT;1;(Return code of 13 is out of bounds)
[1461683424] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (32.0% change >= 20.0% threshold)
[1461683424] SERVICE FLAPPING ALERT: hallefirst1;Device Status;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1461683424] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1461683424] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1461683424] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1461683424] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1461683466] HOST ALERT: ADCA1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1461683466] HOST ALERT: WINKMS;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.36 ms
[1461683487] HOST ALERT: idmrbs.emich.edu;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1461683532] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461683543] SERVICE ALERT: ADCA1;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1461683543] SERVICE ALERT: WINKMS;Service: RPC;OK;SOFT;2;RpcSs: Started
[1461683543] SERVICE ALERT: idmrbs.emich.edu;eDirectory Service;OK;SOFT;2;OK: ndsd status is good.
[1461683592] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461683652] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461683757] SERVICE ALERT: fanghorn;Service: SSH;CRITICAL;SOFT;1;Server answer:
[1461683877] SERVICE ALERT: fanghorn;Service: SSH;CRITICAL;SOFT;2;Server answer:
[1461683920] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (31.1% change >= 20.0% threshold)
[1461683920] SERVICE FLAPPING ALERT: hallefirst1;Device Status;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1461683920] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1461683920] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1461683920] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1461683920] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1461683935] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (31.1% change >= 20.0% threshold)
[1461683935] SERVICE FLAPPING ALERT: hallefirst1;Device Status;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1461683935] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1461683935] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1461683935] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1461683935] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1461684252] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Clearing paper path
[1461684252] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1461684852] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461684912] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461684962] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 79% (15 min average)
[1461684972] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461685562] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461685682] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 77% (15 min average)
[1461686066] SERVICE ALERT: eateriespk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.59.141. Is the printer off?  Is the network down?
[1461686096] HOST ALERT: eateriespk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1461686126] SERVICE ALERT: eateriespk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.59.141. Is the printer off?  Is the network down?
[1461686156] HOST ALERT: eateriespk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1461686186] SERVICE ALERT: eateriespk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.59.141. Is the printer off?  Is the network down?
[1461686216] HOST ALERT: eateriespk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1461686282] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461686306] HOST ALERT: eateriespk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1461686342] SERVICE ALERT: eateriespk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.59.141. Is the printer off?  Is the network down?
[1461686372] HOST ALERT: eateriespk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1461686402] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461686425] SERVICE ALERT: eateriespk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.59.141. Is the printer off?  Is the network down?
[1461686455] HOST ALERT: eateriespk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1461686522] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461686545] HOST ALERT: eateriespk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1461686615] HOST ALERT: eateriespk;UP;SOFT;8;PING OK - Packet loss = 54%, RTA = 0.51 ms
[1461686780] SERVICE ALERT: eateriespk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 36% - OK!
[1461686936] SERVICE ALERT: eateriespk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1461687013] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461687019] SERVICE ALERT: eateriespk;Device Status;OK;HARD;3;TRAY 4 EMPTY PLAIN 11X17   Ready
[1461687072] SERVICE ALERT: hallefirst1;Device Status;OK;SOFT;2;Clearing paper path
[1461687372] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Ready
[1461688272] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461688332] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461688392] SERVICE ALERT: hallefirst1;Device Status;OK;SOFT;3;stack are straight. 54 6F 20 65 78 69 74 20 70 72 65 73 73 20 1F 2E
[1461689172] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461689232] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461689292] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461689535] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 1 Trays Empty
[1461689596] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! All Trays Still Full
[1461689642] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461689762] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 61% (15 min average)
[1461690672] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1461690733] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1461690793] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1461690793] SERVICE FLAPPING ALERT: hallefirst1;Device Status;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1461691092] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Ready
[1461691862] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1461691922] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1461691981] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1461693182] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1461693721] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 80% (15 min average)
[1461693792] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461693852] SERVICE ALERT: hallefirst1;Device Status;OK;SOFT;2;Clearing paper path
[1461693871] SERVICE ALERT: INTLAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1461693991] SERVICE ALERT: INTLAPP;CPU Load;OK;SOFT;2;CPU Load 6% (15 min average)
[1461694321] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461694441] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461694560] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461697861] SERVICE ALERT: TABLEAU;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461697868] SERVICE ALERT: TABLEAU;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461697878] SERVICE ALERT: TABLEAU;Service: User Profile;CRITICAL;SOFT;1;connect to address 164.76.99.232 and port 12489: Connection refused
[1461697891] HOST ALERT: TABLEAU;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1461697952] SERVICE ALERT: TABLEAU;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461697953] SERVICE ALERT: TABLEAU;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461697953] SERVICE ALERT: TABLEAU;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461697981] SERVICE ALERT: TABLEAU;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1461697981] HOST ALERT: TABLEAU;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1461697988] SERVICE ALERT: TABLEAU;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1461697992] HOST ALERT: TABLEAU;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.36 ms
[1461697998] SERVICE ALERT: TABLEAU;Service: User Profile;UNKNOWN;SOFT;2;Failed to open service manager: 1115: A system shutdown is in progress.
[1461698030] HOST DOWNTIME ALERT: TABLEAU;STARTED; Host has entered a period of scheduled downtime
[1461698054] SERVICE ALERT: TABLEAU;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461698054] SERVICE ALERT: TABLEAU;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1461698054] SERVICE ALERT: TABLEAU;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1461698081] SERVICE ALERT: TABLEAU;Service: RPC;OK;SOFT;3;RpcSs: Started
[1461698084] HOST ALERT: TABLEAU;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1461698089] SERVICE ALERT: TABLEAU;Service: NetLogin;OK;SOFT;3;Netlogon: Started
[1461698101] HOST ALERT: TABLEAU;UP;SOFT;2;PING OK - Packet loss = 61%, RTA = 0.38 ms
[1461698101] HOST FLAPPING ALERT: TABLEAU;STARTED; Host appears to have started flapping (20.9% change > 20.0% threshold)
[1461698115] SERVICE ALERT: INTLTESTAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1461698119] SERVICE ALERT: TABLEAU;Service: User Profile;OK;SOFT;3;ProfSvc: Started
[1461698155] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461698173] SERVICE ALERT: TABLEAU;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461698235] SERVICE ALERT: INTLTESTAPP;CPU Load;OK;SOFT;2;CPU Load 0% (15 min average)
[1461698652] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461698712] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461698755] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461698761] HOST FLAPPING ALERT: TABLEAU;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1461698772] SERVICE ALERT: hallefirst1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461698773] SERVICE ALERT: TABLEAU;CPU Load;OK;HARD;3;CPU Load 0% (15 min average)
[1461699372] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;Ready
[1461699741] HOST DOWNTIME ALERT: TABLEAU;CANCELLED; Scheduled downtime for host has been cancelled.
[1461699973] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461700032] SERVICE ALERT: hallefirst1;Device Status;OK;SOFT;2;INPUT PAPER PATH OPEN 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1461700581] SERVICE ALERT: cava2;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461700581] SERVICE ALERT: cava2;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461700611] SERVICE ALERT: cava2;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461700611] HOST ALERT: cava2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1461700624] SERVICE ALERT: cava2;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1461700644] HOST ALERT: cava2;UP;SOFT;2;PING OK - Packet loss = 76%, RTA = 0.38 ms
[1461700681] SERVICE ALERT: cava2;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1461700682] SERVICE ALERT: cava2;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1461700711] SERVICE ALERT: cava2;Service: RPC;OK;SOFT;2;RpcSs: Started
[1461700734] SERVICE ALERT: cava2;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1461700751] SERVICE ALERT: cava2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1461700871] SERVICE ALERT: cava2;CPU Load;OK;SOFT;2;CPU Load 0% (15 min average)
[1461701290] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1461701755] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 90% (15 min average)
[1461701832] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461701881] SERVICE ALERT: cava2;Service: Group Policy;CRITICAL;SOFT;1;gpsvc: Stopped
[1461701892] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461701931] SERVICE ALERT: cava2;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461701931] SERVICE ALERT: cava2;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461701933] SERVICE ALERT: cava2;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461701952] SERVICE ALERT: hallefirst1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461701954] SERVICE ALERT: cava2;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461701961] HOST ALERT: cava2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1461702021] SERVICE ALERT: cava2;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1461702031] SERVICE ALERT: cava2;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461702051] HOST ALERT: cava2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1461702051] SERVICE ALERT: cava2;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1461702051] SERVICE ALERT: cava2;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1461702053] SERVICE ALERT: cava2;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1461702074] SERVICE ALERT: cava2;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1461702081] HOST ALERT: cava2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1461702081] SERVICE ALERT: cava2;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461702091] SERVICE ALERT: cava2;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461702111] HOST ALERT: cava2;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1461702114] SERVICE ALERT: cava2;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461702141] SERVICE ALERT: cava2;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1461702144] HOST ALERT: cava2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1461702151] SERVICE ALERT: cava2;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1461702171] SERVICE ALERT: cava2;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1461702171] SERVICE ALERT: cava2;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1461702173] SERVICE ALERT: cava2;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1461702181] HOST ALERT: cava2;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1461702194] SERVICE ALERT: cava2;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1461702201] SERVICE ALERT: cava2;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1461702211] SERVICE ALERT: cava2;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1461702224] HOST ALERT: cava2;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1461702234] SERVICE ALERT: cava2;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1461702261] SERVICE ALERT: cava2;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461702264] HOST ALERT: cava2;DOWN;SOFT;8;(Host check timed out after 30.01 seconds)
[1461702271] SERVICE ALERT: cava2;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1461702301] HOST ALERT: cava2;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1461702321] SERVICE ALERT: cava2;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1461702331] SERVICE ALERT: cava2;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1461702351] HOST ALERT: cava2;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1461702354] SERVICE ALERT: cava2;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1461702355] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461702381] SERVICE ALERT: cava2;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1461702448] SERVICE ALERT: cava2;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461702455] SERVICE ALERT: cava2;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461702475] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461702501] SERVICE ALERT: cava2;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461702501] SERVICE ALERT: cava2;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461702501] SERVICE ALERT: cava2;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1461702568] SERVICE ALERT: cava2;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1461702575] SERVICE ALERT: cava2;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1461702595] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461702621] SERVICE ALERT: cava2;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1461702621] SERVICE ALERT: cava2;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1461702688] SERVICE ALERT: cava2;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1461702695] SERVICE ALERT: cava2;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1461702741] SERVICE ALERT: cava2;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1461702741] SERVICE ALERT: cava2;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1461703751] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;Processing job from tray 3
[1461703751] SERVICE FLAPPING ALERT: hallefirst1;Device Status;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1461703795] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 88% (15 min average)
[1461704290] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461704350] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461704395] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461704410] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461704515] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461704561] SERVICE ALERT: owen101lab3;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461704621] SERVICE ALERT: owen101lab3;Device Status;OK;SOFT;2;Genuine HP cartridge installed
[1461704635] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461705380] SERVICE ALERT: emu-fndn-fe7;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461705381] SERVICE ALERT: emu-fndn-fe7;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461705410] HOST ALERT: emu-fndn-fe7;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1461705436] SERVICE ALERT: emu-fndn-fe7;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461705466] HOST ALERT: emu-fndn-fe7;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1461705471] SERVICE ALERT: emu-fndn-fe7;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461705500] SERVICE ALERT: emu-fndn-fe7;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1461705501] SERVICE ALERT: emu-fndn-fe7;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1461705501] HOST ALERT: emu-fndn-fe7;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1461705528] SERVICE ALERT: emu-fndn-fe7;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461705556] SERVICE ALERT: emu-fndn-fe7;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1461705558] HOST ALERT: emu-fndn-fe7;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1461705561] SERVICE ALERT: emu-fndn-fe7;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461705591] SERVICE ALERT: emu-fndn-fe7;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1461705591] HOST ALERT: emu-fndn-fe7;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1461705603] SERVICE ALERT: emu-fndn-fe7;Uptime;CRITICAL;HARD;3;connect to address 164.76.2.109 and port 12489: Connection refused
[1461705604] SERVICE ALERT: emu-fndn-fe7;NSClient++ Version;CRITICAL;HARD;3;connect to address 164.76.2.109 and port 12489: Connection refused
[1461705607] HOST ALERT: emu-fndn-fe7;UP;SOFT;6;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1461705610] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Open right  door.
[1461705613] SERVICE ALERT: emu-fndn-fe7;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.2.109 and port 12489: Connection refused
[1461705628] SERVICE ALERT: emu-fndn-fe7;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1461705656] SERVICE ALERT: emu-fndn-fe7;Service: DNS Client;OK;SOFT;3;Dnscache: Started
[1461705661] SERVICE ALERT: emu-fndn-fe7;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1461705691] SERVICE ALERT: emu-fndn-fe7;Service: Firewall;OK;SOFT;3;MpsSvc: Started
[1461705733] SERVICE ALERT: emu-fndn-fe7;Service: RPC;OK;SOFT;2;RpcSs: Started
[1461706200] SERVICE ALERT: emu-fndn-fe7;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 11 minute(s)
[1461706201] SERVICE ALERT: emu-fndn-fe7;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1461706573] SERVICE ALERT: TABLEAU;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1461706693] SERVICE ALERT: TABLEAU;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1461706813] SERVICE ALERT: TABLEAU;CPU Load;OK;SOFT;3;CPU Load 9% (15 min average)
[1461706909] HOST ALERT: cava2;UP;HARD;10;PING OK - Packet loss = 77%, RTA = 0.38 ms
[1461706921] SERVICE ALERT: cava2;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 1 minute(s)
[1461706921] SERVICE ALERT: cava2;Service: WMI;OK;HARD;3;Winmgmt: Started
[1461706921] SERVICE ALERT: cava2;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1461706951] SERVICE ALERT: cava2;VMWare Tools;OK;HARD;3;VMTools: Started
[1461706951] SERVICE ALERT: cava2;Service: RPC;OK;HARD;3;RpcSs: Started
[1461706953] SERVICE ALERT: cava2;C:\ Drive Space;OK;HARD;3;c:\ - total: 49.66 Gb - used: 26.28 Gb (53%) - free 23.38 Gb (47%)
[1461706974] SERVICE ALERT: cava2;Service: Windows Time;OK;HARD;3;W32Time: Started
[1461707035] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1461707051] SERVICE ALERT: cava2;Memory Usage;OK;HARD;3;Memory usage: total:43391.92 MB - used: 7279.72 MB (17%) - free: 36112.20 MB (83%)
[1461707101] SERVICE ALERT: cava2;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1461707111] SERVICE ALERT: cava2;CPU Load;OK;HARD;3;CPU Load 2% (15 min average)
[1461707134] SERVICE ALERT: cava2;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1461707281] SERVICE ALERT: cava2;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1461707468] SERVICE ALERT: cava2;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1461707475] SERVICE ALERT: cava2;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1461707553] SERVICE ALERT: emu-fndn-fe7;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461707583] HOST ALERT: emu-fndn-fe7;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1461707635] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461707645] SERVICE ALERT: emu-fndn-fe7;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461707673] SERVICE ALERT: emu-fndn-fe7;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1461707673] HOST ALERT: emu-fndn-fe7;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1461707675] SERVICE ALERT: emu-fndn-fe7;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461707690] SERVICE ALERT: emu-fndn-fe7;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461707705] HOST ALERT: emu-fndn-fe7;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1461707738] SERVICE ALERT: emu-fndn-fe7;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461707755] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461707765] SERVICE ALERT: emu-fndn-fe7;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1461707768] HOST ALERT: emu-fndn-fe7;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1461707772] SERVICE ALERT: emu-fndn-fe7;Service: RPC;CRITICAL;HARD;3;connect to address 164.76.2.109 and port 12489: Connection refused
[1461707775] SERVICE ALERT: emu-fndn-fe7;Service: Windows Time;CRITICAL;SOFT;2;connect to address 164.76.2.109 and port 12489: Connection refused
[1461707777] HOST ALERT: emu-fndn-fe7;UP;SOFT;5;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1461707790] SERVICE ALERT: emu-fndn-fe7;CPU Load;OK;SOFT;2;CPU Load 16% (15 min average)
[1461707838] SERVICE ALERT: emu-fndn-fe7;Memory Usage;OK;SOFT;2;Memory usage: total:18781.43 MB - used: 1760.49 MB (9%) - free: 17020.94 MB (91%)
[1461707865] SERVICE ALERT: emu-fndn-fe7;C:\ Drive Space;OK;SOFT;3;c:\ - total: 698.53 Gb - used: 43.10 Gb (6%) - free 655.43 Gb (94%)
[1461707875] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461707895] SERVICE ALERT: emu-fndn-fe7;Service: Windows Time;OK;SOFT;3;W32Time: Started
[1461708010] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461708070] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461708130] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461708372] SERVICE ALERT: emu-fndn-fe7;Service: RPC;OK;HARD;3;RpcSs: Started
[1461708670] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;WARNING;SOFT;1;Cyan Cartridge HP CE261A is at WARNING level!
[1461708731] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;WARNING;SOFT;2;Cyan Cartridge HP CE261A is at WARNING level!
[1461708790] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;WARNING;HARD;3;Cyan Cartridge HP CE261A is at WARNING level!
[1461709390] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;OK;HARD;3;Cyan Cartridge HP CE261A is at 100% - OK!
[1461710161] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1461710275] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461710281] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1461710401] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 94% (15 min average)
[1461710875] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461712330] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;MENUS  RETRIEVE JOB  INFORMATION  PAPER HANDLING
[1461712330] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1461712675] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461713275] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461713950] SERVICE FLAPPING ALERT: hallefirst1;Device Status;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1461714551] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461714610] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461714669] SERVICE ALERT: hallefirst1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461715269] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;Ready
[1461715869] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461715929] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461715989] SERVICE ALERT: hallefirst1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461717788] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;Clearing paper path
[1461717788] SERVICE FLAPPING ALERT: hallefirst1;Device Status;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1461719016] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;SOFT;1;e:\ - total: 499.87 Gb - used: 401.00 Gb (80%) - free 98.87 Gb (20%)
[1461719136] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;SOFT;2;e:\ - total: 499.87 Gb - used: 401.00 Gb (80%) - free 98.87 Gb (20%)
[1461719256] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;HARD;3;e:\ - total: 499.87 Gb - used: 401.00 Gb (80%) - free 98.87 Gb (20%)
[1461719275] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1461719876] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1461719995] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461720115] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461720130] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1461721056] SERVICE ALERT: AD-RMAD;E:\ Drive Space;OK;HARD;3;e:\ - total: 499.87 Gb - used: 390.35 Gb (78%) - free 109.53 Gb (22%)
[1461721315] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461721915] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461724315] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 90% (15 min average)
[1461724915] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461725035] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461725155] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461725389] SERVICE ALERT: VC;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461725489] SERVICE ALERT: VC;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1461725755] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461726355] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461726601] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1461726721] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1461726841] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1461727441] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1461727987] SERVICE FLAPPING ALERT: hallefirst1;Device Status;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1461728041] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1461728155] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461728161] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 92% (15 min average)
[1461728755] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461729355] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)