[1461211580] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 69% (15 min average)
[1461211687] SERVICE ALERT: ADDC4;AD Replication;OK;HARD;3;OK: replication is up and running
[1461212662] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461213000] HOST DOWNTIME ALERT: IMAGES;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: OPIE;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: LANSWEEPER;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: ADDC11;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: MBAMDB;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: ADDC66;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: NESSUS;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: ADWSUSCLIENTS;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: NETENG-PA;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: OFFICEKMS;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: ADDC2;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: REFWIN;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: PNCAPP-2012;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: ADDC55;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: ADWU1;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: PNCDB-2012;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: ADDC22;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: CBORD;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: PNCWEB-2012;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: AD-RMAD;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: PROCARE;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: R25-ALPHA2012;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: CEDARON;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: DIRECTACCESS;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: R25-ZETA2012;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: R25-DELTA2012;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: COE-WEBDAV;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: ADDC3;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: SP2010;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: ADARC;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: SANDBOX-DB;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: ADDC33;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: DPS-SECURE02;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: SP2010DB;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: UHSPHARM-2012;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: TABLEAU;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: TI22;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: ADDC5;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: EMU-VALUE01;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: TI22-COE;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: EMUSCRIPT;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: TI22-PSYCLINIC;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: ADDC4;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: EMUSQLEXPRESS;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: WINKMS;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: XDM;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: ADDC44;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: MBAM;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: Echo-01;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: HDBACKUPS;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: FTK;STOPPED; Host has exited from a period of scheduled downtime
[1461213000] HOST DOWNTIME ALERT: ADAUTH-RMAD;STOPPED; Host has exited from a period of scheduled downtime
[1461213262] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461213862] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461214462] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461215062] SERVICE FLAPPING ALERT: hallefirst1;Device Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1461217579] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1461217699] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1461218899] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1461219019] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1461223460] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 85% (15 min average)
[1461224060] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461224180] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461224300] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461224419] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1461224539] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461224659] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 80% (15 min average)
[1461225500] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461225859] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1461225979] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1461226100] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 86% (15 min average)
[1461226700] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461226820] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461226940] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1461229340] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 69% (15 min average)
[1461229940] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1461230060] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 88% (15 min average)
[1461232460] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461232580] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461232700] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 90% (15 min average)
[1461233300] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1461233420] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1461233540] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461234140] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 75% (15 min average)
[1461234740] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461234860] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1461234980] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 67% (15 min average)
[1461235580] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1461235700] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461235820] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 91% (15 min average)
[1461237020] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461237139] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461237259] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1461237885] SERVICE ALERT: ADWSUSSERVERS;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461237985] SERVICE ALERT: ADWSUSSERVERS;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1461238952] SERVICE ALERT: ADWSUSSERVERS;Service: RPC;WARNING;SOFT;1;RpcSs: Error
[1461239072] SERVICE ALERT: ADWSUSSERVERS;Service: RPC;OK;SOFT;2;RpcSs: Started
[1461239505] SERVICE ALERT: ADDC11;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461239535] HOST ALERT: ADDC11;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1461239554] SERVICE ALERT: ADDC11;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461239578] SERVICE ALERT: ADDC11;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461239584] HOST ALERT: ADDC11;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1461239599] SERVICE ALERT: ADDC11;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461239625] SERVICE ALERT: ADDC11;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1461239629] HOST ALERT: ADDC11;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1461239644] SERVICE ALERT: ADDC11;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461239674] SERVICE ALERT: ADDC11;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1461239674] HOST ALERT: ADDC11;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1461239692] SERVICE ALERT: ADDC11;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461239697] SERVICE ALERT: ADDC11;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1461239719] SERVICE ALERT: ADDC11;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1461239722] HOST ALERT: ADDC11;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1461239738] SERVICE ALERT: ADDC11;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461239745] SERVICE ALERT: ADDC11;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1461239764] SERVICE ALERT: ADDC11;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1461239768] HOST ALERT: ADDC11;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1461239786] SERVICE ALERT: ADDC11;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461239794] SERVICE ALERT: ADDC11;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1461239812] SERVICE ALERT: ADDC11;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1461239816] HOST ALERT: ADDC11;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1461239817] SERVICE ALERT: ADDC11;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1461239834] SERVICE ALERT: ADDC11;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461239839] SERVICE ALERT: ADDC11;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1461239847] HOST ALERT: ADDC11;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1461239858] SERVICE ALERT: ADDC11;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1461239878] SERVICE ALERT: ADDC11;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461239884] SERVICE ALERT: ADDC11;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1461239888] HOST ALERT: ADDC11;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1461239906] SERVICE ALERT: ADDC11;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1461239915] SERVICE ALERT: ADDC11;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461239928] SERVICE ALERT: ADDC11;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461239932] SERVICE ALERT: ADDC11;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1461239936] HOST ALERT: ADDC11;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1461239954] SERVICE ALERT: ADDC11;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1461239966] HOST ALERT: owen101lab3;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.49 ms
[1461239967] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (24.0% change >= 20.0% threshold)
[1461239967] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1461239967] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1461239967] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1461239967] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1461239998] SERVICE ALERT: ADDC11;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1461240022] SERVICE ALERT: idm1.emich.edu;NSDS Errors;UNKNOWN;SOFT;1;(No output returned from plugin)
[1461240026] SERVICE ALERT: ADDC11;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1461240032] SERVICE ALERT: idm2.emich.edu;NSDS Errors;UNKNOWN;SOFT;1;(No output returned from plugin)
[1461240033] SERVICE ALERT: ADDC11;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1461240035] SERVICE ALERT: ADDC11;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1461240048] SERVICE ALERT: ADDC11;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1461240052] SERVICE ALERT: owen101lab3;Device Status;OK;HARD;3;Warming up
[1461240058] SERVICE ALERT: ADDC11;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461240074] SERVICE ALERT: ADDC11;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1461240077] SERVICE ALERT: idm4.emich.edu;NSDS Errors;UNKNOWN;SOFT;1;(No output returned from plugin)
[1461240080] SERVICE ALERT: ADDC11;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461240082] SERVICE ALERT: idm3.emich.edu;NSDS Errors;UNKNOWN;SOFT;1;(No output returned from plugin)
[1461240118] SERVICE ALERT: ADDC11;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1461240142] SERVICE ALERT: idm1.emich.edu;NSDS Errors;UNKNOWN;SOFT;2;(No output returned from plugin)
[1461240144] SERVICE ALERT: idmrbstest.emich.edu;NSDS Errors;UNKNOWN;SOFT;1;(No output returned from plugin)
[1461240152] SERVICE ALERT: idm2.emich.edu;NSDS Errors;UNKNOWN;SOFT;2;(No output returned from plugin)
[1461240155] SERVICE ALERT: ADDC11;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1461240168] SERVICE ALERT: ADDC11;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1461240178] SERVICE ALERT: ADDC11;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1461240190] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (24.0% change >= 20.0% threshold)
[1461240190] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1461240190] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1461240190] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1461240190] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1461240197] SERVICE ALERT: idm4.emich.edu;NSDS Errors;UNKNOWN;SOFT;2;(No output returned from plugin)
[1461240202] SERVICE ALERT: idm3.emich.edu;NSDS Errors;UNKNOWN;SOFT;2;(No output returned from plugin)
[1461240256] SERVICE ALERT: ADDC11;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1461240262] SERVICE ALERT: idm1.emich.edu;NSDS Errors;UNKNOWN;HARD;3;(No output returned from plugin)
[1461240264] SERVICE ALERT: idmrbstest.emich.edu;NSDS Errors;UNKNOWN;SOFT;2;(No output returned from plugin)
[1461240272] SERVICE ALERT: idm2.emich.edu;NSDS Errors;UNKNOWN;HARD;3;(No output returned from plugin)
[1461240298] SERVICE ALERT: ADDC11;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1461240317] SERVICE ALERT: idm4.emich.edu;NSDS Errors;UNKNOWN;HARD;3;(No output returned from plugin)
[1461240322] SERVICE ALERT: idm3.emich.edu;NSDS Errors;UNKNOWN;HARD;3;(No output returned from plugin)
[1461240367] SERVICE ALERT: idmrbs.emich.edu;NSDS Errors;UNKNOWN;SOFT;1;(No output returned from plugin)
[1461240376] SERVICE ALERT: ADDC11;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1461240384] SERVICE ALERT: idmrbstest.emich.edu;NSDS Errors;UNKNOWN;HARD;3;(No output returned from plugin)
[1461240393] SERVICE ALERT: owen101lab3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1461240448] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1461240448] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1461240448] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1461240448] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1461240487] SERVICE ALERT: idmrbs.emich.edu;NSDS Errors;OK;SOFT;2;OK:
[1461240556] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 32% - OK!
[1461240862] SERVICE ALERT: idm1.emich.edu;NSDS Errors;OK;HARD;3;OK: 0 Errors Reported
[1461240870] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1461240870] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1461240870] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1461240870] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1461240872] SERVICE ALERT: idm2.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1461240892] SERVICE ALERT: ADWSUSSERVERS;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461240917] SERVICE ALERT: idm4.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1461240922] SERVICE ALERT: idm3.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1461240984] SERVICE ALERT: idmrbstest.emich.edu;NSDS Errors;OK;HARD;3;OK:
[1461240992] SERVICE ALERT: ADWSUSSERVERS;Service: RPC;OK;SOFT;2;RpcSs: Started
[1461240992] SERVICE ALERT: ADWU1;Service: Group Policy;CRITICAL;SOFT;1;gpsvc: Stopped
[1461240995] SERVICE ALERT: ADWU1;Windows Update Services Event 10022;UNKNOWN;SOFT;1;Failed to process command : EvtQuery failed: 1717: The interface is unknown.
[1461241112] SERVICE ALERT: ADWU1;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1461241115] SERVICE ALERT: ADWU1;Windows Update Services Event 10022;OK;SOFT;2;OK: No entries found
[1461241408] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1461241420] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10042;UNKNOWN;SOFT;1;Could not open the 'application' event log: 1717: The interface is unknown.
[1461241540] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10042;OK;SOFT;2;Eventlog check ok
[1461241732] SERVICE ALERT: ADWSUSSERVERS;Service: WMI;WARNING;SOFT;1;Winmgmt: Error
[1461241852] SERVICE ALERT: ADWSUSSERVERS;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1461241907] SERVICE ALERT: ADWU1;Windows Update Services Event 364;UNKNOWN;SOFT;1;Failed to process command : EvtQuery failed: 1717: The interface is unknown.
[1461241909] SERVICE ALERT: ADWSUSSERVERS;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461241909] SERVICE ALERT: ADWSUSSERVERS;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461242008] SERVICE ALERT: ADWSUSSERVERS;CPU Load;OK;SOFT;2;CPU Load 8% (15 min average)
[1461242008] SERVICE ALERT: ADWSUSSERVERS;VMWare Tools;OK;SOFT;2;VMTools: Started
[1461242027] SERVICE ALERT: ADWU1;Windows Update Services Event 364;OK;SOFT;2;OK: No entries found
[1461242129] SERVICE ALERT: ADWSUSSERVERS;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461242138] SERVICE ALERT: ADWSUSSERVERS;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461242170] SERVICE ALERT: ADWU1;Service: RPC;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1461242229] SERVICE ALERT: ADWSUSSERVERS;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1461242238] SERVICE ALERT: ADWSUSSERVERS;Memory Usage;OK;SOFT;2;Memory usage: total:32767.77 MB - used: 4068.70 MB (12%) - free: 28699.07 MB (88%)
[1461242290] SERVICE ALERT: ADWU1;Service: RPC;OK;SOFT;2;RpcSs: Started
[1461242638] HOST ALERT: ADDC11;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1461242656] SERVICE ALERT: ADDC11;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1461242656] SERVICE ALERT: ADDC11;CPU Load;OK;HARD;3;CPU Load 0% (15 min average)
[1461242656] SERVICE ALERT: ADDC11;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1461242656] SERVICE ALERT: ADDC11;C:\ Drive Space;OK;HARD;3;c:\ - total: 79.66 Gb - used: 33.28 Gb (42%) - free 46.38 Gb (58%)
[1461242656] SERVICE ALERT: ADDC11;Memory Usage;OK;HARD;3;Memory usage: total:20479.73 MB - used: 789.64 MB (4%) - free: 19690.08 MB (96%)
[1461242656] SERVICE ALERT: ADDC11;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1461242656] SERVICE ALERT: ADDC11;Service: WMI;OK;HARD;3;Winmgmt: Started
[1461242656] SERVICE ALERT: ADDC11;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1461242656] SERVICE ALERT: ADDC11;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1461242656] SERVICE ALERT: ADDC11;Service: Windows Time;OK;HARD;3;W32Time: Started
[1461242656] SERVICE ALERT: ADDC11;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1461242656] SERVICE ALERT: ADDC11;VMWare Tools;OK;HARD;3;VMTools: Started
[1461242656] SERVICE ALERT: ADDC11;Service: RPC;OK;HARD;3;RpcSs: Started
[1461242656] SERVICE ALERT: ADDC11;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1461243208] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1461243269] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1461243329] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1461243568] SERVICE ALERT: PROCARE;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1461243688] SERVICE ALERT: PROCARE;CPU Load;OK;SOFT;2;CPU Load 4% (15 min average)
[1461243858] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (15 min average)
[1461244239] SERVICE ALERT: WINKMS;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1461244355] SERVICE ALERT: WINKMS;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1461244458] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1461244578] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1461244698] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461244698] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (25.2% change >= 20.0% threshold)
[1461245019] SERVICE ALERT: studentcenterpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1461245298] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461245402] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (30.8% change >= 20.0% threshold)
[1461245402] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1461245402] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1461245402] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1461245402] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1461245898] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461246498] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 52% (15 min average)
[1461247098] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461247218] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461247283] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1461247338] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 93% (15 min average)
[1461247342] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1461247403] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1461247938] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1461248002] SERVICE ALERT: crossroadspk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1461248058] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461248178] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461249238] SERVICE ALERT: AD-RMAD;E:\ Drive Space;OK;HARD;3;e:\ - total: 499.87 Gb - used: 367.36 Gb (73%) - free 132.51 Gb (27%)
[1461249378] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461249978] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461250483] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1461250542] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 9% - OK!
[1461250578] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 77% (15 min average)
[1461251128] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1461251178] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461251298] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461251418] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461251472] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1461251532] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 27% - OK!
[1461251727] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1461251787] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1461251846] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1461251846] SERVICE FLAPPING ALERT: alexanderpk;Paper Tray Status;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1461253218] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461253647] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1461254418] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461255018] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461257418] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461257806] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors is too high! 1173 found since last run. securityErrors OK. wholeSubtreeSearchOps is too high! 1034 found since last run. oneLevelSearchOps OK. searchOps is too high! 1035 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1461257926] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 830 found since last run. oneLevelSearchOps OK. searchOps is too high! 831 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1461258046] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1461258646] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 202 found since last run. oneLevelSearchOps OK. searchOps is too high! 203 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1461258672] SERVICE ALERT: prayharroldpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461258732] SERVICE ALERT: prayharroldpk;Device Status;OK;SOFT;2;CLOSE RIGHT DOOR  46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1461258766] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1461259218] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461259932] SERVICE ALERT: prayharroldpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461259992] SERVICE ALERT: prayharroldpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461260052] SERVICE ALERT: prayharroldpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461260418] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461260653] SERVICE ALERT: prayharroldpk;Device Status;OK;HARD;3;Warming up
[1461261447] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1461261507] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1461261566] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1461262167] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1461262218] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1461262968] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 135099 found since last run. oneLevelSearchOps OK. searchOps is too high! 135100 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1461263088] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 57325 found since last run. oneLevelSearchOps OK. searchOps is too high! 57326 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1461263208] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 56432 found since last run. oneLevelSearchOps OK. searchOps is too high! 56433 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1461265606] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1461266053] SERVICE ALERT: prayharroldpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461266112] SERVICE ALERT: prayharroldpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461266172] SERVICE ALERT: prayharroldpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461266772] SERVICE ALERT: prayharroldpk;Device Status;OK;HARD;3;ORDER CARTRIDGE LESS THAN 2200 PAGES 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1461266772] SERVICE FLAPPING ALERT: prayharroldpk;Device Status;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1461266807] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 242 found since last run. oneLevelSearchOps OK. searchOps is too high! 243 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1461266927] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1461267526] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors is too high! 47353 found since last run. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1461267566] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1461267626] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1461267646] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1461267685] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1461268218] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461268246] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors is too high! 4751 found since last run. securityErrors OK. wholeSubtreeSearchOps is too high! 463 found since last run. oneLevelSearchOps OK. searchOps is too high! 464 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 501 found since last run.
[1461268366] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1461268818] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461269418] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461271218] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461271532] SERVICE ALERT: owen101lab2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461271592] SERVICE ALERT: owen101lab2;Device Status;OK;SOFT;2;Clearing paper path
[1461271818] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 88% (15 min average)
[1461272418] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1461272538] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461272658] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461274286] SERVICE FLAPPING ALERT: alexanderpk;Paper Tray Status;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1461275658] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 63% (15 min average)
[1461276278] SERVICE ALERT: owen101labcolor;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 1 Trays Empty
[1461276338] SERVICE ALERT: owen101labcolor;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 1 Trays Empty
[1461276398] SERVICE ALERT: owen101labcolor;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 1 Trays Empty
[1461276858] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461276978] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461277098] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461277262] SERVICE ALERT: hallefirst2;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1461277323] SERVICE ALERT: hallefirst2;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1461277383] SERVICE ALERT: hallefirst2;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1461278172] SERVICE FLAPPING ALERT: prayharroldpk;Device Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1461278898] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 51% (15 min average)
[1461279085] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1461279498] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1461279618] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461279738] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461280338] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461282738] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461283338] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461285428] SERVICE ALERT: crossroadspk;Black Cartridge HP CC364X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1461285430] HOST ALERT: crossroadspk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1461285478] SERVICE ALERT: crossroadspk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1461285488] SERVICE ALERT: crossroadspk;Black Cartridge HP CC364X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1461285508] HOST ALERT: crossroadspk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1461285538] SERVICE ALERT: crossroadspk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1461285548] SERVICE ALERT: crossroadspk;Black Cartridge HP CC364X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1461285568] HOST ALERT: crossroadspk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1461285598] SERVICE ALERT: crossroadspk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1461285628] HOST ALERT: crossroadspk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1461285718] HOST ALERT: crossroadspk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1461285738] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1461285808] HOST ALERT: crossroadspk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1461285808] SERVICE ALERT: crossroadspk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1461285838] HOST ALERT: crossroadspk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1461285868] SERVICE ALERT: crossroadspk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1461285898] HOST ALERT: crossroadspk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1461285928] SERVICE ALERT: crossroadspk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1461285958] HOST ALERT: crossroadspk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1461286048] HOST ALERT: crossroadspk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1461286336] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1461286456] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1461286576] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461286983] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1461287176] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461287776] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461288040] SERVICE ALERT: hallegroundcolor;Cyan Cartridge HP CE261A;CRITICAL;SOFT;1;Cyan Cartridge HP CE261A is at CRITICAL level!
[1461288099] SERVICE ALERT: hallegroundcolor;Cyan Cartridge HP CE261A;CRITICAL;SOFT;2;Cyan Cartridge HP CE261A is at CRITICAL level!
[1461288159] SERVICE ALERT: hallegroundcolor;Cyan Cartridge HP CE261A;CRITICAL;HARD;3;Cyan Cartridge HP CE261A is at CRITICAL level!
[1461288376] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461288976] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461290776] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461291376] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461294375] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1461294376] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 81% (15 min average)
[1461294496] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461294615] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461294976] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1461295096] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461295215] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1461295216] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461296416] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1461297016] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461297136] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461297256] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)