[1464062771] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 462 found since last run. oneLevelSearchOps OK. searchOps is too high! 465 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 544 found since last run.
[1464062891] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464063322] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 207.45 Gb (94%) - free 12.45 Gb (6%)
[1464063491] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 393 found since last run. oneLevelSearchOps OK. searchOps is too high! 541 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 504 found since last run.
[1464063611] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464064211] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 409 found since last run. oneLevelSearchOps OK. searchOps is too high! 477 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 625 found since last run.
[1464064332] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464064522] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 209.86 Gb (95%) - free 10.04 Gb (5%)
[1464064931] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 125 found since last run. wholeSubtreeSearchOps is too high! 233 found since last run. oneLevelSearchOps OK. searchOps is too high! 255 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 672 found since last run.
[1464065050] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464065649] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 315 found since last run. oneLevelSearchOps OK. searchOps is too high! 328 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 525 found since last run.
[1464065769] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464066322] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 207.46 Gb (94%) - free 12.44 Gb (6%)
[1464066369] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 393 found since last run. oneLevelSearchOps OK. searchOps is too high! 394 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 626 found since last run.
[1464066489] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464066857] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1464067089] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 252 found since last run. oneLevelSearchOps OK. searchOps is too high! 258 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464067209] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464067457] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464067522] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 209.71 Gb (95%) - free 10.19 Gb (5%)
[1464067577] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1464067697] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464067809] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 109 found since last run. wholeSubtreeSearchOps is too high! 445 found since last run. oneLevelSearchOps OK. searchOps is too high! 446 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 618 found since last run.
[1464067929] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464068529] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 418 found since last run. oneLevelSearchOps OK. searchOps is too high! 438 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 502 found since last run.
[1464068649] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464068897] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1464069249] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 365 found since last run. oneLevelSearchOps OK. searchOps is too high! 453 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464069322] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 207.50 Gb (94%) - free 12.40 Gb (6%)
[1464069369] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464069497] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464069969] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 421 found since last run. oneLevelSearchOps OK. searchOps is too high! 425 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 551 found since last run.
[1464070089] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464070140] SERVICE ALERT: ADRODC-LIV;AD Replication;OK;HARD;3;OK: replication is up and running
[1464070689] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 364 found since last run. oneLevelSearchOps OK. searchOps is too high! 461 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464070697] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1464070809] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464071297] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1464071409] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 306 found since last run. oneLevelSearchOps OK. searchOps is too high! 308 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464071529] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464071722] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 212.21 Gb (97%) - free 7.69 Gb (3%)
[1464072129] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 280 found since last run. oneLevelSearchOps OK. searchOps is too high! 281 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 515 found since last run.
[1464072249] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464072497] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1464072849] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 302 found since last run. oneLevelSearchOps OK. searchOps is too high! 305 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 715 found since last run.
[1464072921] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 207.55 Gb (94%) - free 12.35 Gb (6%)
[1464072969] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464073097] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464073569] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 240 found since last run. oneLevelSearchOps OK. searchOps is too high! 291 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 554 found since last run.
[1464073689] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464074121] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 209.32 Gb (95%) - free 10.58 Gb (5%)
[1464074889] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 511 found since last run.
[1464075009] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464075609] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 252 found since last run. oneLevelSearchOps OK. searchOps is too high! 265 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 570 found since last run.
[1464075729] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464075921] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 207.52 Gb (94%) - free 12.38 Gb (6%)
[1464077529] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 230 found since last run. oneLevelSearchOps OK. searchOps is too high! 231 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464077649] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464077721] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 212.64 Gb (97%) - free 7.26 Gb (3%)
[1464077897] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1464078249] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 225 found since last run. oneLevelSearchOps OK. searchOps is too high! 226 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464078369] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464078969] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 321 found since last run. oneLevelSearchOps OK. searchOps is too high! 323 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464079089] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464079097] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1464079217] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 85% (15 min average)
[1464079521] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 207.52 Gb (94%) - free 12.38 Gb (6%)
[1464079689] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 383 found since last run. oneLevelSearchOps OK. searchOps is too high! 384 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464079809] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464079817] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1464079937] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464080057] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1464080409] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 337 found since last run. oneLevelSearchOps OK. searchOps is too high! 338 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 552 found since last run.
[1464080529] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464081129] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 363 found since last run. oneLevelSearchOps OK. searchOps is too high! 365 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 527 found since last run.
[1464081249] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464081257] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 86% (15 min average)
[1464081849] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 809 found since last run. oneLevelSearchOps OK. searchOps is too high! 810 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 886 found since last run.
[1464081969] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464082521] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 210.58 Gb (96%) - free 9.32 Gb (4%)
[1464082569] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 299 found since last run. oneLevelSearchOps OK. searchOps is too high! 301 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 598 found since last run.
[1464082689] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464083289] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 473 found since last run. oneLevelSearchOps OK. searchOps is too high! 475 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 537 found since last run.
[1464083409] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464083570] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;1;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1464083657] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1464083690] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;2;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1464083777] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1464083810] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;HARD;3;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1464084009] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 469 found since last run. oneLevelSearchOps OK. searchOps is too high! 520 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464084129] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464084321] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 207.51 Gb (94%) - free 12.39 Gb (6%)
[1464084377] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464084497] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464084617] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1464084729] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 327 found since last run. oneLevelSearchOps OK. searchOps is too high! 350 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 511 found since last run.
[1464084849] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464085217] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464085449] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 311 found since last run. oneLevelSearchOps OK. searchOps is too high! 320 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 563 found since last run.
[1464085521] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 211.94 Gb (96%) - free 7.96 Gb (4%)
[1464085569] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464085817] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1464086169] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 382 found since last run. oneLevelSearchOps OK. searchOps is too high! 435 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 568 found since last run.
[1464086289] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464086417] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464086889] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 460 found since last run. oneLevelSearchOps OK. searchOps is too high! 540 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 540 found since last run.
[1464087009] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464087321] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 207.48 Gb (94%) - free 12.42 Gb (6%)
[1464087609] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 548 found since last run. oneLevelSearchOps OK. searchOps is too high! 572 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 604 found since last run.
[1464087617] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1464087729] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 488 found since last run. oneLevelSearchOps OK. searchOps is too high! 503 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464087849] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464087921] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 209.03 Gb (95%) - free 10.87 Gb (5%)
[1464088217] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464088337] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464088449] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 461 found since last run. oneLevelSearchOps OK. searchOps is too high! 503 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 536 found since last run.
[1464088457] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1464088569] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 357 found since last run. oneLevelSearchOps OK. searchOps is too high! 383 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464088689] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464089289] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 3973 found since last run. oneLevelSearchOps OK. searchOps is too high! 4017 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 3387 found since last run.
[1464089409] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 670 found since last run. oneLevelSearchOps OK. searchOps is too high! 683 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464089529] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464090129] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 241 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464090249] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464090275] HOST ALERT: hallefirstcolor;UP;HARD;10;PING OK - Packet loss = 70%, RTA = 0.41 ms
[1464090297] SERVICE ALERT: NETIDADPROXY;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464090299] SERVICE ALERT: NETIDADPROXY;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464090306] SERVICE ALERT: NETIDADPROXY;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464090321] HOST ALERT: NETIDADPROXY;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1464090321] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 207.48 Gb (94%) - free 12.42 Gb (6%)
[1464090338] SERVICE ALERT: NETIDADPROXY;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464090368] HOST ALERT: NETIDADPROXY;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1464090417] SERVICE ALERT: NETIDADPROXY;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464090419] SERVICE ALERT: NETIDADPROXY;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464090426] SERVICE ALERT: NETIDADPROXY;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464090447] HOST ALERT: NETIDADPROXY;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1464090457] SERVICE ALERT: NETIDAD;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464090458] SERVICE ALERT: NETIDADPROXY;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464090481] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1464090483] SERVICE ALERT: NETIDAD;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464090483] SERVICE ALERT: NETIDAD;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464090483] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;OK;HARD;3;Cyan Cartridge HP CE261A is at 71% - OK!
[1464090485] SERVICE ALERT: NETIDADPROXY;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464090487] HOST ALERT: NETIDAD;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1464090488] HOST ALERT: NETIDADPROXY;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1464090528] SERVICE ALERT: NETIDADPROXY;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464090537] SERVICE ALERT: NETIDADPROXY;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464090539] SERVICE ALERT: NETIDADPROXY;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464090545] SERVICE ALERT: NETIDADPROXY;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464090558] HOST ALERT: NETIDADPROXY;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1464090577] SERVICE ALERT: NETIDAD;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464090577] HOST ALERT: NETIDAD;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1464090578] SERVICE ALERT: NETIDADPROXY;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464090591] SERVICE ALERT: NETIDADPROXY;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464090592] SERVICE ALERT: hallefirstcolor;Image Fuser Kit HP 110V-CE246A;OK;HARD;3;Image Fuser Kit HP 110V-CE246A, 220V-CE247A is at 100% - OK!
[1464090603] SERVICE ALERT: NETIDAD;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464090603] SERVICE ALERT: NETIDAD;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464090605] SERVICE ALERT: NETIDADPROXY;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464090608] HOST ALERT: NETIDADPROXY;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1464090633] HOST ALERT: NETIDAD;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1464090643] SERVICE ALERT: hallefirstcolor;Device Status;OK;HARD;3;41.7 PRINTER ERROR To continue press OK 41.7 PRINTER ERROR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1464090645] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;OK;HARD;3;Yellow Cartridge HP CE262A is at 63% - OK!
[1464090648] SERVICE ALERT: NETIDADPROXY;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464090663] SERVICE ALERT: NETIDAD;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464090671] SERVICE ALERT: NETIDADPROXY;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464090672] SERVICE ALERT: NETIDAD;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464090672] SERVICE ALERT: NETIDADPROXY;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464090676] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;OK;HARD;3;Image Transfer Kit HP CE249A is at 98% - OK!
[1464090678] HOST ALERT: NETIDADPROXY;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1464090683] SERVICE ALERT: NETIDAD;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464090683] SERVICE ALERT: NETIDADPROXY;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464090683] SERVICE ALERT: NETIDAD;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464090683] SERVICE ALERT: NETIDADPROXY;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464090693] HOST ALERT: NETIDAD;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1464090697] SERVICE ALERT: NETIDAD;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464090698] SERVICE ALERT: NETIDAD;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464090711] SERVICE ALERT: NETIDAD;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464090711] SERVICE ALERT: NETIDADPROXY;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464090713] HOST ALERT: NETIDADPROXY;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1464090713] SERVICE ALERT: NETIDAD;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464090723] SERVICE ALERT: NETIDAD;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464090723] SERVICE ALERT: NETIDAD;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464090723] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;OK;HARD;3;Black Cartridge HP CE260X is at 38% - OK!
[1464090725] SERVICE ALERT: NETIDADPROXY;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464090727] HOST ALERT: NETIDAD;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1464090755] HOST ALERT: NETIDADPROXY;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1464090763] SERVICE ALERT: NETIDAD;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464090767] SERVICE ALERT: NETIDADPROXY;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464090768] SERVICE ALERT: NETIDADPROXY;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464090768] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;OK;HARD;3;Magenta Cartridge HP CE263A is at 78% - OK!
[1464090783] SERVICE ALERT: NETIDAD;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464090791] SERVICE ALERT: NETIDADPROXY;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464090791] SERVICE ALERT: NETIDADPROXY;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464090792] SERVICE ALERT: NETIDAD;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464090793] HOST ALERT: NETIDAD;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1464090797] HOST ALERT: NETIDADPROXY;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1464090803] SERVICE ALERT: NETIDAD;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464090803] SERVICE ALERT: NETIDADPROXY;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464090803] SERVICE ALERT: NETIDAD;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464090803] SERVICE ALERT: NETIDADPROXY;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464090818] SERVICE ALERT: NETIDAD;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464090831] SERVICE ALERT: NETIDAD;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464090831] SERVICE ALERT: NETIDADPROXY;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464090833] HOST ALERT: NETIDAD;DOWN;SOFT;7;(Host check timed out after 30.00 seconds)
[1464090833] SERVICE ALERT: NETIDAD;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464090849] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 248 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464090852] SERVICE ALERT: NETIDAD;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464090853] SERVICE ALERT: NETIDAD;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464090853] SERVICE ALERT: hallefirstcolor;Toner Collection Unit HP CE265A;OK;HARD;3;
[1464090853] SERVICE ALERT: NETIDADPROXY;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464090863] HOST ALERT: NETIDAD;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1464090881] SERVICE ALERT: NETIDAD;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464090883] SERVICE ALERT: NETIDAD;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464090887] SERVICE ALERT: NETIDADPROXY;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464090903] SERVICE ALERT: NETIDAD;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464090911] SERVICE ALERT: NETIDADPROXY;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464090911] HOST ALERT: NETIDAD;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1464090911] SERVICE ALERT: NETIDADPROXY;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464090912] SERVICE ALERT: NETIDAD;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464090923] SERVICE ALERT: NETIDAD;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464090923] SERVICE ALERT: NETIDADPROXY;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464090923] SERVICE ALERT: NETIDAD;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464090923] SERVICE ALERT: NETIDADPROXY;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464090938] SERVICE ALERT: NETIDAD;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464090942] HOST ALERT: NETIDAD;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1464090951] SERVICE ALERT: NETIDAD;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464090953] SERVICE ALERT: NETIDAD;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464090970] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464090972] SERVICE ALERT: NETIDAD;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464090973] SERVICE ALERT: NETIDAD;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464090973] SERVICE ALERT: NETIDADPROXY;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464091001] SERVICE ALERT: NETIDAD;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464091003] SERVICE ALERT: NETIDAD;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464091007] SERVICE ALERT: NETIDADPROXY;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464091010] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;OK;HARD;3;OK: No entries found
[1464091092] SERVICE ALERT: NETIDAD;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464091093] SERVICE ALERT: NETIDAD;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464091093] SERVICE ALERT: NETIDADPROXY;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464091096] HOST ALERT: NETIDAD;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1464091101] SERVICE ALERT: NETIDAD;Service: Group Policy;OK;SOFT;3;gpsvc: Started
[1464091117] SERVICE ALERT: NETIDADPROXY;Service: RPC;OK;HARD;3;RpcSs: Started
[1464091119] HOST ALERT: NETIDADPROXY;UP;HARD;10;PING WARNING - Packet loss = 80%, RTA = 0.37 ms
[1464091119] SERVICE ALERT: NETIDADPROXY;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1464091125] SERVICE ALERT: NETIDADPROXY;CPU Load;OK;HARD;3;CPU Load 5% (15 min average)
[1464091158] SERVICE ALERT: NETIDADPROXY;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1464091277] SERVICE ALERT: NETIDAD;Service: RPC;OK;HARD;3;RpcSs: Started
[1464091303] SERVICE ALERT: NETIDAD;VMWare Tools;OK;HARD;3;VMTools: Started
[1464091303] SERVICE ALERT: NETIDAD;C:\ Drive Space;OK;HARD;3;c:\ - total: 79.66 Gb - used: 14.13 Gb (18%) - free 65.52 Gb (82%)
[1464091305] SERVICE ALERT: NETIDADPROXY;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 11 minute(s)
[1464091348] SERVICE ALERT: NETIDADPROXY;VMWare Tools;OK;HARD;3;VMTools: Started
[1464091411] SERVICE ALERT: NETIDADPROXY;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1464091483] SERVICE ALERT: NETIDAD;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1464091491] SERVICE ALERT: NETIDADPROXY;C:\ Drive Space;OK;HARD;3;c:\ - total: 79.66 Gb - used: 13.09 Gb (16%) - free 66.57 Gb (84%)
[1464091492] SERVICE ALERT: NETIDAD;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 16 minute(s)
[1464091503] SERVICE ALERT: NETIDAD;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1464091503] SERVICE ALERT: NETIDAD;Memory Usage;OK;HARD;3;Memory usage: total:10943.64 MB - used: 1347.96 MB (12%) - free: 9595.69 MB (88%)
[1464091503] SERVICE ALERT: NETIDADPROXY;Memory Usage;OK;HARD;3;Memory usage: total:4799.59 MB - used: 1170.15 MB (24%) - free: 3629.45 MB (76%)
[1464091503] SERVICE ALERT: NETIDADPROXY;Service: WMI;OK;HARD;3;Winmgmt: Started
[1464091517] SERVICE ALERT: NETIDAD;CPU Load;OK;HARD;3;CPU Load 13% (15 min average)
[1464091530] SERVICE ALERT: NETIDAD;Service: Windows Time;OK;HARD;3;W32Time: Started
[1464091533] SERVICE ALERT: NETIDAD;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1464091569] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 420 found since last run. oneLevelSearchOps OK. searchOps is too high! 672 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464091583] SERVICE ALERT: NETIDAD;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1464091587] SERVICE ALERT: NETIDADPROXY;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1464091672] SERVICE ALERT: NETIDAD;Service: WMI;OK;HARD;3;Winmgmt: Started
[1464091673] SERVICE ALERT: NETIDAD;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1464091673] SERVICE ALERT: NETIDADPROXY;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1464091688] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 377 found since last run. oneLevelSearchOps OK. searchOps is too high! 450 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464091808] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 227 found since last run. oneLevelSearchOps OK. searchOps is too high! 273 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464092057] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1464092091] SERVICE ALERT: NETIDADPROXY;Service: Windows Time;OK;HARD;3;W32Time: Started
[1464092121] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 212.02 Gb (96%) - free 7.88 Gb (4%)
[1464093857] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1464093921] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 207.50 Gb (94%) - free 12.40 Gb (6%)
[1464093977] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1464094097] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1464094521] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 210.05 Gb (96%) - free 9.85 Gb (4%)
[1464095297] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 78% (15 min average)
[1464095897] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1464096017] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1464096137] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 74% (15 min average)
[1464096737] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464096857] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1464096921] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 207.49 Gb (94%) - free 12.41 Gb (6%)
[1464096977] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1464097485] SERVICE ALERT: SNAP;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464097485] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.40)
[1464097488] SERVICE ALERT: SNAP;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464097491] SERVICE ALERT: SNAP;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464097491] SERVICE ALERT: SNAP;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464097494] SERVICE ALERT: SNAP;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464097507] HOST ALERT: SNAP;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1464097521] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 209.55 Gb (95%) - free 10.34 Gb (5%)
[1464097594] SERVICE ALERT: SNAP;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1464097595] SERVICE ALERT: SNAP;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1464097599] SERVICE ALERT: SNAP;Service: RPC;OK;SOFT;2;RpcSs: Started
[1464097600] SERVICE ALERT: SNAP;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1464097611] SERVICE ALERT: SNAP;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1464097837] SERVICE ALERT: TABLEAU;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1464097957] SERVICE ALERT: TABLEAU;CPU Load;OK;SOFT;2;CPU Load 66% (15 min average)
[1464098177] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 83% (15 min average)
[1464098777] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1464098897] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464099017] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1464099320] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 207.51 Gb (94%) - free 12.39 Gb (6%)
[1464100958] SERVICE ALERT: TABLEAU;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1464101027] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 48906 found since last run. oneLevelSearchOps OK. searchOps is too high! 48909 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464101077] SERVICE ALERT: TABLEAU;CPU Load;OK;SOFT;2;CPU Load 91% (15 min average)
[1464101120] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 210.13 Gb (96%) - free 9.77 Gb (4%)
[1464101147] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464101160] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.40)
[1464101197] SERVICE ALERT: SNAP;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464101197] SERVICE ALERT: SNAP;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464101200] SERVICE ALERT: SNAP;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464101200] HOST ALERT: SNAP;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.40)
[1464101203] SERVICE ALERT: SNAP;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464101256] SERVICE ALERT: SNAP;Service: User Profile;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464101261] HOST ALERT: SNAP;UP;SOFT;3;PING OK - Packet loss = 16%, RTA = 185.51 ms
[1464101314] SERVICE ALERT: SNAP;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1464101315] SERVICE ALERT: SNAP;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1464101319] SERVICE ALERT: SNAP;Service: RPC;OK;SOFT;2;RpcSs: Started
[1464101320] SERVICE ALERT: SNAP;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1464101364] SERVICE ALERT: SNAP;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1464101417] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 87% (15 min average)
[1464101823] SERVICE ALERT: SNAP;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464101826] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.40)
[1464101860] SERVICE ALERT: SNAP;CPU Load;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464101863] HOST ALERT: SNAP;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.40)
[1464101863] SERVICE ALERT: SNAP;Uptime;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464101899] HOST ALERT: SNAP;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1464101951] SERVICE ALERT: SNAP;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464101958] SERVICE ALERT: SNAP;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464101958] SERVICE ALERT: SNAP;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464101962] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.40)
[1464101967] SERVICE ALERT: SNAP;Service: User Profile;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464101970] HOST ALERT: SNAP;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.40)
[1464101980] SERVICE ALERT: SNAP;CPU Load;CRITICAL;SOFT;2;connect to address 164.76.102.40 and port 12489: No route to host
[1464101983] SERVICE ALERT: SNAP;Uptime;CRITICAL;SOFT;2;connect to address 164.76.102.40 and port 12489: No route to host
[1464101986] HOST ALERT: SNAP;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.102.40)
[1464102017] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1464102057] HOST ALERT: SNAP;DOWN;SOFT;4;CRITICAL - Host Unreachable (164.76.102.40)
[1464102060] SERVICE ALERT: SNAP;Service: Windows Time;CRITICAL;HARD;3;connect to address 164.76.102.40 and port 12489: No route to host
[1464102060] SERVICE ALERT: SNAP;C:\ Drive Space;CRITICAL;SOFT;2;connect to address 164.76.102.40 and port 12489: No route to host
[1464102060] SERVICE ALERT: SNAP;Service: Group Policy;CRITICAL;SOFT;2;connect to address 164.76.102.40 and port 12489: No route to host
[1464102087] SERVICE ALERT: SNAP;Service: User Profile;CRITICAL;SOFT;2;connect to address 164.76.102.40 and port 12489: No route to host
[1464102091] HOST ALERT: SNAP;UP;SOFT;5;PING OK - Packet loss = 0%, RTA = 601.75 ms
[1464102097] SERVICE ALERT: SNAP;CPU Load;OK;SOFT;3;CPU Load 0% (15 min average)
[1464102103] SERVICE ALERT: SNAP;Uptime;OK;SOFT;3;System Uptime - 3 day(s) 13 hour(s) 53 minute(s)
[1464102137] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464102177] SERVICE ALERT: SNAP;Service: Group Policy;OK;SOFT;3;gpsvc: Started
[1464102177] SERVICE ALERT: SNAP;C:\ Drive Space;OK;SOFT;3;c:\ - total: 1116.65 Gb - used: 36.08 Gb (3%) - free 1080.57 Gb (97%)
[1464102204] SERVICE ALERT: SNAP;Service: User Profile;OK;SOFT;3;ProfSvc: Started
[1464102257] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1464102651] SERVICE ALERT: SNAP;Service: Windows Time;OK;HARD;3;W32Time: Started
[1464102789] SERVICE ALERT: SNAP;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464102789] SERVICE ALERT: SNAP;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464102792] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.40)
[1464102807] SERVICE ALERT: SNAP;Service: User Profile;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464102810] HOST ALERT: SNAP;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.40)
[1464102844] HOST ALERT: SNAP;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1464102897] SERVICE ALERT: SNAP;C:\ Drive Space;OK;SOFT;2;c:\ - total: 1116.65 Gb - used: 36.08 Gb (3%) - free 1080.57 Gb (97%)
[1464102897] SERVICE ALERT: SNAP;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1464102924] SERVICE ALERT: SNAP;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1464103107] SERVICE ALERT: SNAP;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464103110] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.40)
[1464103117] SERVICE ALERT: SNAP;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464103117] SERVICE ALERT: SNAP;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464103120] SERVICE ALERT: SNAP;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464103227] SERVICE ALERT: SNAP;Service: Firewall;CRITICAL;SOFT;2;connect to address 164.76.102.40 and port 12489: No route to host
[1464103230] HOST ALERT: SNAP;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.40)
[1464103237] SERVICE ALERT: SNAP;Service: WMI;CRITICAL;SOFT;2;connect to address 164.76.102.40 and port 12489: No route to host
[1464103237] SERVICE ALERT: SNAP;Service: DNS Client;CRITICAL;SOFT;2;connect to address 164.76.102.40 and port 12489: No route to host
[1464103240] SERVICE ALERT: SNAP;Service: RPC;OK;SOFT;2;RpcSs: Started
[1464103255] HOST ALERT: SNAP;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.39 ms
[1464103335] SERVICE ALERT: SNAP;Service: Firewall;OK;SOFT;3;MpsSvc: Started
[1464103354] SERVICE ALERT: SNAP;Service: DNS Client;OK;SOFT;3;Dnscache: Started
[1464103355] SERVICE ALERT: SNAP;Service: WMI;OK;SOFT;3;Winmgmt: Started
[1464103457] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1464103509] SERVICE ALERT: SNAP;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464103509] SERVICE ALERT: SNAP;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464103509] SERVICE ALERT: SNAP;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464103512] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.40)
[1464103520] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 207.52 Gb (94%) - free 12.38 Gb (6%)
[1464103530] HOST ALERT: SNAP;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.41 ms
[1464103616] SERVICE ALERT: SNAP;C:\ Drive Space;OK;SOFT;2;c:\ - total: 1116.65 Gb - used: 36.08 Gb (3%) - free 1080.57 Gb (97%)
[1464103619] SERVICE ALERT: SNAP;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1464103623] SERVICE ALERT: SNAP;Memory Usage;OK;SOFT;2;Memory usage: total:16362.75 MB - used: 1794.75 MB (11%) - free: 14568.00 MB (89%)
[1464103841] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.40)
[1464103841] HOST FLAPPING ALERT: SNAP;STARTED; Host appears to have started flapping (21.2% change > 20.0% threshold)
[1464103841] SERVICE ALERT: SNAP;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464103854] SERVICE ALERT: SNAP;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464103857] HOST ALERT: SNAP;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.40)
[1464103900] SERVICE ALERT: SNAP;CPU Load;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464103903] HOST ALERT: SNAP;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.102.40)
[1464103903] SERVICE ALERT: SNAP;Uptime;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464103938] SERVICE ALERT: SNAP;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464103941] HOST ALERT: SNAP;DOWN;SOFT;4;CRITICAL - Host Unreachable (164.76.102.40)
[1464103966] SERVICE ALERT: SNAP;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464103966] SERVICE ALERT: SNAP;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464103969] SERVICE ALERT: SNAP;Service: RPC;CRITICAL;SOFT;2;connect to address 164.76.102.40 and port 12489: No route to host
[1464103969] HOST ALERT: SNAP;DOWN;SOFT;5;CRITICAL - Host Unreachable (164.76.102.40)
[1464103972] SERVICE ALERT: SNAP;Service: Windows Time;CRITICAL;SOFT;2;connect to address 164.76.102.40 and port 12489: No route to host
[1464104020] SERVICE ALERT: SNAP;CPU Load;CRITICAL;SOFT;2;connect to address 164.76.102.40 and port 12489: No route to host
[1464104023] SERVICE ALERT: SNAP;Uptime;CRITICAL;SOFT;2;connect to address 164.76.102.40 and port 12489: No route to host
[1464104023] HOST FLAPPING ALERT: SNAP;STOPPED; Host appears to have stopped flapping (4.3% change < 5.0% threshold)
[1464104023] HOST ALERT: SNAP;DOWN;SOFT;6;CRITICAL - Host Unreachable (164.76.102.40)
[1464104044] HOST ALERT: SNAP;UP;SOFT;7;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1464104055] SERVICE ALERT: SNAP;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1464104057] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1464104094] SERVICE ALERT: SNAP;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464104094] SERVICE ALERT: SNAP;Service: Windows Time;OK;SOFT;3;W32Time: Started
[1464104095] SERVICE ALERT: SNAP;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464104099] SERVICE ALERT: SNAP;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464104120] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 210.07 Gb (96%) - free 9.83 Gb (4%)
[1464104137] SERVICE ALERT: SNAP;CPU Load;OK;SOFT;3;CPU Load 0% (15 min average)
[1464104143] SERVICE ALERT: SNAP;Uptime;OK;SOFT;3;System Uptime - 3 day(s) 14 hour(s) 27 minute(s)
[1464104194] SERVICE ALERT: SNAP;Service: DNS Client;OK;SOFT;3;Dnscache: Started
[1464104195] SERVICE ALERT: SNAP;Service: WMI;OK;SOFT;3;Winmgmt: Started
[1464104406] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.40)
[1464104477] HOST ALERT: SNAP;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.40)
[1464104541] HOST ALERT: SNAP;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.46 ms
[1464104657] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1464104679] SERVICE ALERT: SNAP;Service: RPC;OK;HARD;3;RpcSs: Started
[1464105257] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464105347] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 3305 found since last run. oneLevelSearchOps OK. searchOps is too high! 3306 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464105467] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464105920] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 207.58 Gb (94%) - free 12.32 Gb (6%)
[1464106372] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.40)
[1464106443] SERVICE ALERT: SNAP;Service: NetLogin;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464106443] HOST ALERT: SNAP;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.40)
[1464106485] HOST ALERT: SNAP;UP;SOFT;3;PING WARNING - Packet loss = 93%, RTA = 0.38 ms
[1464106499] SERVICE ALERT: SNAP;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464106511] SERVICE ALERT: SNAP;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464106560] SERVICE ALERT: SNAP;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1464106595] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1464106599] SERVICE ALERT: SNAP;Service: RPC;OK;SOFT;2;RpcSs: Started
[1464106611] SERVICE ALERT: SNAP;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1464106654] SERVICE ALERT: marshallpk;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! All Trays Still Full
[1464107120] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 212.99 Gb (97%) - free 6.91 Gb (3%)
[1464107332] SERVICE ALERT: SNAP;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464107440] SERVICE ALERT: SNAP;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1464108823] SERVICE ALERT: studentcenterpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.150.252. Is the printer off?  Is the network down?
[1464108855] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1464108884] SERVICE ALERT: studentcenterpk;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! All Trays Still Full
[1464108920] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 207.61 Gb (94%) - free 12.29 Gb (6%)
[1464109167] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.40)
[1464109230] HOST ALERT: SNAP;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.40)
[1464109294] HOST ALERT: SNAP;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.44 ms
[1464109454] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464109467] SERVICE ALERT: SNAP;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464109470] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.40)
[1464109526] SERVICE ALERT: SNAP;Service: User Profile;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464109529] HOST ALERT: SNAP;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.40)
[1464109540] SERVICE ALERT: SNAP;CPU Load;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464109547] HOST ALERT: SNAP;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1464109547] HOST FLAPPING ALERT: SNAP;STARTED; Host appears to have started flapping (22.1% change > 20.0% threshold)
[1464109575] SERVICE ALERT: SNAP;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1464109631] SERVICE ALERT: SNAP;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464109635] SERVICE ALERT: SNAP;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464109636] SERVICE ALERT: SNAP;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464109643] SERVICE ALERT: SNAP;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1464109643] SERVICE ALERT: SNAP;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464109657] SERVICE ALERT: SNAP;CPU Load;OK;SOFT;2;CPU Load 1% (15 min average)
[1464109731] SERVICE ALERT: SNAP;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1464109742] SERVICE ALERT: SNAP;C:\ Drive Space;CRITICAL;SOFT;2;connect to address 164.76.102.40 and port 12489: No route to host
[1464109745] SERVICE ALERT: SNAP;Service: Group Policy;CRITICAL;SOFT;2;connect to address 164.76.102.40 and port 12489: No route to host
[1464109745] HOST FLAPPING ALERT: SNAP;STOPPED; Host appears to have stopped flapping (4.3% change < 5.0% threshold)
[1464109745] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.40)
[1464109748] SERVICE ALERT: SNAP;Memory Usage;CRITICAL;SOFT;2;connect to address 164.76.102.40 and port 12489: No route to host
[1464109843] SERVICE ALERT: SNAP;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464109846] HOST ALERT: SNAP;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.40)
[1464109856] SERVICE ALERT: SNAP;Service: Group Policy;OK;SOFT;3;gpsvc: Started
[1464109856] SERVICE ALERT: SNAP;C:\ Drive Space;OK;SOFT;3;c:\ - total: 1116.65 Gb - used: 36.07 Gb (3%) - free 1080.58 Gb (97%)
[1464109863] SERVICE ALERT: SNAP;Memory Usage;OK;SOFT;3;Memory usage: total:16362.75 MB - used: 1808.32 MB (11%) - free: 14554.43 MB (89%)
[1464109867] HOST ALERT: SNAP;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.44 ms
[1464109960] SERVICE ALERT: SNAP;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1464110343] SERVICE ALERT: SNAP;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464110346] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.40)
[1464110417] HOST ALERT: SNAP;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.40)
[1464110451] SERVICE ALERT: SNAP;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1464110455] HOST ALERT: SNAP;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1464110720] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 210.35 Gb (96%) - free 9.55 Gb (4%)
[1464110854] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1464110867] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 217 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464110915] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1464110975] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1464110987] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 253 found since last run. oneLevelSearchOps OK. searchOps is too high! 254 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464111062] SERVICE ALERT: SNAP;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464111062] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.40)
[1464111065] SERVICE ALERT: SNAP;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464111065] SERVICE ALERT: SNAP;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464111068] SERVICE ALERT: SNAP;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464111107] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464111164] HOST ALERT: SNAP;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1464111171] SERVICE ALERT: SNAP;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1464111175] SERVICE ALERT: SNAP;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1464111176] SERVICE ALERT: SNAP;C:\ Drive Space;OK;SOFT;2;c:\ - total: 1116.65 Gb - used: 36.08 Gb (3%) - free 1080.58 Gb (97%)
[1464111183] SERVICE ALERT: SNAP;Memory Usage;OK;SOFT;2;Memory usage: total:16362.75 MB - used: 1845.41 MB (11%) - free: 14517.34 MB (89%)
[1464111254] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1464111853] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464111995] SERVICE ALERT: SNAP;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464111997] SERVICE ALERT: SNAP;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464111997] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.40)
[1464112000] SERVICE ALERT: SNAP;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464112000] SERVICE ALERT: SNAP;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464112055] SERVICE ALERT: SNAP;Service: User Profile;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464112058] SERVICE ALERT: SNAP;CPU Load;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464112058] HOST ALERT: SNAP;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.40)
[1464112107] SERVICE ALERT: SNAP;Service: Firewall;CRITICAL;SOFT;2;connect to address 164.76.102.40 and port 12489: No route to host
[1464112110] HOST ALERT: SNAP;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.102.40)
[1464112117] SERVICE ALERT: SNAP;Service: WMI;CRITICAL;SOFT;2;connect to address 164.76.102.40 and port 12489: No route to host
[1464112117] SERVICE ALERT: SNAP;Service: DNS Client;CRITICAL;SOFT;2;connect to address 164.76.102.40 and port 12489: No route to host
[1464112120] SERVICE ALERT: SNAP;Service: RPC;CRITICAL;SOFT;2;connect to address 164.76.102.40 and port 12489: No route to host
[1464112163] SERVICE ALERT: SNAP;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1464112167] HOST ALERT: SNAP;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1464112176] SERVICE ALERT: SNAP;CPU Load;OK;SOFT;2;CPU Load 9% (15 min average)
[1464112215] SERVICE ALERT: SNAP;Service: Firewall;OK;SOFT;3;MpsSvc: Started
[1464112234] SERVICE ALERT: SNAP;Service: DNS Client;OK;SOFT;3;Dnscache: Started
[1464112235] SERVICE ALERT: SNAP;Service: WMI;OK;SOFT;3;Winmgmt: Started
[1464112239] SERVICE ALERT: SNAP;Service: RPC;OK;SOFT;3;RpcSs: Started
[1464112372] SERVICE ALERT: SNAP;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464112372] SERVICE ALERT: SNAP;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464112375] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.40)
[1464112446] HOST ALERT: SNAP;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.40)
[1464112480] SERVICE ALERT: SNAP;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1464112484] HOST ALERT: SNAP;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1464112491] SERVICE ALERT: SNAP;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1464112520] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 207.64 Gb (94%) - free 12.26 Gb (6%)
[1464113053] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 82% (15 min average)
[1464113435] SERVICE ALERT: SNAP;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464113547] SERVICE ALERT: SNAP;Service: Firewall;CRITICAL;SOFT;2;connect to address 164.76.102.40 and port 12489: No route to host
[1464113550] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.40)
[1464113579] HOST ALERT: SNAP;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.40 ms
[1464113653] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464113655] SERVICE ALERT: SNAP;Service: Firewall;OK;SOFT;3;MpsSvc: Started
[1464113720] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 211.25 Gb (96%) - free 8.65 Gb (4%)
[1464113773] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464113893] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464114428] SERVICE ALERT: idm1.emich.edu;Google Driver;CRITICAL;SOFT;1;CRITICAL: cn=gdap,cn=driverset1,o=system is stopped!
[1464114433] SERVICE ALERT: idm1.emich.edu;Oracle Driver;CRITICAL;SOFT;1;CRITICAL: cn=oradap,cn=driverset1,o=system is stopped!
[1464114549] SERVICE ALERT: idm1.emich.edu;Google Driver;CRITICAL;SOFT;2;CRITICAL: cn=gdap,cn=driverset1,o=system is stopped!
[1464114554] SERVICE ALERT: idm1.emich.edu;Oracle Driver;CRITICAL;SOFT;2;CRITICAL: cn=oradap,cn=driverset1,o=system is stopped!
[1464114669] SERVICE ALERT: idm1.emich.edu;Google Driver;CRITICAL;HARD;3;CRITICAL: cn=gdap,cn=driverset1,o=system is stopped!
[1464114674] SERVICE ALERT: idm1.emich.edu;Oracle Driver;CRITICAL;HARD;3;CRITICAL: cn=oradap,cn=driverset1,o=system is stopped!
[1464114955] SERVICE ALERT: SNAP;Uptime;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464114958] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.40)
[1464114963] SERVICE ALERT: SNAP;Service: NetLogin;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464115063] SERVICE ALERT: SNAP;Uptime;OK;SOFT;2;System Uptime - 3 day(s) 17 hour(s) 29 minute(s)
[1464115067] HOST ALERT: SNAP;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.41 ms
[1464115080] SERVICE ALERT: SNAP;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1464115093] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1464115143] SERVICE ALERT: idm1.emich.edu;Loopback Driver;CRITICAL;SOFT;1;CRITICAL: cn=Generic Loopback,cn=driverset1,o=system is stopped!
[1464115264] SERVICE ALERT: idm1.emich.edu;Loopback Driver;OK;SOFT;2;OK: cn=Generic Loopback,cn=driverset1,o=system is running.
[1464115269] SERVICE ALERT: idm1.emich.edu;Google Driver;OK;HARD;3;OK: cn=gdap,cn=driverset1,o=system is running.
[1464115274] SERVICE ALERT: idm1.emich.edu;Oracle Driver;OK;HARD;3;OK: cn=oradap,cn=driverset1,o=system is running.
[1464115483] SERVICE ALERT: SNAP;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464115486] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.40)
[1464115520] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 207.63 Gb (94%) - free 12.27 Gb (6%)
[1464115550] HOST ALERT: SNAP;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.46 ms
[1464115550] HOST FLAPPING ALERT: SNAP;STARTED; Host appears to have started flapping (20.4% change > 20.0% threshold)
[1464115600] SERVICE ALERT: SNAP;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1464115693] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 84% (15 min average)
[1464115775] SERVICE ALERT: SNAP;Service: User Profile;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464115775] SERVICE ALERT: marshallpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1464115778] SERVICE ALERT: SNAP;CPU Load;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464115778] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.40)
[1464115837] HOST ALERT: SNAP;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1464115864] SERVICE ALERT: idm1.emich.edu;Loopback Driver;CRITICAL;SOFT;1;CRITICAL: cn=Generic Loopback,cn=driverset1,o=system is stopped!
[1464115868] SERVICE ALERT: idm1.emich.edu;Google Driver;CRITICAL;SOFT;1;CRITICAL: cn=gdap,cn=driverset1,o=system is stopped!
[1464115882] SERVICE ALERT: SNAP;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1464115894] SERVICE ALERT: SNAP;CPU Load;OK;SOFT;2;CPU Load 0% (15 min average)
[1464115907] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 292 found since last run. oneLevelSearchOps OK. searchOps is too high! 297 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464115983] SERVICE ALERT: idm1.emich.edu;Loopback Driver;CRITICAL;SOFT;2;CRITICAL: cn=Generic Loopback,cn=driverset1,o=system is stopped!
[1464115988] SERVICE ALERT: idm1.emich.edu;Google Driver;CRITICAL;SOFT;2;CRITICAL: cn=gdap,cn=driverset1,o=system is stopped!
[1464116027] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464116104] SERVICE ALERT: idm1.emich.edu;Loopback Driver;CRITICAL;HARD;3;CRITICAL: cn=Generic Loopback,cn=driverset1,o=system is stopped!
[1464116109] SERVICE ALERT: idm1.emich.edu;Google Driver;CRITICAL;HARD;3;CRITICAL: cn=gdap,cn=driverset1,o=system is stopped!
[1464116483] HOST FLAPPING ALERT: SNAP;STOPPED; Host appears to have stopped flapping (4.0% change < 5.0% threshold)
[1464116538] SERVICE ALERT: EMUPRINTPS;Print Spooler;CRITICAL;SOFT;1;Spooler: Stopped
[1464116658] SERVICE ALERT: EMUPRINTPS;Print Spooler;CRITICAL;SOFT;2;Spooler: Stopped
[1464116720] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 212.22 Gb (97%) - free 7.68 Gb (3%)
[1464116778] SERVICE ALERT: EMUPRINTPS;Print Spooler;CRITICAL;HARD;3;Spooler: Stopped
[1464116893] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464117013] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464117133] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464117227] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 263 found since last run. oneLevelSearchOps OK. searchOps is too high! 264 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464117292] SERVICE ALERT: EMUPRINTPS;Print Spooler;OK;HARD;3;Spooler: Started
[1464117347] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464117733] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1464117904] SERVICE ALERT: idm1.emich.edu;Loopback Driver;OK;HARD;3;OK: cn=Generic Loopback,cn=driverset1,o=system is running.
[1464118333] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464118520] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 207.62 Gb (94%) - free 12.28 Gb (6%)
[1464118745] SERVICE ALERT: idm2.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 13
[1464118773] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1464118787] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 13
[1464118834] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1464118865] SERVICE ALERT: idm2.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 16
[1464118873] SERVICE ALERT: idm1.emich.edu;Oracle Driver;CRITICAL;SOFT;1;Unknown: Error reported?
[1464118893] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1464118906] SERVICE ALERT: idm1.emich.edu;NSDS Errors;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 60 seconds.
[1464118907] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 16
[1464118940] SERVICE ALERT: idm1.emich.edu;OfficeAD Driver;CRITICAL;SOFT;1;Unknown: Error reported?
[1464118984] SERVICE ALERT: idm1.emich.edu;AD Driver;CRITICAL;SOFT;1;Unknown: Error reported?
[1464118985] SERVICE ALERT: idm2.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 16
[1464118994] SERVICE ALERT: idm1.emich.edu;Oracle Driver;CRITICAL;SOFT;2;Unknown: Error reported?
[1464119026] SERVICE ALERT: idm1.emich.edu;NSDS Errors;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 60 seconds.
[1464119027] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 8
[1464119061] SERVICE ALERT: idm1.emich.edu;OfficeAD Driver;CRITICAL;SOFT;2;CRITICAL: cn=OfficeAD,cn=driverset1,o=system is stopped!
[1464119103] SERVICE ALERT: idm1.emich.edu;AD Driver;CRITICAL;SOFT;2;CRITICAL: cn=Campus AD,cn=driverset1,o=system is stopped!
[1464119104] SERVICE ALERT: idm1.emich.edu;Loopback Driver;CRITICAL;SOFT;1;CRITICAL: cn=Generic Loopback,cn=driverset1,o=system is stopped!
[1464119113] SERVICE ALERT: idm1.emich.edu;Oracle Driver;CRITICAL;HARD;3;CRITICAL: cn=oradap,cn=driverset1,o=system is stopped!
[1464119120] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 209.12 Gb (95%) - free 10.78 Gb (5%)
[1464119129] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;CRITICAL;SOFT;1;CRITICAL: cn=ADAUTH Driver,cn=driverset1,o=system is stopped!
[1464119139] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 2
[1464119146] SERVICE ALERT: idm1.emich.edu;NSDS Errors;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 60 seconds.
[1464119147] SERVICE ALERT: idm1.emich.edu;LDAP Stats;UNKNOWN;SOFT;1;Difference is negative. LDAP recently rebooted.
[1464119180] SERVICE ALERT: idm1.emich.edu;OfficeAD Driver;CRITICAL;HARD;3;CRITICAL: cn=OfficeAD,cn=driverset1,o=system is stopped!
[1464119223] SERVICE ALERT: idm1.emich.edu;Loopback Driver;OK;SOFT;2;OK: cn=Generic Loopback,cn=driverset1,o=system is running.
[1464119223] SERVICE ALERT: idm1.emich.edu;AD Driver;OK;SOFT;3;OK: cn=Campus AD,cn=driverset1,o=system is running.
[1464119248] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;OK;SOFT;2;OK: cn=ADAUTH Driver,cn=driverset1,o=system is running.
[1464119258] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 1
[1464119267] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464119378] SERVICE ALERT: idm4.emich.edu;NSDS Errors;OK;SOFT;3;OK: Total errors: 0
[1464119585] SERVICE ALERT: idm2.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1464119627] SERVICE ALERT: idm3.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1464119686] SERVICE ALERT: idm1.emich.edu;NSDS Errors;OK;HARD;3;OK: 0 Errors Reported
[1464119708] SERVICE ALERT: idm1.emich.edu;Google Driver;OK;HARD;3;OK: cn=gdap,cn=driverset1,o=system is running.
[1464119708] SERVICE FLAPPING ALERT: idm1.emich.edu;Google Driver;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1464119714] SERVICE ALERT: idm1.emich.edu;Oracle Driver;OK;HARD;3;OK: cn=oradap,cn=driverset1,o=system is running.
[1464119714] SERVICE FLAPPING ALERT: idm1.emich.edu;Oracle Driver;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1464119823] SERVICE ALERT: idm1.emich.edu;Loopback Driver;CRITICAL;SOFT;1;Unknown: Error reported?
[1464119824] SERVICE ALERT: idm1.emich.edu;AD Driver;CRITICAL;SOFT;1;Unknown: Error reported?
[1464119843] SERVICE ALERT: marshallpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1464119848] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;CRITICAL;SOFT;1;Unknown: Error reported?
[1464119867] SERVICE ALERT: idm1.emich.edu;LDAP Stats;CRITICAL;SOFT;1;CRITICAL: Failed to connect to LDAP!
[1464119903] SERVICE ALERT: marshallpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1464119943] SERVICE ALERT: idm1.emich.edu;AD Driver;CRITICAL;SOFT;2;Unknown: Error reported?
[1464119944] SERVICE ALERT: idm1.emich.edu;Loopback Driver;CRITICAL;SOFT;2;Unknown: Error reported?
[1464119963] SERVICE ALERT: marshallpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1464119969] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;CRITICAL;SOFT;2;Unknown: Error reported?
[1464119978] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 16
[1464119987] SERVICE ALERT: idm1.emich.edu;LDAP Stats;CRITICAL;SOFT;2;CRITICAL: Failed to connect to LDAP!
[1464120063] SERVICE ALERT: idm1.emich.edu;Loopback Driver;CRITICAL;HARD;3;Unknown: Error reported?
[1464120063] SERVICE ALERT: idm1.emich.edu;AD Driver;CRITICAL;HARD;3;Unknown: Error reported?
[1464120088] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;CRITICAL;HARD;3;Unknown: Error reported?
[1464120098] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 16
[1464120107] SERVICE ALERT: idm1.emich.edu;LDAP Stats;CRITICAL;HARD;3;CRITICAL: Failed to connect to LDAP!
[1464120133] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1464120185] SERVICE ALERT: idm2.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 6
[1464120218] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 3
[1464120227] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 3
[1464120305] SERVICE ALERT: idm2.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 1
[1464120308] SERVICE ALERT: idm1.emich.edu;Google Driver;CRITICAL;SOFT;1;CRITICAL: cn=gdap,cn=driverset1,o=system is stopped!
[1464120313] SERVICE ALERT: idm1.emich.edu;Oracle Driver;CRITICAL;SOFT;1;CRITICAL: cn=oradap,cn=driverset1,o=system is stopped!
[1464120346] SERVICE ALERT: idm1.emich.edu;NSDS Errors;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 60 seconds.
[1464120347] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 1
[1464120380] SERVICE ALERT: idm1.emich.edu;OfficeAD Driver;OK;HARD;3;OK: cn=OfficeAD,cn=driverset1,o=system is running.
[1464120406] SERVICE ALERT: idm1.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 1
[1464120426] SERVICE ALERT: idm2.emich.edu;NSDS Errors;OK;SOFT;3;OK: Total errors: 0
[1464120428] SERVICE ALERT: idm1.emich.edu;Google Driver;OK;SOFT;2;OK: cn=gdap,cn=driverset1,o=system is running.
[1464120434] SERVICE ALERT: idm1.emich.edu;Oracle Driver;OK;SOFT;2;OK: cn=oradap,cn=driverset1,o=system is running.
[1464120466] SERVICE ALERT: idm3.emich.edu;NSDS Errors;OK;SOFT;3;OK: Total errors: 0
[1464120526] SERVICE ALERT: idm1.emich.edu;NSDS Errors;OK;SOFT;3;OK: 0 Errors Reported
[1464120664] SERVICE ALERT: idm1.emich.edu;Loopback Driver;OK;HARD;3;OK: cn=Generic Loopback,cn=driverset1,o=system is running.
[1464120664] SERVICE FLAPPING ALERT: idm1.emich.edu;Loopback Driver;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1464120664] SERVICE ALERT: idm1.emich.edu;AD Driver;OK;HARD;3;OK: cn=Campus AD,cn=driverset1,o=system is running.
[1464120707] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464120733] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464120818] SERVICE ALERT: idm4.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1464120920] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 207.61 Gb (94%) - free 12.29 Gb (6%)
[1464121034] SERVICE ALERT: idm1.emich.edu;Oracle Driver;CRITICAL;SOFT;1;CRITICAL: cn=oradap,cn=driverset1,o=system is stopped!
[1464121154] SERVICE ALERT: idm1.emich.edu;Oracle Driver;CRITICAL;SOFT;2;CRITICAL: cn=oradap,cn=driverset1,o=system is stopped!
[1464121163] SERVICE ALERT: marshallpk;Device Status;OK;HARD;3;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1464121273] SERVICE ALERT: idm1.emich.edu;Oracle Driver;CRITICAL;HARD;3;CRITICAL: cn=oradap,cn=driverset1,o=system is stopped!
[1464121333] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1464121763] SERVICE ALERT: marshallpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1464121824] SERVICE ALERT: marshallpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1464121883] SERVICE ALERT: marshallpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1464121933] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464122120] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 210.21 Gb (96%) - free 9.69 Gb (4%)
[1464122483] SERVICE ALERT: marshallpk;Device Status;OK;HARD;3;Clearing paper path
[1464122483] SERVICE FLAPPING ALERT: marshallpk;Device Status;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1464122493] SERVICE ALERT: marshallpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1464122493] SERVICE FLAPPING ALERT: marshallpk;Paper Tray Status;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1464122533] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 81% (15 min average)
[1464123133] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464123253] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1464123373] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1464123381] SERVICE ALERT: idm1.emich.edu;OfficeAD Driver;CRITICAL;SOFT;1;CRITICAL: cn=OfficeAD,cn=driverset1,o=system is stopped!
[1464123500] SERVICE ALERT: idm1.emich.edu;OfficeAD Driver;CRITICAL;SOFT;2;CRITICAL: cn=OfficeAD,cn=driverset1,o=system is stopped!
[1464123621] SERVICE ALERT: idm1.emich.edu;OfficeAD Driver;CRITICAL;HARD;3;CRITICAL: cn=OfficeAD,cn=driverset1,o=system is stopped!
[1464123663] SERVICE ALERT: idm1.emich.edu;Loopback Driver;CRITICAL;SOFT;1;CRITICAL: cn=Generic Loopback,cn=driverset1,o=system is stopped!
[1464123784] SERVICE ALERT: idm1.emich.edu;Loopback Driver;CRITICAL;SOFT;2;CRITICAL: cn=Generic Loopback,cn=driverset1,o=system is stopped!
[1464123903] SERVICE ALERT: idm1.emich.edu;Loopback Driver;CRITICAL;HARD;3;CRITICAL: cn=Generic Loopback,cn=driverset1,o=system is stopped!
[1464123920] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 207.67 Gb (94%) - free 12.23 Gb (6%)
[1464124203] SERVICE ALERT: AD-RMAD;E:\ Drive Space;CRITICAL;HARD;3;e:\ - total: 599.87 Gb - used: 543.99 Gb (91%) - free 55.88 Gb (9%)
[1464124294] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1464124354] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1464124413] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1464124821] SERVICE ALERT: idm1.emich.edu;OfficeAD Driver;OK;HARD;3;OK: cn=OfficeAD,cn=driverset1,o=system is running.
[1464124821] SERVICE FLAPPING ALERT: idm1.emich.edu;OfficeAD Driver;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1464124873] SERVICE ALERT: idm1.emich.edu;Oracle Driver;OK;HARD;3;OK: cn=oradap,cn=driverset1,o=system is running.
[1464125103] SERVICE ALERT: idm1.emich.edu;Loopback Driver;OK;HARD;3;OK: cn=Generic Loopback,cn=driverset1,o=system is running.
[1464125118] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 211.96 Gb (96%) - free 7.94 Gb (4%)
[1464125613] SERVICE ALERT: marshallpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1464126233] SERVICE DOWNTIME ALERT: idm1.emich.edu;AD Auth Driver;STARTED; Service has entered a period of scheduled downtime
[1464126683] SERVICE ALERT: marshallpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1464126743] SERVICE ALERT: marshallpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1464126803] SERVICE ALERT: marshallpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1464126918] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 207.72 Gb (94%) - free 12.18 Gb (6%)
[1464127078] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 1 Trays Empty
[1464127138] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 1 Trays Empty
[1464127198] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 1 Trays Empty
[1464128173] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1464128227] SERVICE FLAPPING ALERT: idm1.emich.edu;Google Driver;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1464128718] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 211.53 Gb (96%) - free 8.37 Gb (4%)
[1464128771] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464130354] SERVICE ALERT: sillpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1464130413] SERVICE ALERT: sillpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1464130473] SERVICE ALERT: sillpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1464130518] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 207.72 Gb (94%) - free 12.18 Gb (6%)
[1464131118] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 210.64 Gb (96%) - free 9.26 Gb (4%)
[1464132273] SERVICE ALERT: sillpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 48% - OK!
[1464133274] SERVICE FLAPPING ALERT: idm1.emich.edu;Oracle Driver;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1464133518] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 207.72 Gb (94%) - free 12.18 Gb (6%)
[1464134604] SERVICE FLAPPING ALERT: marshallpk;Device Status;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1464134718] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 210.13 Gb (96%) - free 9.77 Gb (4%)
[1464135620] SERVICE FLAPPING ALERT: idm1.emich.edu;OfficeAD Driver;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1464135903] SERVICE FLAPPING ALERT: idm1.emich.edu;Loopback Driver;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1464136413] SERVICE FLAPPING ALERT: marshallpk;Paper Tray Status;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1464137118] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 207.71 Gb (94%) - free 12.19 Gb (6%)
[1464138970] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1464139570] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464140718] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 210.52 Gb (96%) - free 9.38 Gb (4%)
[1464140770] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1464141370] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1464142570] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1464143118] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 207.78 Gb (94%) - free 12.12 Gb (6%)
[1464143770] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464144318] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 210.49 Gb (96%) - free 9.41 Gb (4%)
[1464144370] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1464144970] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464145570] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 69% (15 min average)
[1464146170] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464146290] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464146410] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 80% (15 min average)
[1464146718] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 207.79 Gb (94%) - free 12.11 Gb (6%)
[1464147010] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1464147130] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464147250] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464147918] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 211.33 Gb (96%) - free 8.57 Gb (4%)
[1464148450] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)