[1464322806] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 81% (15 min average)
[1464323406] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464323526] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 73% (15 min average)
[1464324126] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1464324246] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1464324366] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464327101] HOST ALERT: porterpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1464327191] HOST ALERT: porterpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1464327281] HOST ALERT: porterpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1464327311] HOST ALERT: porterpk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1464327401] HOST ALERT: porterpk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1464327491] HOST ALERT: porterpk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1464327581] HOST ALERT: porterpk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1464327671] HOST ALERT: porterpk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1464327701] HOST ALERT: porterpk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1464327791] HOST ALERT: porterpk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1464329166] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1464331566] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464335766] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1464336366] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464336966] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 40% (15 min average)
[1464343234] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;1;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1464343354] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;2;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1464343474] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;HARD;3;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1464343564] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464343684] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464343804] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464347404] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 34% (15 min average)
[1464348004] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464348124] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464348244] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464348844] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 57% (15 min average)
[1464349183] SERVICE ALERT: NETIDAD;Active Directory Federation Services;CRITICAL;SOFT;1;adfssrv: Stopped
[1464349303] SERVICE ALERT: NETIDAD;Active Directory Federation Services;OK;SOFT;2;adfssrv: Started
[1464349444] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1464349564] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 79% (15 min average)
[1464349807] SERVICE ALERT: ADDC4;AD Replication;WARNING;SOFT;1;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1464349926] SERVICE ALERT: ADDC4;AD Replication;WARNING;SOFT;2;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1464349982] SERVICE ALERT: ADDC2;AD Replication;WARNING;SOFT;1;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1464350046] SERVICE ALERT: ADDC4;AD Replication;WARNING;HARD;3;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1464350074] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;OK;HARD;3;OK: No entries found
[1464350101] SERVICE ALERT: ADDC2;AD Replication;WARNING;SOFT;2;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1464350108] SERVICE ALERT: ADDC5;GPO Replication;UNKNOWN;SOFT;1;CHECK_NRPE: Invalid packet type received from server.
[1464350221] SERVICE ALERT: ADDC2;AD Replication;WARNING;HARD;3;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1464350289] SERVICE ALERT: ADDC1;AD Replication;WARNING;SOFT;1;WARNING: Replication with ADDC4-OLD failed!
[1464350407] SERVICE ALERT: ADDC5;GPO Replication;UNKNOWN;SOFT;2;CHECK_NRPE: Invalid packet type received from server.
[1464350410] SERVICE ALERT: ADDC1;AD Replication;WARNING;SOFT;2;WARNING: Replication with ADDC4-OLD failed!
[1464350529] SERVICE ALERT: ADDC1;AD Replication;WARNING;HARD;3;WARNING: Replication with ADDC4-OLD failed!
[1464350530] SERVICE ALERT: ADDC6;AD Replication;WARNING;SOFT;1;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1464350609] SERVICE ALERT: ADDC3;AD Replication;WARNING;SOFT;1;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1464350649] SERVICE ALERT: ADDC6;AD Replication;WARNING;SOFT;2;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1464350706] SERVICE ALERT: ADDC5;GPO Replication;UNKNOWN;HARD;3;CHECK_NRPE: Invalid packet type received from server.
[1464350730] SERVICE ALERT: ADDC3;AD Replication;WARNING;SOFT;2;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1464350762] SERVICE ALERT: ADDC5;AD Replication;WARNING;SOFT;1;WARNING: Replication with ADDC1-OLD failed!
[1464350769] SERVICE ALERT: ADDC6;AD Replication;WARNING;HARD;3;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1464350849] SERVICE ALERT: ADDC3;AD Replication;WARNING;HARD;3;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1464350882] SERVICE ALERT: ADDC5;AD Replication;WARNING;SOFT;2;WARNING: Replication with ADDC1-OLD failed!
[1464350930] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;SOFT;1;WARNING: Replication with ADDC5-OLD failed!
[1464351002] SERVICE ALERT: ADDC5;AD Replication;WARNING;HARD;3;WARNING: Replication with ADDC1-OLD failed!
[1464351050] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;SOFT;2;WARNING: Replication with ADDC5-OLD failed!
[1464351169] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;HARD;3;WARNING: Replication with ADDC5-OLD failed!
[1464351511] SERVICE ALERT: NETIDADPROXY;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1464351631] SERVICE ALERT: NETIDADPROXY;CPU Load;OK;SOFT;2;CPU Load 62% (15 min average)
[1464352107] SERVICE ALERT: TABLEAU;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1464352133] SERVICE ALERT: NETIDAD;Service: DNS Client;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1464352142] SERVICE ALERT: NETIDADPROXY;Service: Group Policy;CRITICAL;SOFT;1;gpsvc: Stopped
[1464352159] SERVICE ALERT: NETIDAD;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.99.247 and port 12489: Connection refused
[1464352164] SERVICE ALERT: NETIDAD;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464352227] SERVICE ALERT: TABLEAU;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1464352253] SERVICE ALERT: NETIDAD;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1464352263] SERVICE ALERT: NETIDADPROXY;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1464352264] SERVICE ALERT: NETIDAD;CPU Load;OK;SOFT;2;CPU Load 2% (15 min average)
[1464352279] SERVICE ALERT: NETIDAD;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1464352301] SERVICE ALERT: NETIDAD;Active Directory Federation Services;CRITICAL;SOFT;1;adfssrv: Starting
[1464352347] SERVICE ALERT: TABLEAU;CPU Load;OK;SOFT;3;CPU Load 0% (15 min average)
[1464352421] SERVICE ALERT: NETIDAD;Active Directory Federation Services;OK;SOFT;2;adfssrv: Started
[1464352452] SERVICE ALERT: NETIDADPROXY;ADFS Web Proxy Service;CRITICAL;SOFT;1;appproxysvc: Stopped
[1464352564] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1464352572] SERVICE ALERT: NETIDADPROXY;ADFS Web Proxy Service;CRITICAL;SOFT;2;appproxysvc: Stopped
[1464352684] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464352692] SERVICE ALERT: NETIDADPROXY;ADFS Web Proxy Service;CRITICAL;HARD;3;appproxysvc: Stopped
[1464352718] SERVICE ALERT: NETIDADPROXY;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1464352804] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 85% (15 min average)
[1464352838] SERVICE ALERT: NETIDADPROXY;Service: Windows Time;CRITICAL;SOFT;2;W32Time: Stopped
[1464352957] SERVICE ALERT: NETIDADPROXY;Service: Windows Time;CRITICAL;HARD;3;W32Time: Stopped
[1464353292] SERVICE ALERT: NETIDADPROXY;ADFS Web Proxy Service;OK;HARD;3;appproxysvc: Started
[1464354146] SERVICE ALERT: TABLEAU;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1464354266] SERVICE ALERT: TABLEAU;CPU Load;OK;SOFT;2;CPU Load 0% (15 min average)
[1464354569] SERVICE ALERT: marshallpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1464354821] SERVICE ALERT: NETIDAD;Active Directory Federation Services;CRITICAL;SOFT;1;adfssrv: Stopped
[1464354941] SERVICE ALERT: NETIDAD;Active Directory Federation Services;OK;SOFT;2;adfssrv: Started
[1464355658] SERVICE ALERT: IMAGES;Service: MDT Monitor;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464355688] HOST ALERT: IMAGES;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1464355693] SERVICE ALERT: IMAGES;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464355694] SERVICE ALERT: IMAGES;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464355695] SERVICE ALERT: IMAGES;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464355698] SERVICE ALERT: IMAGES;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464355705] SERVICE ALERT: IMAGES;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464355706] SERVICE ALERT: IMAGES;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464355707] SERVICE ALERT: IMAGES;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464355707] SERVICE ALERT: IMAGES;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464355708] SERVICE ALERT: IMAGES;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464355709] SERVICE ALERT: IMAGES;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464355709] SERVICE ALERT: IMAGES;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464355711] SERVICE ALERT: IMAGES;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464355718] SERVICE ALERT: IMAGES;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464355723] HOST ALERT: IMAGES;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1464355787] HOST ALERT: IMAGES;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1464355793] SERVICE ALERT: IMAGES;C:\ Drive Space;OK;SOFT;2;c:\ - total: 1157.56 Gb - used: 548.65 Gb (47%) - free 608.92 Gb (53%)
[1464355794] SERVICE ALERT: IMAGES;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1464355795] SERVICE ALERT: IMAGES;Service: Windows Time;CRITICAL;SOFT;2;W32Time: Stopped
[1464355798] SERVICE ALERT: IMAGES;Memory Usage;OK;SOFT;2;Memory usage: total:37586.86 MB - used: 1066.31 MB (3%) - free: 36520.55 MB (97%)
[1464355805] SERVICE ALERT: IMAGES;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1464355806] SERVICE ALERT: IMAGES;CPU Load;OK;SOFT;2;CPU Load 0% (15 min average)
[1464355807] SERVICE ALERT: IMAGES;Service: RPC;OK;SOFT;2;RpcSs: Started
[1464355807] SERVICE ALERT: IMAGES;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1464355808] SERVICE ALERT: IMAGES;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1464355809] SERVICE ALERT: IMAGES;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 1 minute(s)
[1464355809] SERVICE ALERT: IMAGES;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1464355811] SERVICE ALERT: IMAGES;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1464355818] SERVICE ALERT: IMAGES;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1464355915] SERVICE ALERT: IMAGES;Service: Windows Time;OK;SOFT;3;W32Time: Started
[1464355937] SERVICE ALERT: IMAGES;Service: PXE Boot;OK;HARD;3;WDSServer: Started
[1464356055] SERVICE ALERT: IMAGES;PXE Boot Service;OK;HARD;3;WDSServer: Started
[1464356238] SERVICE ALERT: IMAGES;Service: MDT Monitor;UNKNOWN;HARD;3;Failed to open service MDT_Monitor: 1060: The specified service does not exist as an installed service.
[1464356404] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1464356524] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464356644] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 77% (15 min average)
[1464357131] SERVICE ALERT: MBAM;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1464357250] SERVICE ALERT: ADDC4;AD Replication;OK;HARD;3;OK: replication is up and running
[1464357251] SERVICE ALERT: MBAM;CPU Load;OK;SOFT;2;CPU Load 0% (15 min average)
[1464357341] SERVICE ALERT: NETIDAD;Active Directory Federation Services;CRITICAL;SOFT;1;adfssrv: Stopped
[1464357374] SERVICE ALERT: ADDC6;AD Replication;OK;HARD;3;OK: replication is up and running
[1464357426] SERVICE ALERT: ADDC2;AD Replication;OK;HARD;3;OK: replication is up and running
[1464357453] SERVICE ALERT: ADDC3;AD Replication;OK;HARD;3;OK: replication is up and running
[1464357461] SERVICE ALERT: NETIDAD;Active Directory Federation Services;OK;SOFT;2;adfssrv: Started
[1464357606] SERVICE ALERT: ADDC5;AD Replication;OK;HARD;3;OK: replication is up and running
[1464357722] SERVICE ALERT: ADDC1;AD Replication;OK;HARD;3;OK: replication is up and running
[1464357821] SERVICE ALERT: ADDC5;GPO Replication;OK;HARD;3;OK! All GPOs are properly replicated.
[1464358444] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464358564] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 74% (15 min average)
[1464358600] SERVICE ALERT: SNAP;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1464358603] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.40)
[1464358652] HOST ALERT: SNAP;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1464358690] SERVICE ALERT: SNAP;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464358702] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.40)
[1464358711] SERVICE ALERT: SNAP;NSClient++ Version;CRITICAL;SOFT;2;connect to address 164.76.102.40 and port 12489: No route to host
[1464358714] HOST ALERT: SNAP;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.40)
[1464358785] HOST ALERT: SNAP;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.102.40)
[1464358790] SERVICE ALERT: SNAP;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1464358794] HOST ALERT: SNAP;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.44 ms
[1464358794] HOST FLAPPING ALERT: SNAP;STARTED; Host appears to have started flapping (22.2% change > 20.0% threshold)
[1464358828] SERVICE ALERT: SNAP;NSClient++ Version;OK;SOFT;3;NSClient++ 0.4.4.15 2015-11-25
[1464359402] HOST FLAPPING ALERT: SNAP;STOPPED; Host appears to have stopped flapping (4.3% change < 5.0% threshold)
[1464359861] SERVICE ALERT: NETIDAD;Active Directory Federation Services;CRITICAL;SOFT;1;adfssrv: Stopped
[1464359981] SERVICE ALERT: NETIDAD;Active Directory Federation Services;OK;SOFT;2;adfssrv: Started
[1464360964] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1464361375] SERVICE ALERT: ADRODC-LIV;AD Replication;OK;HARD;3;OK: replication is up and running
[1464361848] SERVICE ALERT: VCENTERSSO;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464361848] SERVICE ALERT: VC;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464361848] SERVICE ALERT: VC;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464361878] HOST ALERT: VCENTERSSO;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1464361878] HOST ALERT: VC;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1464361908] SERVICE ALERT: VC;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464361933] SERVICE ALERT: VCENTERSSO;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464361934] SERVICE ALERT: VCENTERSSO;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464361938] HOST ALERT: VC;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1464361963] HOST ALERT: VCENTERSSO;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1464361967] SERVICE ALERT: VCENTERSSO;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464361968] SERVICE ALERT: VC;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464361968] SERVICE ALERT: VC;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464361998] HOST ALERT: VCENTERSSO;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1464361998] HOST ALERT: VC;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1464362018] SERVICE ALERT: VC;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464362019] SERVICE ALERT: VCENTERSSO;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464362028] SERVICE ALERT: VC;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464362048] HOST ALERT: VC;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1464362049] HOST ALERT: VCENTERSSO;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1464362053] SERVICE ALERT: VCENTERSSO;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464362054] SERVICE ALERT: VCENTERSSO;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464362055] SERVICE ALERT: VCENTERSSO;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464362083] HOST ALERT: VCENTERSSO;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1464362088] SERVICE ALERT: VCENTERSSO;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464362088] SERVICE ALERT: VC;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464362088] SERVICE ALERT: VC;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464362103] SERVICE ALERT: VC;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464362105] SERVICE ALERT: VCENTERSSO;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464362118] HOST ALERT: VCENTERSSO;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1464362118] HOST ALERT: VC;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1464362138] SERVICE ALERT: VC;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464362139] SERVICE ALERT: VCENTERSSO;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464362148] SERVICE ALERT: VC;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464362152] SERVICE ALERT: VC;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464362164] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464362168] HOST ALERT: VC;DOWN;SOFT;6;(Host check timed out after 30.00 seconds)
[1464362169] HOST ALERT: VCENTERSSO;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1464362173] SERVICE ALERT: VCENTERSSO;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464362174] SERVICE ALERT: VCENTERSSO;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464362175] SERVICE ALERT: VC;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464362175] SERVICE ALERT: VCENTERSSO;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464362186] SERVICE ALERT: VCENTERSSO;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464362195] HOST ALERT: vspheredb;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.22)
[1464362203] HOST ALERT: VCENTERSSO;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1464362205] HOST ALERT: VC;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1464362223] SERVICE ALERT: VC;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464362225] SERVICE ALERT: VCENTERSSO;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464362253] HOST ALERT: VC;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1464362255] HOST ALERT: VCENTERSSO;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1464362256] SERVICE ALERT: VC;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464362258] HOST ALERT: vspheredb;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.22)
[1464362259] SERVICE ALERT: VCENTERSSO;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464362270] SERVICE ALERT: VCENTERSSO;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464362271] SERVICE ALERT: VC;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464362272] SERVICE ALERT: VC;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464362276] SERVICE ALERT: VCENTERSSO;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464362284] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464362286] HOST ALERT: VC;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1464362289] HOST ALERT: VCENTERSSO;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1464362295] SERVICE ALERT: VC;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464362295] SERVICE ALERT: VCENTERSSO;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464362305] SERVICE ALERT: VC;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464362306] SERVICE ALERT: VCENTERSSO;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464362320] SERVICE ALERT: VCENTERSSO;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464362321] HOST ALERT: vspheredb;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.102.22)
[1464362325] HOST ALERT: VC;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1464362331] SERVICE ALERT: VCENTERSSO;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464362332] SERVICE ALERT: VC;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464362332] SERVICE ALERT: VCENTERSSO;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464362343] SERVICE ALERT: VC;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464362345] SERVICE ALERT: VCENTERSSO;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464362358] SERVICE ALERT: VC;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464362360] SERVICE ALERT: VCENTERSSO;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464362362] SERVICE ALERT: VC;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464362384] HOST ALERT: vspheredb;DOWN;SOFT;4;CRITICAL - Host Unreachable (164.76.102.22)
[1464362390] SERVICE ALERT: VCENTERSSO;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464362391] SERVICE ALERT: VC;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464362392] SERVICE ALERT: VCENTERSSO;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464362392] SERVICE ALERT: VC;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464362395] SERVICE ALERT: VC;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464362395] SERVICE ALERT: VC;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464362396] SERVICE ALERT: VCENTERSSO;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464362404] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464362415] SERVICE ALERT: VC;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464362425] SERVICE ALERT: VC;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464362426] SERVICE ALERT: VCENTERSSO;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464362440] SERVICE ALERT: VCENTERSSO;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464362447] HOST ALERT: vspheredb;DOWN;SOFT;5;CRITICAL - Host Unreachable (164.76.102.22)
[1464362451] SERVICE ALERT: VCENTERSSO;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464362452] SERVICE ALERT: VC;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464362452] SERVICE ALERT: VCENTERSSO;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464362477] SERVICE ALERT: VC;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464362480] SERVICE ALERT: VCENTERSSO;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464362482] SERVICE ALERT: VC;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464362510] HOST ALERT: vspheredb;DOWN;SOFT;6;CRITICAL - Host Unreachable (164.76.102.22)
[1464362510] SERVICE ALERT: VCENTERSSO;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464362511] SERVICE ALERT: VC;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464362512] SERVICE ALERT: VCENTERSSO;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464362515] SERVICE ALERT: VC;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464362515] SERVICE ALERT: VC;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464362516] SERVICE ALERT: VCENTERSSO;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464362545] SERVICE ALERT: VC;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464362560] SERVICE ALERT: VCENTERSSO;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464362571] SERVICE ALERT: VCENTERSSO;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464362572] SERVICE ALERT: VC;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464362572] SERVICE ALERT: VCENTERSSO;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464362573] HOST ALERT: vspheredb;DOWN;SOFT;7;CRITICAL - Host Unreachable (164.76.102.22)
[1464362596] SERVICE ALERT: VC;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464362600] SERVICE ALERT: VCENTERSSO;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464362602] SERVICE ALERT: VC;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464362632] SERVICE ALERT: VCENTERSSO;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464362635] SERVICE ALERT: VC;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464362635] SERVICE ALERT: VC;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464362636] HOST ALERT: vspheredb;DOWN;SOFT;8;CRITICAL - Host Unreachable (164.76.102.22)
[1464362699] HOST ALERT: vspheredb;DOWN;SOFT;9;CRITICAL - Host Unreachable (164.76.102.22)
[1464362762] HOST ALERT: vspheredb;DOWN;HARD;10;CRITICAL - Host Unreachable (164.76.102.22)
[1464362837] HOST DOWNTIME ALERT: VC;STARTED; Host has entered a period of scheduled downtime
[1464362863] HOST DOWNTIME ALERT: VUM;STARTED; Host has entered a period of scheduled downtime
[1464362884] HOST DOWNTIME ALERT: VCENTERSSO;STARTED; Host has entered a period of scheduled downtime
[1464363604] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 74% (15 min average)
[1464364885] HOST ALERT: porterpk;UP;HARD;10;PING OK - Packet loss = 72%, RTA = 0.66 ms
[1464365258] SERVICE ALERT: porterpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 44% - OK!
[1464365414] HOST ALERT: porterpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1464365498] HOST ALERT: porterpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1464365566] HOST ALERT: porterpk;UP;SOFT;3;PING OK - Packet loss = 44%, RTA = 2.19 ms
[1464365710] HOST ALERT: porterpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1464365710] HOST FLAPPING ALERT: porterpk;STARTED; Host appears to have started flapping (22.9% change > 20.0% threshold)
[1464365774] HOST ALERT: porterpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.62 ms
[1464365864] SERVICE ALERT: porterpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1464365924] SERVICE ALERT: porterpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1464365984] SERVICE ALERT: porterpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1464366578] SERVICE ALERT: porterpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 44% - OK!
[1464366602] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1464366698] HOST ALERT: porterpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1464366722] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 72% (15 min average)
[1464366762] HOST ALERT: porterpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.55 ms
[1464367184] SERVICE ALERT: porterpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1464367244] SERVICE ALERT: porterpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1464367304] SERVICE ALERT: porterpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1464367304] SERVICE FLAPPING ALERT: porterpk;Black Print Cartridge HP C8543X;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1464368084] HOST FLAPPING ALERT: porterpk;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1464370322] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1464370442] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1464370562] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464372362] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1464372362] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1464373562] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1464373682] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1464373802] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 91% (15 min average)
[1464374402] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1464374522] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1464374642] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1464376442] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 45% (15 min average)
[1464377581] SERVICE ALERT: ADAUTHDC1;AD Replication;UNKNOWN;SOFT;1;CHECK_NRPE: Invalid packet type received from server.
[1464377642] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1464377699] SERVICE ALERT: ADAUTHDC1;AD Replication;UNKNOWN;SOFT;2;CHECK_NRPE: Invalid packet type received from server.
[1464377762] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 86% (15 min average)
[1464377830] SERVICE ALERT: ADAUTHDC1;AD Replication;UNKNOWN;HARD;3;CHECK_NRPE: Invalid packet type received from server.
[1464378704] SERVICE FLAPPING ALERT: porterpk;Black Print Cartridge HP C8543X;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1464379019] SERVICE ALERT: ADAUTHDC1;AD Replication;WARNING;HARD;3;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1464379623] SERVICE ALERT: ADAUTHDC1;AD Replication;OK;HARD;3;OK: replication is up and running
[1464380137] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 621 found since last run. oneLevelSearchOps OK. searchOps is too high! 622 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464380257] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 625 found since last run. oneLevelSearchOps OK. searchOps is too high! 626 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464380377] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464381962] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1464382082] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 63% (15 min average)
[1464383282] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464383377] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 949 found since last run. oneLevelSearchOps OK. searchOps is too high! 950 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464383402] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1464383497] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 573 found since last run. oneLevelSearchOps OK. searchOps is too high! 574 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464383522] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 73% (15 min average)
[1464383617] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464384722] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464384842] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464384962] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464385048] SERVICE ALERT: INTLAPP-ISSMAP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1464385168] SERVICE ALERT: INTLAPP-ISSMAP;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464385288] SERVICE ALERT: INTLAPP-ISSMAP;CPU Load;OK;SOFT;3;CPU Load 1% (15 min average)
[1464389411] HOST ALERT: porterpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1464389497] HOST ALERT: porterpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1464389534] HOST ALERT: porterpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1464389624] HOST ALERT: porterpk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1464389707] HOST ALERT: porterpk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1464389797] HOST ALERT: porterpk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1464389887] HOST ALERT: porterpk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1464389977] HOST ALERT: porterpk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1464390067] HOST ALERT: porterpk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1464390097] HOST ALERT: porterpk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1464390362] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1464390962] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1464391562] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464391682] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464391802] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 94% (15 min average)
[1464392402] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464392522] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464392642] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1464393242] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 59% (15 min average)
[1464393842] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464393962] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1464394082] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1464398044] SERVICE ALERT: TARGETX-2012;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1464398164] SERVICE ALERT: TARGETX-2012;CPU Load;OK;SOFT;2;CPU Load 61% (15 min average)
[1464400078] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 82% (15 min average)
[1464400240] HOST DOWNTIME ALERT: idm1.emich.edu;STARTED; Host has entered a period of scheduled downtime
[1464400408] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 12
[1464400417] SERVICE ALERT: idm1.emich.edu;LDAP Stats;CRITICAL;SOFT;1;CRITICAL: Failed to connect to LDAP!
[1464400458] SERVICE ALERT: idm1.emich.edu;OfficeAD Driver;CRITICAL;HARD;3;Unknown: Error reported?
[1464400497] SERVICE ALERT: idm1.emich.edu;Loopback Driver;CRITICAL;SOFT;1;Unknown: Error reported?
[1464400528] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 12
[1464400537] SERVICE ALERT: idm1.emich.edu;LDAP Stats;CRITICAL;SOFT;2;CRITICAL: Failed to connect to LDAP!
[1464400616] SERVICE ALERT: idm2.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 12
[1464400618] SERVICE ALERT: idm1.emich.edu;Loopback Driver;CRITICAL;SOFT;2;Unknown: Error reported?
[1464400648] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 12
[1464400657] SERVICE ALERT: idm1.emich.edu;LDAP Stats;CRITICAL;HARD;3;CRITICAL: Failed to connect to LDAP!
[1464400658] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 4
[1464400678] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464400710] SERVICE ALERT: idm1.emich.edu;NSDS Errors;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1464400736] SERVICE ALERT: idm2.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 11
[1464400778] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 11
[1464400796] SERVICE ALERT: idm1.emich.edu;Loopback Driver;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 60 seconds.
[1464400798] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464400826] HOST ALERT: idm1.emich.edu;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1464400838] SERVICE ALERT: idm1.emich.edu;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1464400856] SERVICE ALERT: idm2.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 11
[1464400868] HOST ALERT: idm1.emich.edu;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1464400890] SERVICE ALERT: idm1.emich.edu;NSDS Errors;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 60 seconds.
[1464400899] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 11
[1464400910] HOST ALERT: idm1.emich.edu;UP;SOFT;3;PING OK - Packet loss = 76%, RTA = 0.29 ms
[1464400912] SERVICE ALERT: idm1.emich.edu;Oracle Driver;CRITICAL;SOFT;1;Unknown: Error reported?
[1464400916] SERVICE ALERT: idm1.emich.edu;AD Driver;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 60 seconds.
[1464400918] SERVICE ALERT: idm1.emich.edu;Google Driver;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 60 seconds.
[1464400918] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464400952] SERVICE ALERT: idm1.emich.edu;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.36 ms
[1464400976] SERVICE ALERT: idm1.emich.edu;AD Driver;CRITICAL;SOFT;2;Unknown: Error reported?
[1464400978] SERVICE ALERT: idm1.emich.edu;Google Driver;CRITICAL;SOFT;2;Unknown: Error reported?
[1464400984] SERVICE ALERT: idm1.emich.edu;Oracle Driver;CRITICAL;SOFT;2;Unknown: Error reported?
[1464401010] SERVICE ALERT: idm1.emich.edu;NSDS Errors;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 60 seconds.
[1464401097] SERVICE ALERT: idm1.emich.edu;AD Driver;CRITICAL;HARD;3;CRITICAL: cn=Campus AD,cn=driverset1,o=system is stopped!
[1464401099] SERVICE ALERT: idm1.emich.edu;Google Driver;CRITICAL;HARD;3;CRITICAL: cn=gdap,cn=driverset1,o=system is stopped!
[1464401105] SERVICE ALERT: idm1.emich.edu;Oracle Driver;CRITICAL;HARD;3;CRITICAL: cn=oradap,cn=driverset1,o=system is stopped!
[1464401248] SERVICE ALERT: idm4.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1464401258] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464401336] SERVICE ALERT: idm1.emich.edu;Loopback Driver;OK;HARD;3;OK: cn=Generic Loopback,cn=driverset1,o=system is running.
[1464401456] SERVICE ALERT: idm2.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1464401498] SERVICE ALERT: idm3.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1464401518] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1464401550] SERVICE ALERT: idm1.emich.edu;NSDS Errors;OK;HARD;3;OK: 0 Errors Reported
[1464401657] SERVICE ALERT: idm1.emich.edu;OfficeAD Driver;OK;HARD;3;OK: cn=OfficeAD,cn=driverset1,o=system is running.
[1464401696] SERVICE ALERT: idm1.emich.edu;AD Driver;OK;HARD;3;OK: cn=Campus AD,cn=driverset1,o=system is running.
[1464401698] SERVICE ALERT: idm1.emich.edu;Google Driver;OK;HARD;3;OK: cn=gdap,cn=driverset1,o=system is running.
[1464401704] SERVICE ALERT: idm1.emich.edu;Oracle Driver;OK;HARD;3;OK: cn=oradap,cn=driverset1,o=system is running.
[1464402118] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464402718] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1464403318] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 82% (15 min average)
[1464403829] HOST DOWNTIME ALERT: idm1.emich.edu;STOPPED; Host has exited from a period of scheduled downtime
[1464403918] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464404038] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464404158] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464407158] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 82% (15 min average)
[1464407758] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464407878] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1464407998] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)