[1460001802] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460001922] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460002042] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460003401] HOST DOWNTIME ALERT: OPIE;STOPPED; Host has exited from a period of scheduled downtime
[1460003401] HOST DOWNTIME ALERT: ADWU1;STOPPED; Host has exited from a period of scheduled downtime
[1460003401] HOST DOWNTIME ALERT: PNCWEB-2012;STOPPED; Host has exited from a period of scheduled downtime
[1460003401] HOST DOWNTIME ALERT: PNCAPP-2012;STOPPED; Host has exited from a period of scheduled downtime
[1460003401] HOST DOWNTIME ALERT: WINKMS;STOPPED; Host has exited from a period of scheduled downtime
[1460003401] HOST DOWNTIME ALERT: PROCARE;STOPPED; Host has exited from a period of scheduled downtime
[1460003401] HOST DOWNTIME ALERT: OFFICEKMS;STOPPED; Host has exited from a period of scheduled downtime
[1460003401] HOST DOWNTIME ALERT: COE-WEBDAV;STOPPED; Host has exited from a period of scheduled downtime
[1460003401] HOST DOWNTIME ALERT: NETENG-PA;STOPPED; Host has exited from a period of scheduled downtime
[1460003401] HOST DOWNTIME ALERT: ADAUTH-RMAD;STOPPED; Host has exited from a period of scheduled downtime
[1460003401] HOST DOWNTIME ALERT: NESSUS;STOPPED; Host has exited from a period of scheduled downtime
[1460003401] HOST DOWNTIME ALERT: PNCDB-2012;STOPPED; Host has exited from a period of scheduled downtime
[1460003401] HOST DOWNTIME ALERT: TI22-PSYCLINIC;STOPPED; Host has exited from a period of scheduled downtime
[1460003401] HOST DOWNTIME ALERT: TI22-COE;STOPPED; Host has exited from a period of scheduled downtime
[1460003401] HOST DOWNTIME ALERT: REFWIN;STOPPED; Host has exited from a period of scheduled downtime
[1460003401] HOST DOWNTIME ALERT: TI22;STOPPED; Host has exited from a period of scheduled downtime
[1460003401] HOST DOWNTIME ALERT: ADWSUSCLIENTS;STOPPED; Host has exited from a period of scheduled downtime
[1460003401] HOST DOWNTIME ALERT: SANDBOX-DB;STOPPED; Host has exited from a period of scheduled downtime
[1460003401] HOST DOWNTIME ALERT: DIRECTACCESS;STOPPED; Host has exited from a period of scheduled downtime
[1460003401] HOST DOWNTIME ALERT: CBORD;STOPPED; Host has exited from a period of scheduled downtime
[1460003401] HOST DOWNTIME ALERT: IMAGES;STOPPED; Host has exited from a period of scheduled downtime
[1460003401] HOST DOWNTIME ALERT: LANSWEEPER;STOPPED; Host has exited from a period of scheduled downtime
[1460003401] HOST DOWNTIME ALERT: HDBACKUPS;STOPPED; Host has exited from a period of scheduled downtime
[1460003401] HOST DOWNTIME ALERT: UHSPHARM-2012;STOPPED; Host has exited from a period of scheduled downtime
[1460003401] HOST DOWNTIME ALERT: TABLEAU;STOPPED; Host has exited from a period of scheduled downtime
[1460003401] HOST DOWNTIME ALERT: MBAMDB;STOPPED; Host has exited from a period of scheduled downtime
[1460003401] HOST DOWNTIME ALERT: XDM;STOPPED; Host has exited from a period of scheduled downtime
[1460003401] HOST DOWNTIME ALERT: MBAM;STOPPED; Host has exited from a period of scheduled downtime
[1460003401] HOST DOWNTIME ALERT: CEDARON;STOPPED; Host has exited from a period of scheduled downtime
[1460003401] HOST DOWNTIME ALERT: DPS-SECURE02;STOPPED; Host has exited from a period of scheduled downtime
[1460003401] HOST DOWNTIME ALERT: Echo-01;STOPPED; Host has exited from a period of scheduled downtime
[1460003401] HOST DOWNTIME ALERT: EMUSQLEXPRESS;STOPPED; Host has exited from a period of scheduled downtime
[1460003401] HOST DOWNTIME ALERT: EMUSCRIPT;STOPPED; Host has exited from a period of scheduled downtime
[1460003401] HOST DOWNTIME ALERT: FTK;STOPPED; Host has exited from a period of scheduled downtime
[1460003401] HOST DOWNTIME ALERT: EMU-VALUE01;STOPPED; Host has exited from a period of scheduled downtime
[1460003401] HOST DOWNTIME ALERT: AD-RMAD;STOPPED; Host has exited from a period of scheduled downtime
[1460003842] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1460004442] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1460004442] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1460008149] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1460008269] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 95% (5 min average)
[1460008389] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 96% (5 min average)
[1460009589] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1460009589] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1460010189] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460010302] SERVICE ALERT: ADCA2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460010422] SERVICE ALERT: ADCA2;CPU Load;OK;SOFT;2;CPU Load 1% (5 min average)
[1460010756] SERVICE ALERT: SANDBOX-APP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460010876] SERVICE ALERT: SANDBOX-APP;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1460010996] SERVICE ALERT: SANDBOX-APP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1460011759] SERVICE ALERT: INTLDB;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460011879] SERVICE ALERT: INTLDB;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1460012196] SERVICE ALERT: SANDBOX-APP;CPU Load;OK;HARD;3;CPU Load 5% (5 min average)
[1460012589] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 92% (5 min average)
[1460013189] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1460014989] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 92% (5 min average)
[1460015589] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 95% (5 min average)
[1460015842] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1460016189] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 87% (5 min average)
[1460018242] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 65% (5 min average)
[1460020042] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1460020162] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 74% (5 min average)
[1460023162] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1460023282] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460023401] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460024600] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 83% (5 min average)
[1460025188] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460025307] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 50% (5 min average)
[1460025800] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460025919] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460026039] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1460026039] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1460026639] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1460027705] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1460030239] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 71% (5 min average)
[1460030839] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460030959] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 97% (5 min average)
[1460031079] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460031679] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 62% (5 min average)
[1460032279] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460032368] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1460032399] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 92% (5 min average)
[1460032422] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;OK;SOFT;2;Black Cartridge HP CE260X is at 31% - OK!
[1460032519] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 60% (5 min average)
[1460034038] SERVICE ALERT: markjeffersonpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1460034068] HOST ALERT: markjeffersonpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460034093] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! 2 of 3 Trays Still Full
[1460034097] HOST ALERT: markjeffersonpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.55 ms
[1460034319] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460034439] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460034559] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 70% (5 min average)
[1460035159] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1460035279] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 51% (5 min average)
[1460035665] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460035726] SERVICE ALERT: markjeffersonpk;Device Status;OK;SOFT;2;Clearing paper path
[1460036479] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460036599] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1460036719] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460037302] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (34.1% change >= 20.0% threshold)
[1460037302] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1460037302] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1460038502] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (28.7% change >= 20.0% threshold)
[1460038502] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1460038502] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1460038519] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 81% (5 min average)
[1460038726] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (34.1% change >= 20.0% threshold)
[1460038726] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1460038726] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1460039119] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460039239] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460039333] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1460039333] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1460039334] SERVICE ALERT: marshallpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1460039337] SERVICE ALERT: strongpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1460039338] HOST ALERT: marshallpk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.61 ms
[1460039341] HOST ALERT: strongpk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.53 ms
[1460039359] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460039359] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (35.9% change >= 20.0% threshold)
[1460039408] SERVICE ALERT: owen101lab3;Printer Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1460039426] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (35.9% change >= 20.0% threshold)
[1460039426] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1460039426] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1460039467] SERVICE ALERT: owen101lab3;Printer Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1460039484] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (35.9% change >= 20.0% threshold)
[1460039484] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1460039484] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1460039513] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (35.9% change >= 20.0% threshold)
[1460039513] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1460039513] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1460039527] SERVICE ALERT: owen101lab3;Printer Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1460039548] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (35.9% change >= 20.0% threshold)
[1460039548] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1460039548] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1460039591] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (35.9% change >= 20.0% threshold)
[1460039591] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1460039591] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1460039659] SERVICE ALERT: marshallpk;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460039719] SERVICE ALERT: marshallpk;Printer Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460039752] SERVICE ALERT: hallefirst3;Maintenance Kit HP 110V-C9152A;CRITICAL;SOFT;1;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 0% - CRITICAL!
[1460039779] SERVICE ALERT: marshallpk;Printer Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460039805] SERVICE ALERT: hallethirdsouth;Printer Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.107.213. Is the printer off?  Is the network down?
[1460039812] SERVICE ALERT: hallefirst3;Maintenance Kit HP 110V-C9152A;CRITICAL;SOFT;2;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 0% - CRITICAL!
[1460039865] SERVICE ALERT: hallethirdsouth;Printer Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.107.213. Is the printer off?  Is the network down?
[1460039873] SERVICE ALERT: hallefirst3;Maintenance Kit HP 110V-C9152A;CRITICAL;HARD;3;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 0% - CRITICAL!
[1460039925] SERVICE ALERT: hallethirdsouth;Printer Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.107.213. Is the printer off?  Is the network down?
[1460040025] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (31.3% change >= 20.0% threshold)
[1460040025] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1460040025] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1460040060] HOST ALERT: prayharrold205lab;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460040124] SERVICE ALERT: prayharrold205lab;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.205.183. Is the printer off?  Is the network down?
[1460040150] HOST ALERT: prayharrold205lab;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460040171] SERVICE ALERT: prayharrold205lab;Printer Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.205.183. Is the printer off?  Is the network down?
[1460040184] SERVICE ALERT: prayharrold205lab;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.205.183. Is the printer off?  Is the network down?
[1460040201] HOST ALERT: prayharrold205lab;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460040231] SERVICE ALERT: prayharrold205lab;Printer Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.205.183. Is the printer off?  Is the network down?
[1460040244] SERVICE ALERT: prayharrold205lab;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.205.183. Is the printer off?  Is the network down?
[1460040261] HOST ALERT: prayharrold205lab;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460040291] SERVICE ALERT: prayharrold205lab;Printer Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.205.183. Is the printer off?  Is the network down?
[1460040305] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 94% (5 min average)
[1460040321] HOST ALERT: prayharrold205lab;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460040379] SERVICE ALERT: marshallpk;Printer Status;OK;HARD;3;Processing job from tray 3
[1460040411] HOST ALERT: prayharrold205lab;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1460040425] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 67% (5 min average)
[1460040501] HOST ALERT: prayharrold205lab;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1460040591] HOST ALERT: prayharrold205lab;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1460040678] SERVICE ALERT: markjeffersonpk;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1460040681] HOST ALERT: prayharrold205lab;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1460040698] SERVICE ALERT: markjeffersonpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1460040708] HOST ALERT: markjeffersonpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460040725] SERVICE ALERT: markjeffersonpk;Printer Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1460040735] HOST ALERT: markjeffersonpk;UP;SOFT;2;PING OK - Packet loss = 54%, RTA = 0.69 ms
[1460040740] SERVICE ALERT: markjeffersonpk;Maintenance Kit HP 110V-C9152A;OK;SOFT;2;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 57% - OK!
[1460040751] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! 1 of 3 Trays Still Full
[1460040771] HOST ALERT: prayharrold205lab;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1460040779] SERVICE ALERT: markjeffersonpk;Printer Status;OK;SOFT;2;TRAY 2 EMPTY ANY TYPE LETTER   Processing job
[1460041075] SERVICE ALERT: BOEAPPTEST;CPU Load;WARNING;SOFT;1;CPU Load 94% (5 min average)
[1460041195] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;SOFT;2;CPU Load 95% (5 min average)
[1460041315] SERVICE ALERT: BOEAPPTEST;CPU Load;OK;SOFT;3;CPU Load 5% (5 min average)
[1460041379] SERVICE ALERT: markjeffersonpk;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460041439] SERVICE ALERT: markjeffersonpk;Printer Status;OK;SOFT;2;Open right  door.
[1460041713] HOST ALERT: markjeffersonpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460041777] HOST ALERT: markjeffersonpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.54 ms
[1460041777] HOST FLAPPING ALERT: markjeffersonpk;STARTED; Host appears to have started flapping (20.9% change > 20.0% threshold)
[1460042359] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 72% (5 min average)
[1460042640] SERVICE ALERT: markjeffersonpk;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460042699] SERVICE ALERT: markjeffersonpk;Printer Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460042759] SERVICE ALERT: markjeffersonpk;Printer Status;OK;SOFT;3;Clearing paper path
[1460042934] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1460042959] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460042993] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1460043002] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1460043053] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1460043062] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1460043079] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460043121] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1460043152] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1460043197] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460043212] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1460043212] HOST FLAPPING ALERT: markjeffersonpk;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1460043272] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1460043722] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1460043792] SERVICE ALERT: hallegroundcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1460043959] SERVICE ALERT: markjeffersonpk;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460044019] SERVICE ALERT: markjeffersonpk;Printer Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460044079] SERVICE ALERT: markjeffersonpk;Printer Status;OK;SOFT;3;CLOSE RIGHT DOOR  46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1460044315] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1460044321] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1460044371] SERVICE ALERT: hallefirst1;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1460044382] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1460044401] HOST ALERT: hallefirst1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460044418] SERVICE ALERT: hallefirst1;Printer Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1460044431] SERVICE ALERT: hallefirst1;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1460044435] SERVICE ALERT: BOEAPPTEST;CPU Load;OK;SOFT;2;CPU Load 76% (5 min average)
[1460044442] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1460044448] HOST ALERT: hallefirst1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460044472] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1460044478] SERVICE ALERT: hallefirst1;Printer Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1460044491] SERVICE ALERT: hallefirst1;Maintenance Kit HP 110V-C9152A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1460044508] HOST ALERT: hallefirst1;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460044538] SERVICE ALERT: hallefirst1;Printer Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1460044568] HOST ALERT: hallefirst1;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460044632] HOST ALERT: hallefirst1;UP;SOFT;5;PING OK - Packet loss = 0%, RTA = 0.57 ms
[1460044834] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1460044893] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1460044922] SERVICE ALERT: hallefirst3;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1460044953] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1460044966] SERVICE ALERT: halleground2;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460044978] SERVICE ALERT: hallefirst3;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! 2 of 3 Trays Still Full
[1460045026] SERVICE ALERT: halleground2;Printer Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460045085] SERVICE ALERT: hallefirst1;Maintenance Kit HP 110V-C9152A;OK;HARD;3;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 45% - OK!
[1460045086] SERVICE ALERT: halleground2;Printer Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460045095] HOST ALERT: prayharrold205lab;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.57 ms
[1460045132] SERVICE ALERT: hallefirst1;Printer Status;OK;HARD;3;Ready
[1460045336] SERVICE ALERT: strongpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1460045452] SERVICE ALERT: marshallpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1460045547] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 34% - OK!
[1460045639] SERVICE ALERT: prayharrold205lab;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1460045685] SERVICE ALERT: prayharrold205lab;Printer Status;OK;HARD;3;Clearing paper path
[1460045705] SERVICE ALERT: hallefirst2;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460045765] SERVICE ALERT: hallefirst2;Printer Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460045825] SERVICE ALERT: hallefirst2;Printer Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460046241] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1460046241] SERVICE FLAPPING ALERT: alexanderpk;Paper Tray Status;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1460046244] SERVICE ALERT: prayharrold205lab;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.205.183. Is the printer off?  Is the network down?
[1460046274] HOST ALERT: prayharrold205lab;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460046291] SERVICE ALERT: prayharrold205lab;Printer Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.205.183. Is the printer off?  Is the network down?
[1460046304] SERVICE ALERT: prayharrold205lab;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.205.183. Is the printer off?  Is the network down?
[1460046321] HOST ALERT: prayharrold205lab;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460046351] SERVICE ALERT: prayharrold205lab;Printer Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.205.183. Is the printer off?  Is the network down?
[1460046364] SERVICE ALERT: prayharrold205lab;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.205.183. Is the printer off?  Is the network down?
[1460046381] HOST ALERT: prayharrold205lab;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460046411] SERVICE ALERT: prayharrold205lab;Printer Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.205.183. Is the printer off?  Is the network down?
[1460046441] HOST ALERT: prayharrold205lab;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460046531] HOST ALERT: prayharrold205lab;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460046621] HOST ALERT: prayharrold205lab;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1460046711] HOST ALERT: prayharrold205lab;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1460046801] HOST ALERT: prayharrold205lab;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1460046891] HOST ALERT: prayharrold205lab;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1460046981] HOST ALERT: prayharrold205lab;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1460048221] SERVICE ALERT: ADCA2;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1460048225] SERVICE ALERT: hallefirst2;Printer Status;OK;HARD;3;Ready
[1460048341] SERVICE ALERT: ADCA2;CPU Load;OK;SOFT;2;CPU Load 3% (5 min average)
[1460048403] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1460048403] SERVICE FLAPPING ALERT: alexanderpk;Paper Tray Status;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1460048403] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1460048403] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1460048435] SERVICE ALERT: ADRODC-WEMU;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1460048439] HOST ALERT: ADRODC-WEMU;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1460048449] SERVICE ALERT: ADRODC-WEMU;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1460048593] SERVICE ALERT: ADRODC-WEMU;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1460048640] SERVICE ALERT: ADRODC-WEMU;Service: RPC;OK;HARD;3;RpcSs: Started
[1460048661] SERVICE ALERT: ADRODC-WEMU;Service: Windows Time;OK;HARD;3;W32Time: Started
[1460048683] SERVICE ALERT: ADRODC-WEMU;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1460048706] SERVICE ALERT: ADRODC-WEMU;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 20 minute(s)
[1460048755] SERVICE ALERT: ADRODC-WEMU;C:\ Drive Space;OK;HARD;3;c:\ - total: 97.31 Gb - used: 11.66 Gb (12%) - free 85.65 Gb (88%)
[1460048773] SERVICE ALERT: ADRODC-WEMU;Service: WMI;OK;HARD;3;Winmgmt: Started
[1460048801] SERVICE ALERT: ADRODC-WEMU;CPU Load;OK;HARD;3;CPU Load 0% (5 min average)
[1460048825] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 91% (5 min average)
[1460048851] SERVICE ALERT: ADRODC-WEMU;Memory Usage;OK;HARD;3;Memory usage: total:21654.72 MB - used: 831.39 MB (4%) - free: 20823.33 MB (96%)
[1460048895] SERVICE ALERT: ADRODC-WEMU;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1460048941] SERVICE ALERT: ADRODC-WEMU;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1460048945] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 84% (5 min average)
[1460051086] SERVICE ALERT: halleground2;Printer Status;OK;HARD;3;TRAY 4 EMPTY PLAIN LETTER   Ready
[1460052074] SERVICE ALERT: mckennypk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1460052134] SERVICE ALERT: mckennypk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1460052193] SERVICE ALERT: mckennypk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1460052793] SERVICE ALERT: mckennypk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1460053485] SERVICE ALERT: halleground2;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460053545] SERVICE ALERT: halleground2;Printer Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460053605] SERVICE ALERT: halleground2;Printer Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460053917] SERVICE ALERT: eateriespk;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460053977] SERVICE ALERT: eateriespk;Printer Status;OK;SOFT;2;Ready
[1460054476] SERVICE ALERT: INTLDB;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460054596] SERVICE ALERT: INTLDB;CPU Load;OK;SOFT;2;CPU Load 1% (5 min average)
[1460054596] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1460054944] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1460055064] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 90% (5 min average)
[1460055122] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1460055182] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 100% - OK!
[1460055184] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 72% (5 min average)
[1460055196] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 55% (5 min average)
[1460055406] SERVICE ALERT: halleground2;Printer Status;OK;HARD;3;Ready
[1460055406] SERVICE FLAPPING ALERT: halleground2;Printer Status;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1460055796] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460055916] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1460056036] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1460056440] SERVICE FLAPPING ALERT: alexanderpk;Paper Tray Status;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1460058641] SERVICE ALERT: C4SAS-SE;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460058761] SERVICE ALERT: C4SAS-SE;CPU Load;OK;SOFT;2;CPU Load 11% (5 min average)
[1460061330] SERVICE ALERT: hallefirst1;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460061390] SERVICE ALERT: hallefirst1;Printer Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460061436] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 73% (5 min average)
[1460061450] SERVICE ALERT: hallefirst1;Printer Status;OK;SOFT;3;Check top bin. Remove paper.
[1460061784] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1460061904] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 85% (5 min average)
[1460062036] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460062156] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1460062276] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460062287] HOST ALERT: markjeffersonpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460062337] SERVICE ALERT: markjeffersonpk;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1460062346] HOST ALERT: markjeffersonpk;UP;SOFT;2;PING OK - Packet loss = 50%, RTA = 0.54 ms
[1460062391] SERVICE ALERT: markjeffersonpk;Maintenance Kit HP 110V-C9152A;OK;SOFT;2;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 56% - OK!
[1460062396] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 1 Trays Empty
[1460062456] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 1 Trays Empty
[1460062517] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 1 Trays Empty
[1460062876] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1460063116] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1460063452] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1460063476] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1460063511] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1460063572] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1460063716] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 1 Trays Empty
[1460063776] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 1 Trays Empty
[1460063823] SERVICE ALERT: hallefirst2;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460063837] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 1 Trays Empty
[1460063883] SERVICE ALERT: hallefirst2;Printer Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460063944] SERVICE ALERT: hallefirst2;Printer Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460064436] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1460064436] SERVICE FLAPPING ALERT: hallefirstcolor;Paper Tray Status;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1460064543] SERVICE ALERT: hallefirst2;Printer Status;OK;HARD;3;Ready
[1460065100] SERVICE ALERT: TABLEAU;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1460065220] SERVICE ALERT: TABLEAU;CPU Load;OK;SOFT;2;CPU Load 3% (5 min average)
[1460065504] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1460065604] SERVICE FLAPPING ALERT: halleground2;Printer Status;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1460065624] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 83% (5 min average)
[1460065945] HOST ALERT: prayharrold205lab;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.60 ms
[1460066159] SERVICE ALERT: prayharrold205lab;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1460066205] SERVICE ALERT: prayharrold205lab;Printer Status;OK;HARD;3;Warming up
[1460066805] SERVICE ALERT: prayharrold205lab;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460066865] SERVICE ALERT: prayharrold205lab;Printer Status;OK;SOFT;2;Warming up
[1460067499] SERVICE ALERT: owen101lab2;Printer Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1460067559] SERVICE ALERT: owen101lab2;Printer Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1460067619] SERVICE ALERT: owen101lab2;Printer Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1460067726] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460067756] HOST ALERT: owen101lab1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460067786] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460067790] HOST ALERT: owen101lab1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.68 ms
[1460067840] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;OK;SOFT;3;Black Print Cartridge HP C8543X is at 36% - OK!
[1460068213] SERVICE ALERT: owen101lab2;Printer Status;OK;HARD;3;MENUS  RETRIEVE JOB  INFORMATION  PAPER HANDLING
[1460068550] SERVICE ALERT: owen101lab2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1460068604] SERVICE ALERT: owen101lab1;Printer Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460068610] SERVICE ALERT: owen101lab2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1460068634] HOST ALERT: owen101lab1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460068640] HOST ALERT: owen101lab2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460068650] SERVICE ALERT: owen101lab2;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1460068663] HOST ALERT: owen101lab2;UP;SOFT;2;PING OK - Packet loss = 64%, RTA = 0.52 ms
[1460068664] SERVICE ALERT: owen101lab1;Printer Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460068666] SERVICE ALERT: owen101lab2;Paper Tray Status;OK;SOFT;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1460068668] HOST ALERT: owen101lab1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.59 ms
[1460068668] HOST FLAPPING ALERT: owen101lab1;STARTED; Host appears to have started flapping (21.1% change > 20.0% threshold)
[1460068704] SERVICE ALERT: owen101lab2;Maintenance Kit HP 110V-C9152A;OK;SOFT;2;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 59% - OK!
[1460068718] SERVICE ALERT: owen101lab1;Printer Status;OK;SOFT;3;Processing job from tray 4
[1460068993] HOST ALERT: owen101lab2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460069057] HOST ALERT: owen101lab2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.49 ms
[1460069057] HOST FLAPPING ALERT: owen101lab2;STARTED; Host appears to have started flapping (23.3% change > 20.0% threshold)
[1460069164] SERVICE ALERT: prayharrold205lab;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.205.183. Is the printer off?  Is the network down?
[1460069194] HOST ALERT: prayharrold205lab;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460069224] SERVICE ALERT: prayharrold205lab;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.205.183. Is the printer off?  Is the network down?
[1460069254] HOST ALERT: prayharrold205lab;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460069271] SERVICE ALERT: prayharrold205lab;Printer Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.205.183. Is the printer off?  Is the network down?
[1460069284] SERVICE ALERT: prayharrold205lab;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.205.183. Is the printer off?  Is the network down?
[1460069301] HOST ALERT: prayharrold205lab;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460069331] SERVICE ALERT: prayharrold205lab;Printer Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.205.183. Is the printer off?  Is the network down?
[1460069361] HOST ALERT: prayharrold205lab;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460069391] SERVICE ALERT: prayharrold205lab;Printer Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.205.183. Is the printer off?  Is the network down?
[1460069421] HOST ALERT: prayharrold205lab;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460069476] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 61% (5 min average)
[1460069511] HOST ALERT: prayharrold205lab;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1460069601] HOST ALERT: prayharrold205lab;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1460069651] SERVICE ALERT: owen101labcolor;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 1 Trays Empty
[1460069691] HOST ALERT: prayharrold205lab;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1460069712] SERVICE ALERT: owen101labcolor;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! All Trays Still Full
[1460069781] HOST ALERT: prayharrold205lab;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1460069871] HOST ALERT: prayharrold205lab;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1460070076] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1460070196] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460070316] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 96% (5 min average)
[1460070492] HOST FLAPPING ALERT: owen101lab1;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1460070881] HOST FLAPPING ALERT: owen101lab2;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1460071181] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1460071241] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 100% - OK!
[1460071693] SERVICE ALERT: studentcenterpk;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460071753] SERVICE ALERT: studentcenterpk;Printer Status;OK;SOFT;2;50 72 65 73 73 20 17 20 74 6F 20 73 65 65 20 73  74 65 70  by step information
[1460071841] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1460071901] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1460071961] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1460072353] SERVICE ALERT: studentcenterpk;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460072414] SERVICE ALERT: studentcenterpk;Printer Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460072476] SERVICE ALERT: studentcenterpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.150.252. Is the printer off?  Is the network down?
[1460072479] SERVICE ALERT: studentcenterpk;Printer Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.150.252. Is the printer off?  Is the network down?
[1460072531] SERVICE ALERT: studentcenterpk;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! 2 of 3 Trays Still Full
[1460073073] SERVICE ALERT: studentcenterpk;Printer Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460073230] HOST ALERT: halleground1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460073294] HOST ALERT: halleground1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.52 ms
[1460073445] SERVICE ALERT: hallefirstcolor;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460073505] SERVICE ALERT: hallefirstcolor;Printer Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460073565] SERVICE ALERT: hallefirstcolor;Printer Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460073673] SERVICE ALERT: studentcenterpk;Printer Status;OK;HARD;3;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1460074166] SERVICE ALERT: hallefirstcolor;Printer Status;OK;HARD;3;Open tray indicated.
[1460074362] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 43% - OK!
[1460075473] SERVICE ALERT: studentcenterpk;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460075534] SERVICE ALERT: studentcenterpk;Printer Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460075593] SERVICE ALERT: studentcenterpk;Printer Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460075593] SERVICE FLAPPING ALERT: studentcenterpk;Printer Status;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1460075836] SERVICE FLAPPING ALERT: hallefirstcolor;Paper Tray Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1460077331] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1460077391] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1460077451] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1460077545] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;SOFT;1;e:\ - total: 499.87 Gb - used: 401.28 Gb (80%) - free 98.59 Gb (20%)
[1460077654] SERVICE ALERT: IT-2012;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1460077665] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;SOFT;2;e:\ - total: 499.87 Gb - used: 401.28 Gb (80%) - free 98.59 Gb (20%)
[1460077774] SERVICE ALERT: IT-2012;CPU Load;OK;SOFT;2;CPU Load 86% (5 min average)
[1460077785] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;HARD;3;e:\ - total: 499.87 Gb - used: 401.28 Gb (80%) - free 98.59 Gb (20%)
[1460079585] SERVICE ALERT: AD-RMAD;E:\ Drive Space;OK;HARD;3;e:\ - total: 499.87 Gb - used: 395.43 Gb (79%) - free 104.45 Gb (21%)
[1460080516] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 93% (5 min average)
[1460080623] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1460080743] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 94% (5 min average)
[1460080863] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 96% (5 min average)
[1460081116] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1460081463] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 56% (5 min average)
[1460083155] SERVICE ALERT: VC;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460083255] SERVICE ALERT: VC;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1460084463] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 91% (5 min average)
[1460084583] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 92% (5 min average)
[1460084703] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 93% (5 min average)
[1460084716] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (5 min average)
[1460085303] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 97% (5 min average)
[1460085303] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1460085316] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460085436] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1460085556] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460085793] SERVICE FLAPPING ALERT: studentcenterpk;Printer Status;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1460086393] SERVICE ALERT: studentcenterpk;Printer Status;OK;HARD;3;Turn green knobs.
[1460086451] SERVICE ALERT: studentcenterpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1460086503] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 93% (5 min average)
[1460086993] SERVICE ALERT: studentcenterpk;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460087053] SERVICE ALERT: studentcenterpk;Printer Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460087103] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 78% (5 min average)
[1460087114] SERVICE ALERT: studentcenterpk;Printer Status;WARNING;HARD;3;Possible Paper Jam Detected!