[1458619487] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1458619607] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1458619727] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1458620327] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 94% (5 min average)
[1458620927] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 95% (5 min average)
[1458622126] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;1;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1458622246] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;2;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1458622366] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;HARD;3;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1458623327] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 61% (5 min average)
[1458623927] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458624047] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 88% (5 min average)
[1458624647] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458624767] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1458624887] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1458625487] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 68% (5 min average)
[1458626086] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1458626206] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1458626326] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1458627002] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1458627004] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 84.90 Gb - used: 78.77 Gb (93%) - free 6.13 Gb (7%)
[1458627122] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 97% (5 min average)
[1458627124] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 84.90 Gb - used: 78.77 Gb (93%) - free 6.13 Gb (7%)
[1458627242] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 81% (5 min average)
[1458627244] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 84.90 Gb - used: 78.77 Gb (93%) - free 6.13 Gb (7%)
[1458627842] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1458627962] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 92% (5 min average)
[1458628082] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 79% (5 min average)
[1458628682] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1458628802] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 94% (5 min average)
[1458628922] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1458628982] SERVICE ALERT: INTLDB;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1458629102] SERVICE ALERT: INTLDB;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1458629222] SERVICE ALERT: INTLDB;CPU Load;OK;SOFT;3;CPU Load 27% (5 min average)
[1458629566] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;OK;HARD;3;OK: No entries found
[1458630122] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 97% (5 min average)
[1458631322] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 94% (5 min average)
[1458631922] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 76% (5 min average)
[1458631922] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1458632926] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (5 min average)
[1458633122] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 91% (5 min average)
[1458633242] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 69% (5 min average)
[1458633526] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458633646] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1458633766] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1458635566] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 82% (5 min average)
[1458636166] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1458636286] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 97% (5 min average)
[1458636406] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1458637606] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 36% (5 min average)
[1458638806] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1458638926] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 68% (5 min average)
[1458639526] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1458639646] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1458639766] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1458640966] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 56% (5 min average)
[1458642166] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458642286] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1458642406] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1458642698] SERVICE ALERT: INTLAPP;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1458642818] SERVICE ALERT: INTLAPP;CPU Load;WARNING;SOFT;2;CPU Load 92% (5 min average)
[1458642938] SERVICE ALERT: INTLAPP;CPU Load;OK;SOFT;3;CPU Load 0% (5 min average)
[1458643442] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1458643606] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1458644206] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1458648072] SERVICE ALERT: halleground1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1458648095] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 51% - OK!
[1458648099] HOST ALERT: owen101lab3;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.45 ms
[1458648132] SERVICE ALERT: halleground1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1458648165] SERVICE ALERT: owen101lab3;Device Status;OK;HARD;3;Processing job from tray 4
[1458648192] SERVICE ALERT: halleground1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1458648406] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 64% (5 min average)
[1458648433] SERVICE ALERT: owen101lab3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1458649006] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458649126] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 75% (5 min average)
[1458649134] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1458649254] SERVICE ALERT: BOEAPPTEST;CPU Load;OK;SOFT;2;CPU Load 31% (5 min average)
[1458649726] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1458649846] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1458649966] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1458651166] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 92% (5 min average)
[1458651654] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1458651766] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 51% (5 min average)
[1458651774] SERVICE ALERT: BOEAPPTEST;CPU Load;OK;SOFT;2;CPU Load 33% (5 min average)
[1458653491] SERVICE ALERT: markjeffersonpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1458653501] SERVICE ALERT: markjeffersonpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1458653521] HOST ALERT: markjeffersonpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1458653551] SERVICE ALERT: markjeffersonpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1458653561] SERVICE ALERT: markjeffersonpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1458653581] HOST ALERT: markjeffersonpk;UP;SOFT;2;PING WARNING - Packet loss = 90%, RTA = 0.80 ms
[1458653605] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1458653616] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;SOFT;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1458653642] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;OK;HARD;3;c:\ - total: 84.90 Gb - used: 75.63 Gb (89%) - free 9.27 Gb (11%)
[1458653815] SERVICE ALERT: strongpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1458653876] SERVICE ALERT: strongpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1458653935] SERVICE ALERT: strongpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1458654766] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1458654806] SERVICE ALERT: markjeffersonpk;Device Status;OK;HARD;3;Ready
[1458654886] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 76% (5 min average)
[1458655135] SERVICE ALERT: strongpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1458655486] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1458655606] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 80% (5 min average)
[1458656592] SERVICE ALERT: halleground1;Device Status;OK;HARD;3;Checking printer
[1458657406] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458657526] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 94% (5 min average)
[1458657554] SERVICE ALERT: halleground1;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1458657613] SERVICE ALERT: halleground1;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1458657646] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1458657674] SERVICE ALERT: halleground1;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1458658846] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 87% (5 min average)
[1458658876] SERVICE ALERT: halleground1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1458659011] SERVICE ALERT: markjeffersonpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1458659021] SERVICE ALERT: markjeffersonpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1458659041] HOST ALERT: markjeffersonpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1458659071] SERVICE ALERT: markjeffersonpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1458659081] SERVICE ALERT: markjeffersonpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1458659090] HOST ALERT: markjeffersonpk;UP;SOFT;2;PING OK - Packet loss = 73%, RTA = 0.74 ms
[1458659126] SERVICE ALERT: markjeffersonpk;Device Status;OK;SOFT;3;Ready
[1458659135] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;SOFT;3;Printer Paper Is OK! All Trays Still Full
[1458659446] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1458659566] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1458659686] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1458659818] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1458659877] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! 1 of 2 Trays Still Full
[1458659904] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;SOFT;1;WARNING: Replication with ADDC5 failed!
[1458659955] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1458660015] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1458660023] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;SOFT;2;WARNING: Replication with ADDC5 failed!
[1458660075] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1458660144] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;HARD;3;WARNING: Replication with ADDC5 failed!
[1458660925] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1458660985] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1458661045] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1458661525] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1458661586] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 74% - OK!
[1458661646] SERVICE ALERT: markjeffersonpk;Device Status;OK;HARD;3;Warming up
[1458661646] SERVICE FLAPPING ALERT: markjeffersonpk;Device Status;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1458661875] SERVICE ALERT: studentcenterpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1458661886] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1458661946] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1458662006] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1458663206] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Ready
[1458663286] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 90% (5 min average)
[1458663886] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1458666284] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 85% (5 min average)
[1458666884] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458667005] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1458667124] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1458667762] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1458667822] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1458667883] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1458668324] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1458668384] SERVICE ALERT: hallefirst1;Device Status;OK;SOFT;2;Open front door.
[1458668812] SERVICE ALERT: porterpk;Black Cartridge HP CC364X;WARNING;SOFT;1;Black Cartridge HP CC364X is at 2% - WARNING!
[1458668872] SERVICE ALERT: porterpk;Black Cartridge HP CC364X;WARNING;SOFT;2;Black Cartridge HP CC364X is at 2% - WARNING!
[1458668932] SERVICE ALERT: porterpk;Black Cartridge HP CC364X;WARNING;HARD;3;Black Cartridge HP CC364X is at 2% - WARNING!
[1458669314] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (34.1% change >= 20.0% threshold)
[1458669314] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1458669314] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1458669314] HOST DOWNTIME ALERT: MAGICDBPROD;STARTED; Host has entered a period of scheduled downtime
[1458669314] HOST DOWNTIME ALERT: MAGICAPPTEST;STARTED; Host has entered a period of scheduled downtime
[1458670274] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (28.7% change >= 20.0% threshold)
[1458670274] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1458670274] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1458670274] HOST DOWNTIME ALERT: MAGICDBPROD;STARTED; Host has entered a period of scheduled downtime
[1458670274] HOST DOWNTIME ALERT: MAGICAPPTEST;STARTED; Host has entered a period of scheduled downtime
[1458670290] SERVICE ALERT: idmrbstest.emich.edu;eDirectory Status;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_edir, ...) failed. errno is 2: No such file or directory
[1458670405] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1458670410] SERVICE ALERT: idmrbstest.emich.edu;eDirectory Status;CRITICAL;SOFT;2;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_edir, ...) failed. errno is 2: No such file or directory
[1458670422] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (28.7% change >= 20.0% threshold)
[1458670422] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1458670422] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1458670422] HOST DOWNTIME ALERT: MAGICDBPROD;STARTED; Host has entered a period of scheduled downtime
[1458670422] HOST DOWNTIME ALERT: MAGICAPPTEST;STARTED; Host has entered a period of scheduled downtime
[1458670446] SERVICE ALERT: idmrbstest.emich.edu;eDirectory Status;CRITICAL;HARD;3;CRITICAL: Failed to run the ndsd status command successfully!
[1458670465] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1458670525] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1458670722] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 92% (5 min average)
[1458670732] SERVICE ALERT: porterpk;Black Cartridge HP CC364X;CRITICAL;HARD;3;Black Cartridge HP CC364X is at 1% - CRITICAL!
[1458671296] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1458671322] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1458671356] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1458671416] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1458671462] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 1 Trays Empty
[1458671523] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 1 Trays Empty
[1458671582] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 1 Trays Empty
[1458671725] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Clearing paper path
[1458671725] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1458673785] SERVICE ALERT: idmrbstest.emich.edu;eDirectory Status;OK;HARD;3;OK: ndsd status is good.
[1458674582] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1458674922] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (5 min average)
[1458676122] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458676242] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1458676362] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1458676732] SERVICE ALERT: porterpk;Black Cartridge HP CC364X;OK;HARD;3;Black Cartridge HP CC364X is at 100% - OK!
[1458676882] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1458677427] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (30.1% change >= 20.0% threshold)
[1458677427] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1458677427] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1458677427] HOST DOWNTIME ALERT: MAGICDBPROD;STARTED; Host has entered a period of scheduled downtime
[1458677427] HOST DOWNTIME ALERT: MAGICAPPTEST;STARTED; Host has entered a period of scheduled downtime
[1458678013] SERVICE ALERT: studentcenterpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1458678162] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1458678762] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1458678936] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1458678996] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! 1 of 3 Trays Still Full
[1458679361] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 72% (5 min average)
[1458679961] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458680081] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1458680201] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1458680790] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1458680801] SERVICE ALERT: markjeffersonpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1458680820] HOST ALERT: markjeffersonpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1458680850] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1458680850] SERVICE ALERT: markjeffersonpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1458680861] SERVICE ALERT: markjeffersonpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1458680872] HOST ALERT: markjeffersonpk;UP;SOFT;2;PING OK - Packet loss = 78%, RTA = 0.62 ms
[1458680904] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;OK;SOFT;3;Black Print Cartridge HP C8543X is at 74% - OK!
[1458680904] SERVICE ALERT: markjeffersonpk;Device Status;OK;SOFT;2;Processing job from tray 3
[1458680916] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;SOFT;3;Printer Paper Is OK! All Trays Still Full
[1458681031] SERVICE ALERT: ADAUTH-RMAD;E:\ Drive Space;WARNING;SOFT;1;e:\ - total: 199.87 Gb - used: 162.92 Gb (82%) - free 36.95 Gb (18%)
[1458681151] SERVICE ALERT: ADAUTH-RMAD;E:\ Drive Space;WARNING;SOFT;2;e:\ - total: 199.87 Gb - used: 162.92 Gb (82%) - free 36.95 Gb (18%)
[1458681271] SERVICE ALERT: ADAUTH-RMAD;E:\ Drive Space;WARNING;HARD;3;e:\ - total: 199.87 Gb - used: 162.92 Gb (82%) - free 36.95 Gb (18%)
[1458681748] SERVICE ALERT: ADRODC-LIV;AD Replication;OK;HARD;3;OK: replication is up and running
[1458681764] SERVICE ALERT: marshallpk;Black Cartridge HP CC364X;WARNING;SOFT;1;Black Cartridge HP CC364X is at 2% - WARNING!
[1458681824] SERVICE ALERT: marshallpk;Black Cartridge HP CC364X;WARNING;SOFT;2;Black Cartridge HP CC364X is at 2% - WARNING!
[1458681885] SERVICE ALERT: marshallpk;Black Cartridge HP CC364X;WARNING;HARD;3;Black Cartridge HP CC364X is at 2% - WARNING!
[1458683201] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 85% (5 min average)
[1458683801] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1458683921] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1458684041] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1458689441] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 59% (5 min average)
[1458690041] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458690161] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 73% (5 min average)
[1458690761] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458690881] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1458691001] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1458691601] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 81% (5 min average)
[1458692201] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1458692321] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1458692332] SERVICE ALERT: hallegroundcolor;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 1 Trays Empty
[1458692391] SERVICE ALERT: hallegroundcolor;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 1 Trays Empty
[1458692441] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 95% (5 min average)
[1458692451] SERVICE ALERT: hallegroundcolor;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 1 Trays Empty
[1458695508] SERVICE ALERT: DB01-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1458695628] SERVICE ALERT: DB01-2012;CPU Load;OK;SOFT;2;CPU Load 82% (5 min average)
[1458696144] SERVICE ALERT: TARGETXTEST2012;CPU Load;WARNING;SOFT;1;CPU Load 91% (5 min average)
[1458696249] SERVICE ALERT: EVENTS-2012;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1458696264] SERVICE ALERT: TARGETXTEST2012;CPU Load;OK;SOFT;2;CPU Load 19% (5 min average)
[1458696369] SERVICE ALERT: EVENTS-2012;CPU Load;OK;SOFT;2;CPU Load 84% (5 min average)
[1458698869] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1458698928] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1458698987] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1458698990] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1458699049] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1458699059] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1458699109] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1458699119] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1458699121] HOST ALERT: hallefirst2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1458699139] HOST ALERT: hallefirst1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1458699179] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1458699185] HOST ALERT: hallefirst2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1458699209] HOST ALERT: hallefirst1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1458699277] HOST ALERT: hallefirst1;UP;SOFT;3;PING OK - Packet loss = 44%, RTA = 0.71 ms
[1458699607] HOST ALERT: hallefirst1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1458699671] HOST ALERT: hallefirst1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.58 ms
[1458699671] HOST FLAPPING ALERT: hallefirst1;STARTED; Host appears to have started flapping (23.7% change > 20.0% threshold)
[1458699703] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;Ready
[1458699773] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 9% - OK!
[1458700181] SERVICE ALERT: hallefirst1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1458701044] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10022;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1458701047] HOST ALERT: ADWSUSSERVERS;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.133)
[1458701062] SERVICE ALERT: ADWSUSSERVERS;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.102.133 and port 12489: No route to host
[1458701065] HOST ALERT: ADWSUSSERVERS;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.133)
[1458701095] HOST ALERT: ADWSUSSERVERS;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1458701152] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10022;OK;SOFT;2;Eventlog check ok
[1458701179] SERVICE ALERT: ADWSUSSERVERS;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1458701225] SERVICE ALERT: CBORD;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458701325] SERVICE ALERT: CBORD;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1458701982] HOST FLAPPING ALERT: hallefirst1;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1458702841] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1458702961] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1458703081] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 92% (5 min average)
[1458703681] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 95% (5 min average)
[1458703841] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1458704281] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 66% (5 min average)