[1463976325] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 211.49 Gb (96%) - free 8.41 Gb (4%)
[1463976375] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 235 found since last run. oneLevelSearchOps OK. searchOps is too high! 300 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463976495] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463976501] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463976621] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1463976741] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463977095] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 222 found since last run. oneLevelSearchOps OK. searchOps is too high! 245 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463977215] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463977342] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463977815] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 226 found since last run. oneLevelSearchOps OK. searchOps is too high! 285 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463977935] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463977941] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463978125] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 202.93 Gb (92%) - free 16.97 Gb (8%)
[1463978535] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 214 found since last run. oneLevelSearchOps OK. searchOps is too high! 238 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463978655] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463978697] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;1;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1463978817] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;2;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1463978937] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;HARD;3;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1463979255] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 209 found since last run. oneLevelSearchOps OK. searchOps is too high! 264 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463979375] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463979925] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 211.07 Gb (96%) - free 8.83 Gb (4%)
[1463979975] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 201 found since last run. oneLevelSearchOps OK. searchOps is too high! 207 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463980095] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463980339] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 79% (15 min average)
[1463980696] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 237 found since last run. oneLevelSearchOps OK. searchOps is too high! 287 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463980815] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463980939] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463981059] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1463981177] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463981725] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 202.95 Gb (92%) - free 16.95 Gb (8%)
[1463981777] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463982013] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 236 found since last run. oneLevelSearchOps OK. searchOps is too high! 285 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463982134] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463982377] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463982733] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 299 found since last run. oneLevelSearchOps OK. searchOps is too high! 302 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463982853] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463983577] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463984125] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 211.13 Gb (96%) - free 8.77 Gb (4%)
[1463984177] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463985253] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 241 found since last run. oneLevelSearchOps OK. searchOps is too high! 366 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463985373] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463985925] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 203.09 Gb (92%) - free 16.81 Gb (8%)
[1463985963] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463985973] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 210 found since last run. oneLevelSearchOps OK. searchOps is too high! 220 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463986093] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463986137] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;OK;HARD;3;OK: No entries found
[1463986563] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 224 found since last run. oneLevelSearchOps OK. searchOps is too high! 227 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463986683] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463986693] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 232 found since last run. oneLevelSearchOps OK. searchOps is too high! 245 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463986813] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463987177] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463987283] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 273 found since last run. oneLevelSearchOps OK. searchOps is too high! 276 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 595 found since last run.
[1463987404] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463987413] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 286 found since last run. oneLevelSearchOps OK. searchOps is too high! 291 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463987533] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463987725] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 210.09 Gb (96%) - free 9.81 Gb (4%)
[1463987777] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463988003] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 232 found since last run. oneLevelSearchOps OK. searchOps is too high! 234 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 544 found since last run.
[1463988123] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463988133] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 302 found since last run. oneLevelSearchOps OK. searchOps is too high! 310 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463988253] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463988723] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 232 found since last run. oneLevelSearchOps OK. searchOps is too high! 240 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463988843] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463988853] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 310 found since last run. oneLevelSearchOps OK. searchOps is too high! 313 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463988973] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463989443] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 339 found since last run. oneLevelSearchOps OK. searchOps is too high! 343 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 500 found since last run.
[1463989525] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 203.06 Gb (92%) - free 16.84 Gb (8%)
[1463989563] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463990163] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 328 found since last run. oneLevelSearchOps OK. searchOps is too high! 332 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463990173] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 204 found since last run. oneLevelSearchOps OK. searchOps is too high! 220 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463990283] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463990293] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463990883] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 212 found since last run. oneLevelSearchOps OK. searchOps is too high! 215 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463990893] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 266 found since last run. oneLevelSearchOps OK. searchOps is too high! 269 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463991003] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463991013] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463991377] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 77% (15 min average)
[1463991603] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 268 found since last run. oneLevelSearchOps OK. searchOps is too high! 271 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 556 found since last run.
[1463991613] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 317 found since last run. oneLevelSearchOps OK. searchOps is too high! 321 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463991723] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463991733] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463991925] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 212.00 Gb (96%) - free 7.90 Gb (4%)
[1463992323] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 302 found since last run. oneLevelSearchOps OK. searchOps is too high! 303 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 625 found since last run.
[1463992443] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463992933] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 224 found since last run. oneLevelSearchOps OK. searchOps is too high! 275 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463993043] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 325 found since last run. oneLevelSearchOps OK. searchOps is too high! 328 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 574 found since last run.
[1463993053] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463993125] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 203.16 Gb (92%) - free 16.74 Gb (8%)
[1463993163] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463993177] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463993297] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1463993417] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 94% (15 min average)
[1463993653] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 252 found since last run. oneLevelSearchOps OK. searchOps is too high! 265 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463993763] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 378 found since last run. oneLevelSearchOps OK. searchOps is too high! 381 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 502 found since last run.
[1463993773] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463993883] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463994017] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463994137] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463994257] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1463994373] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 311 found since last run. oneLevelSearchOps OK. searchOps is too high! 317 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463994483] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 320 found since last run. oneLevelSearchOps OK. searchOps is too high! 322 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 563 found since last run.
[1463994493] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463994603] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463995093] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 285 found since last run. oneLevelSearchOps OK. searchOps is too high! 301 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463995203] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 480 found since last run. oneLevelSearchOps OK. searchOps is too high! 482 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 633 found since last run.
[1463995213] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463995323] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463995457] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 54% (15 min average)
[1463995813] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 435 found since last run. oneLevelSearchOps OK. searchOps is too high! 445 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463995923] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 226 found since last run. oneLevelSearchOps OK. searchOps is too high! 229 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 529 found since last run.
[1463995933] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463996043] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463996533] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 423 found since last run. oneLevelSearchOps OK. searchOps is too high! 431 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463996643] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 247 found since last run. oneLevelSearchOps OK. searchOps is too high! 253 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463996653] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463996724] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 210.20 Gb (96%) - free 9.70 Gb (4%)
[1463996763] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463997253] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 342 found since last run. oneLevelSearchOps OK. searchOps is too high! 392 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463997363] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 509 found since last run. oneLevelSearchOps OK. searchOps is too high! 511 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 619 found since last run.
[1463997373] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463997483] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463997537] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;1;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1463997657] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;2;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1463997777] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;HARD;3;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1463997973] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 402 found since last run. oneLevelSearchOps OK. searchOps is too high! 420 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463998083] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 495 found since last run. oneLevelSearchOps OK. searchOps is too high! 500 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 600 found since last run.
[1463998093] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463998203] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463998457] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463998524] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 203.25 Gb (92%) - free 16.65 Gb (8%)
[1463998577] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 78% (15 min average)
[1463998693] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 354 found since last run. oneLevelSearchOps OK. searchOps is too high! 363 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463998803] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 475 found since last run. oneLevelSearchOps OK. searchOps is too high! 477 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 637 found since last run.
[1463998813] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463998923] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463999177] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1463999297] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1463999413] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 458 found since last run. oneLevelSearchOps OK. searchOps is too high! 520 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463999417] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463999523] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 426 found since last run. oneLevelSearchOps OK. searchOps is too high! 429 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 618 found since last run.
[1463999533] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463999643] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 333 found since last run. oneLevelSearchOps OK. searchOps is too high! 336 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 686 found since last run.
[1463999763] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 320 found since last run. oneLevelSearchOps OK. searchOps is too high! 323 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 505 found since last run.
[1464000017] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1464000133] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 2481 found since last run. oneLevelSearchOps OK. searchOps is too high! 2510 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 1308 found since last run.
[1464000253] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 317 found since last run. oneLevelSearchOps OK. searchOps is too high! 326 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 571 found since last run.
[1464000324] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 210.55 Gb (96%) - free 9.35 Gb (4%)
[1464000373] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 261 found since last run. oneLevelSearchOps OK. searchOps is too high! 270 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 756 found since last run.
[1464000617] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 45% (15 min average)
[1464001217] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1464001337] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464001457] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 75% (15 min average)
[1464002057] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1464002124] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 203.23 Gb (92%) - free 16.67 Gb (8%)
[1464002177] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464002297] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1464002870] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 1533 found since last run. oneLevelSearchOps OK. searchOps is too high! 1534 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464002897] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1464002990] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464003497] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1464003924] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 211.36 Gb (96%) - free 8.54 Gb (4%)
[1464004097] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1464004217] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464004337] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1464004377] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;OK;HARD;3;OK: No entries found
[1464005124] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 203.30 Gb (92%) - free 16.60 Gb (8%)
[1464005537] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 74% (15 min average)
[1464006137] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464006257] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 44% (15 min average)
[1464007422] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1464007482] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1464007542] SERVICE ALERT: mckennypk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1464007790] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 492 found since last run. oneLevelSearchOps OK. searchOps is too high! 529 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464007910] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464008124] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 211.71 Gb (96%) - free 8.18 Gb (4%)
[1464009257] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1464009324] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 203.31 Gb (92%) - free 16.58 Gb (8%)
[1464009377] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1464009497] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1464010317] SERVICE ALERT: NETIDAD;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464010333] HOST ALERT: NETIDAD;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1464010403] SERVICE ALERT: NETIDAD;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464010404] SERVICE ALERT: NETIDADPROXY;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464010404] SERVICE ALERT: NETIDADPROXY;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464010405] SERVICE ALERT: NETIDAD;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464010421] SERVICE ALERT: NETIDAD;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464010423] HOST ALERT: NETIDAD;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1464010434] HOST ALERT: NETIDADPROXY;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1464010435] SERVICE ALERT: NETIDAD;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464010437] SERVICE ALERT: NETIDAD;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464010465] HOST ALERT: NETIDAD;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1464010489] SERVICE ALERT: NETIDAD;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464010491] SERVICE ALERT: NETIDADPROXY;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464010501] SERVICE ALERT: NETIDADPROXY;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464010503] SERVICE ALERT: NETIDADPROXY;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464010519] HOST ALERT: NETIDAD;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1464010521] HOST ALERT: NETIDADPROXY;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1464010523] SERVICE ALERT: NETIDAD;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464010524] SERVICE ALERT: NETIDADPROXY;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464010524] SERVICE ALERT: NETIDADPROXY;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464010525] SERVICE ALERT: NETIDAD;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464010541] SERVICE ALERT: NETIDAD;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464010553] HOST ALERT: NETIDAD;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1464010554] HOST ALERT: NETIDADPROXY;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1464010555] SERVICE ALERT: NETIDAD;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464010557] SERVICE ALERT: NETIDAD;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464010573] SERVICE ALERT: NETIDAD;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464010576] SERVICE ALERT: NETIDAD;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464010576] SERVICE ALERT: NETIDADPROXY;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464010585] HOST ALERT: NETIDAD;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1464010604] SERVICE ALERT: NETIDAD;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464010606] HOST ALERT: NETIDADPROXY;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1464010609] SERVICE ALERT: NETIDAD;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464010611] SERVICE ALERT: NETIDADPROXY;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464010621] SERVICE ALERT: NETIDADPROXY;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464010623] SERVICE ALERT: NETIDADPROXY;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464010634] HOST ALERT: NETIDAD;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1464010641] HOST ALERT: NETIDADPROXY;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1464010643] SERVICE ALERT: NETIDAD;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464010644] SERVICE ALERT: NETIDADPROXY;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464010644] SERVICE ALERT: NETIDADPROXY;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464010645] SERVICE ALERT: NETIDAD;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464010660] SERVICE ALERT: NETIDAD;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464010661] SERVICE ALERT: NETIDAD;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464010663] SERVICE ALERT: NETIDADPROXY;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464010673] HOST ALERT: NETIDAD;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1464010674] HOST ALERT: NETIDADPROXY;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1464010675] SERVICE ALERT: NETIDAD;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464010693] SERVICE ALERT: NETIDAD;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464010696] SERVICE ALERT: NETIDAD;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464010696] SERVICE ALERT: NETIDADPROXY;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464010705] HOST ALERT: NETIDAD;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1464010724] SERVICE ALERT: NETIDAD;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464010726] HOST ALERT: NETIDADPROXY;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1464010729] SERVICE ALERT: NETIDAD;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464010731] SERVICE ALERT: NETIDADPROXY;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464010741] SERVICE ALERT: NETIDADPROXY;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464010743] SERVICE ALERT: NETIDADPROXY;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464010746] SERVICE ALERT: NETIDAD;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464010754] HOST ALERT: NETIDAD;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1464010761] HOST ALERT: NETIDADPROXY;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1464010780] SERVICE ALERT: NETIDAD;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464010783] SERVICE ALERT: NETIDADPROXY;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464010806] SERVICE ALERT: NETIDAD;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464010808] SERVICE ALERT: NETIDADPROXY;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464010809] SERVICE ALERT: NETIDAD;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464010813] SERVICE ALERT: NETIDAD;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464010813] HOST ALERT: NETIDADPROXY;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1464010816] SERVICE ALERT: NETIDAD;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464010816] SERVICE ALERT: NETIDADPROXY;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464010844] SERVICE ALERT: NETIDAD;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464010846] HOST ALERT: NETIDADPROXY;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1464010853] SERVICE ALERT: NETIDADPROXY;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464010866] SERVICE ALERT: NETIDAD;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464010876] SERVICE ALERT: NETIDAD;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464010876] SERVICE ALERT: NETIDADPROXY;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464010900] SERVICE ALERT: NETIDAD;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464010903] SERVICE ALERT: NETIDADPROXY;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464010914] SERVICE ALERT: NETIDADPROXY;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464010926] SERVICE ALERT: NETIDAD;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464010928] SERVICE ALERT: NETIDADPROXY;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464010929] SERVICE ALERT: NETIDAD;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464010973] SERVICE ALERT: NETIDADPROXY;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464010986] SERVICE ALERT: NETIDAD;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464010990] SERVICE ALERT: NETIDADPROXY;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464010996] SERVICE ALERT: NETIDAD;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464010996] SERVICE ALERT: NETIDADPROXY;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464011034] SERVICE ALERT: NETIDADPROXY;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464011046] SERVICE ALERT: NETIDAD;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464011048] SERVICE ALERT: NETIDADPROXY;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464011049] SERVICE ALERT: NETIDAD;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464011093] SERVICE ALERT: NETIDADPROXY;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464011110] SERVICE ALERT: NETIDADPROXY;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464011116] SERVICE ALERT: NETIDAD;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464011116] SERVICE ALERT: NETIDADPROXY;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464011124] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 210.89 Gb (96%) - free 9.01 Gb (4%)
[1464011154] SERVICE ALERT: NETIDADPROXY;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464011230] SERVICE ALERT: NETIDADPROXY;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464011597] SERVICE ALERT: idm1.emich.edu;Oracle Driver;OK;HARD;3;OK: cn=oradap,cn=driverset1,o=system is running.
[1464011724] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 203.36 Gb (92%) - free 16.54 Gb (8%)
[1464011897] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 88% (15 min average)
[1464012198] SERVICE ALERT: idm1.emich.edu;Oracle Driver;UNKNOWN;SOFT;1;Warning: cn=oradap,cn=driverset1,o=system is stopping!
[1464012317] SERVICE ALERT: idm1.emich.edu;Oracle Driver;UNKNOWN;SOFT;2;Warning: cn=oradap,cn=driverset1,o=system is stopping!
[1464012324] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 217.88 Gb (99%) - free 2.02 Gb (1%)
[1464012437] SERVICE ALERT: idm1.emich.edu;Oracle Driver;UNKNOWN;HARD;3;Warning: cn=oradap,cn=driverset1,o=system is stopping!
[1464012497] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464012596] SERVICE ALERT: NETIDAD;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1464012600] HOST ALERT: NETIDAD;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1464012605] SERVICE ALERT: NETIDADPROXY;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1464012609] HOST ALERT: NETIDADPROXY;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1464012617] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464012624] SERVICE ALERT: NETIDAD;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1464012680] SERVICE ALERT: NETIDAD;Service: RPC;OK;HARD;3;RpcSs: Started
[1464012683] SERVICE ALERT: NETIDADPROXY;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1464012737] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464012766] SERVICE ALERT: NETIDAD;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1464012826] SERVICE ALERT: NETIDAD;VMWare Tools;OK;HARD;3;VMTools: Started
[1464012828] SERVICE ALERT: NETIDADPROXY;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 3 minute(s)
[1464012829] SERVICE ALERT: NETIDAD;C:\ Drive Space;OK;HARD;3;c:\ - total: 79.66 Gb - used: 14.17 Gb (18%) - free 65.49 Gb (82%)
[1464012923] HOST ALERT: NETIDADPROXY;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1464012924] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 203.47 Gb (93%) - free 16.43 Gb (7%)
[1464012926] HOST ALERT: NETIDAD;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1464012984] HOST ALERT: NETIDADPROXY;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1464012987] HOST ALERT: NETIDAD;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1464013023] SERVICE ALERT: NETIDAD;Memory Usage;OK;HARD;3;Memory usage: total:10943.64 MB - used: 494.16 MB (5%) - free: 10449.48 MB (95%)
[1464013025] SERVICE ALERT: NETIDAD;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1464013027] HOST ALERT: NETIDAD;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1464013038] SERVICE ALERT: idm1.emich.edu;Oracle Driver;CRITICAL;HARD;3;CRITICAL: cn=oradap,cn=driverset1,o=system is stopped!
[1464013041] SERVICE ALERT: NETIDAD;CPU Load;OK;HARD;3;CPU Load 49% (15 min average)
[1464013055] SERVICE ALERT: NETIDAD;Service: Windows Time;OK;HARD;3;W32Time: Started
[1464013060] HOST ALERT: NETIDADPROXY;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1464013109] SERVICE ALERT: NETIDAD;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1464013111] SERVICE ALERT: NETIDADPROXY;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1464013115] HOST ALERT: NETIDADPROXY;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1464013121] SERVICE ALERT: NETIDADPROXY;Service: RPC;OK;HARD;3;RpcSs: Started
[1464013123] SERVICE ALERT: NETIDADPROXY;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1464013193] SERVICE ALERT: NETIDAD;Service: WMI;OK;HARD;3;Winmgmt: Started
[1464013473] SERVICE ALERT: NETIDADPROXY;VMWare Tools;OK;HARD;3;VMTools: Started
[1464013496] SERVICE ALERT: NETIDAD;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 8 minute(s)
[1464013496] SERVICE ALERT: NETIDADPROXY;C:\ Drive Space;OK;HARD;3;c:\ - total: 79.66 Gb - used: 13.16 Gb (17%) - free 66.49 Gb (83%)
[1464013534] SERVICE ALERT: NETIDADPROXY;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1464013537] SERVICE ALERT: NETIDAD;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1464013610] SERVICE ALERT: NETIDADPROXY;CPU Load;OK;HARD;3;CPU Load 52% (15 min average)
[1464013624] SERVICE ALERT: NETIDADPROXY;Memory Usage;OK;HARD;3;Memory usage: total:4799.59 MB - used: 1300.16 MB (27%) - free: 3499.43 MB (73%)
[1464013624] SERVICE ALERT: NETIDADPROXY;Service: WMI;OK;HARD;3;Winmgmt: Started
[1464013937] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 90% (15 min average)
[1464014137] SERVICE ALERT: NETIDAD;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.99.247 and port 12489: Connection refused
[1464014257] SERVICE ALERT: NETIDAD;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1464014537] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464014657] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1464014777] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464015342] SERVICE ALERT: mckennypk;Device Status;OK;HARD;3;50 72 65 73 73 20 17 20 74 6F 20 73 65 65 20 73  74 65 70  by step information
[1464015410] SERVICE ALERT: NETIDADPROXY;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1464015530] SERVICE ALERT: NETIDADPROXY;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1464015650] SERVICE ALERT: NETIDADPROXY;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1464015924] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 209.29 Gb (95%) - free 10.60 Gb (5%)
[1464015942] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1464016002] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1464016062] SERVICE ALERT: mckennypk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1464016224] SERVICE ALERT: sillpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1464016236] SERVICE ALERT: sillpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1464016250] SERVICE ALERT: NETIDADPROXY;CPU Load;OK;HARD;3;CPU Load 53% (15 min average)
[1464016250] SERVICE FLAPPING ALERT: NETIDADPROXY;CPU Load;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1464016283] SERVICE ALERT: sillpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1464016296] SERVICE ALERT: sillpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1464016344] SERVICE ALERT: sillpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1464016356] SERVICE ALERT: sillpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1464016721] SERVICE ALERT: NETIDADPROXY;Service: RPC;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1464016723] SERVICE ALERT: NETIDADPROXY;Service: Firewall;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1464016841] SERVICE ALERT: NETIDADPROXY;Service: RPC;OK;SOFT;2;RpcSs: Started
[1464016843] SERVICE ALERT: NETIDADPROXY;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1464016986] SERVICE ALERT: NETIDAD;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464017086] SERVICE ALERT: NETIDAD;Service: User Profile;UNKNOWN;SOFT;2;Failed to open service manager: 1115: A system shutdown is in progress.
[1464017114] SERVICE ALERT: NETIDADPROXY;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464017115] SERVICE ALERT: NETIDAD;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464017206] SERVICE ALERT: NETIDAD;Service: User Profile;OK;SOFT;3;ProfSvc: Started
[1464017214] SERVICE ALERT: NETIDADPROXY;C:\ Drive Space;OK;SOFT;2;c:\ - total: 79.66 Gb - used: 13.45 Gb (17%) - free 66.20 Gb (83%)
[1464017215] SERVICE ALERT: NETIDAD;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 1 minute(s)
[1464017724] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 203.60 Gb (93%) - free 16.30 Gb (7%)
[1464018361] HOST ALERT: NETIDAD;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1464018377] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1464018410] HOST ALERT: NETIDAD;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1464018528] SERVICE ALERT: NETIDAD;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464018530] SERVICE ALERT: NETIDADPROXY;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464018558] HOST ALERT: NETIDAD;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1464018560] HOST ALERT: NETIDADPROXY;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1464018613] SERVICE ALERT: NETIDAD;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464018616] SERVICE ALERT: NETIDAD;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464018616] SERVICE ALERT: NETIDADPROXY;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464018643] HOST ALERT: NETIDAD;DOWN;SOFT;2;(Host check timed out after 30.00 seconds)
[1464018644] SERVICE ALERT: NETIDAD;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464018646] HOST ALERT: NETIDADPROXY;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1464018648] SERVICE ALERT: NETIDAD;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464018650] SERVICE ALERT: NETIDADPROXY;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464018661] SERVICE ALERT: NETIDADPROXY;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464018663] SERVICE ALERT: NETIDADPROXY;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464018668] SERVICE ALERT: NETIDADPROXY;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464018674] HOST ALERT: NETIDAD;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1464018680] HOST ALERT: NETIDADPROXY;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1464018700] SERVICE ALERT: NETIDAD;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464018703] SERVICE ALERT: NETIDADPROXY;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464018730] HOST ALERT: NETIDAD;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1464018733] SERVICE ALERT: NETIDAD;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464018733] HOST ALERT: NETIDADPROXY;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1464018736] SERVICE ALERT: NETIDAD;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464018736] SERVICE ALERT: NETIDADPROXY;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464018763] HOST ALERT: NETIDAD;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1464018764] SERVICE ALERT: NETIDAD;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464018766] HOST ALERT: NETIDADPROXY;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1464018768] SERVICE ALERT: NETIDAD;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464018770] SERVICE ALERT: NETIDADPROXY;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464018781] SERVICE ALERT: NETIDADPROXY;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464018783] SERVICE ALERT: NETIDADPROXY;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464018788] SERVICE ALERT: NETIDADPROXY;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464018794] HOST ALERT: NETIDAD;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1464018800] HOST ALERT: NETIDADPROXY;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1464018820] SERVICE ALERT: NETIDAD;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464018823] SERVICE ALERT: NETIDADPROXY;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464018845] SERVICE ALERT: NETIDAD;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464018848] SERVICE ALERT: NETIDADPROXY;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464018849] SERVICE ALERT: NETIDAD;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464018850] HOST ALERT: NETIDAD;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1464018853] SERVICE ALERT: NETIDAD;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464018853] HOST ALERT: NETIDADPROXY;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1464018856] SERVICE ALERT: NETIDAD;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464018856] SERVICE ALERT: NETIDADPROXY;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464018883] HOST ALERT: NETIDAD;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1464018884] SERVICE ALERT: NETIDAD;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464018886] HOST ALERT: NETIDADPROXY;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1464018893] SERVICE ALERT: NETIDADPROXY;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464018901] SERVICE ALERT: NETIDADPROXY;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464018903] SERVICE ALERT: NETIDADPROXY;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464018908] SERVICE ALERT: NETIDADPROXY;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464018914] HOST ALERT: NETIDAD;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1464018923] HOST ALERT: NETIDADPROXY;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1464018939] SERVICE ALERT: NETIDAD;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464018943] SERVICE ALERT: NETIDADPROXY;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464018954] SERVICE ALERT: NETIDADPROXY;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464018965] SERVICE ALERT: NETIDAD;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464018968] SERVICE ALERT: NETIDADPROXY;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464018969] SERVICE ALERT: NETIDAD;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464018969] HOST ALERT: NETIDAD;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1464018973] HOST ALERT: NETIDADPROXY;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1464018977] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464019013] SERVICE ALERT: NETIDADPROXY;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464019026] SERVICE ALERT: NETIDAD;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464019034] SERVICE ALERT: NETIDADPROXY;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464019035] SERVICE ALERT: NETIDAD;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464019043] SERVICE ALERT: NETIDAD;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464019043] SERVICE ALERT: NETIDADPROXY;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464019043] SERVICE ALERT: NETIDADPROXY;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464019045] SERVICE ALERT: NETIDAD;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464019061] SERVICE ALERT: NETIDAD;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464019074] SERVICE ALERT: NETIDADPROXY;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464019075] SERVICE ALERT: NETIDAD;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464019076] SERVICE ALERT: NETIDAD;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464019085] SERVICE ALERT: NETIDAD;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464019088] SERVICE ALERT: NETIDADPROXY;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464019089] SERVICE ALERT: NETIDAD;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464019133] SERVICE ALERT: NETIDADPROXY;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464019146] SERVICE ALERT: NETIDAD;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464019154] SERVICE ALERT: NETIDADPROXY;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464019155] SERVICE ALERT: NETIDAD;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464019163] SERVICE ALERT: NETIDAD;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464019163] SERVICE ALERT: NETIDADPROXY;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464019163] SERVICE ALERT: NETIDADPROXY;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464019165] SERVICE ALERT: NETIDAD;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464019181] SERVICE ALERT: NETIDAD;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464019194] SERVICE ALERT: NETIDADPROXY;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464019195] SERVICE ALERT: NETIDAD;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464019196] SERVICE ALERT: NETIDAD;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464019266] SERVICE ALERT: NETIDAD;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464019274] SERVICE ALERT: NETIDADPROXY;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464019275] SERVICE ALERT: NETIDAD;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464019283] SERVICE ALERT: NETIDAD;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464019283] SERVICE ALERT: NETIDADPROXY;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464019283] SERVICE ALERT: NETIDADPROXY;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464019285] SERVICE ALERT: NETIDAD;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464019301] SERVICE ALERT: NETIDAD;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464019315] SERVICE ALERT: NETIDAD;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464019316] SERVICE ALERT: NETIDAD;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464019524] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 209.80 Gb (95%) - free 10.09 Gb (5%)
[1464020510] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 4745 found since last run. oneLevelSearchOps OK. searchOps is too high! 4749 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464020544] SERVICE ALERT: sillpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1464020557] SERVICE ALERT: sillpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 48% - OK!
[1464020630] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464021323] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 203.62 Gb (93%) - free 16.28 Gb (7%)
[1464021377] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1464021646] SERVICE ALERT: NETIDAD;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1464021646] SERVICE FLAPPING ALERT: NETIDAD;Service: User Profile;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1464021650] HOST ALERT: NETIDAD;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1464021655] SERVICE ALERT: NETIDAD;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 2 minute(s)
[1464021655] SERVICE FLAPPING ALERT: NETIDAD;Uptime;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1464021663] SERVICE ALERT: NETIDAD;Memory Usage;OK;HARD;3;Memory usage: total:10943.64 MB - used: 1187.89 MB (11%) - free: 9755.76 MB (89%)
[1464021663] SERVICE FLAPPING ALERT: NETIDAD;Memory Usage;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1464021663] SERVICE ALERT: NETIDADPROXY;Service: WMI;OK;HARD;3;Winmgmt: Started
[1464021663] SERVICE FLAPPING ALERT: NETIDADPROXY;Service: WMI;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1464021663] SERVICE ALERT: NETIDADPROXY;Memory Usage;OK;HARD;3;Memory usage: total:4799.59 MB - used: 510.41 MB (11%) - free: 4289.18 MB (89%)
[1464021663] SERVICE FLAPPING ALERT: NETIDADPROXY;Memory Usage;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1464021663] SERVICE ALERT: NETIDADPROXY;C:\ Drive Space;OK;HARD;3;c:\ - total: 79.66 Gb - used: 12.72 Gb (16%) - free 66.93 Gb (84%)
[1464021663] SERVICE FLAPPING ALERT: NETIDADPROXY;C:\ Drive Space;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1464021665] SERVICE ALERT: NETIDAD;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1464021665] SERVICE FLAPPING ALERT: NETIDAD;Service: DNS Client;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1464021667] HOST ALERT: NETIDADPROXY;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1464021680] SERVICE ALERT: NETIDAD;CPU Load;OK;HARD;3;CPU Load 17% (15 min average)
[1464021680] SERVICE FLAPPING ALERT: NETIDAD;CPU Load;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1464021695] SERVICE ALERT: NETIDAD;Service: Windows Time;OK;HARD;3;W32Time: Started
[1464021695] SERVICE FLAPPING ALERT: NETIDAD;Service: Windows Time;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1464021696] SERVICE ALERT: NETIDAD;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1464021696] SERVICE FLAPPING ALERT: NETIDAD;NSClient++ Version;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1464021748] SERVICE ALERT: NETIDAD;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1464021748] SERVICE FLAPPING ALERT: NETIDAD;Service: Firewall;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1464021750] SERVICE ALERT: NETIDADPROXY;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1464021750] SERVICE FLAPPING ALERT: NETIDADPROXY;NSClient++ Version;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1464021833] SERVICE ALERT: NETIDAD;Service: WMI;OK;HARD;3;Winmgmt: Started
[1464021833] SERVICE FLAPPING ALERT: NETIDAD;Service: WMI;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1464021836] SERVICE ALERT: NETIDAD;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1464021836] SERVICE FLAPPING ALERT: NETIDAD;Service: NetLogin;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1464021836] SERVICE ALERT: NETIDADPROXY;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1464021836] SERVICE FLAPPING ALERT: NETIDADPROXY;Service: DNS Client;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1464021864] SERVICE ALERT: NETIDAD;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1464021864] SERVICE FLAPPING ALERT: NETIDAD;Service: Group Policy;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1464021881] SERVICE ALERT: NETIDADPROXY;Service: RPC;OK;HARD;3;RpcSs: Started
[1464021881] SERVICE FLAPPING ALERT: NETIDADPROXY;Service: RPC;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1464021883] SERVICE ALERT: NETIDADPROXY;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1464021883] SERVICE FLAPPING ALERT: NETIDADPROXY;Service: Firewall;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1464021888] SERVICE ALERT: NETIDADPROXY;CPU Load;OK;HARD;3;CPU Load 1% (15 min average)
[1464021919] SERVICE ALERT: NETIDAD;Service: RPC;OK;HARD;3;RpcSs: Started
[1464021919] SERVICE FLAPPING ALERT: NETIDAD;Service: RPC;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1464021922] SERVICE ALERT: NETIDADPROXY;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1464021922] SERVICE FLAPPING ALERT: NETIDADPROXY;Service: Group Policy;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1464021977] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464021980] HOST ALERT: NETIDAD;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1464022065] SERVICE ALERT: NETIDAD;VMWare Tools;OK;HARD;3;VMTools: Started
[1464022065] SERVICE FLAPPING ALERT: NETIDAD;VMWare Tools;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1464022068] SERVICE ALERT: NETIDADPROXY;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 6 minute(s)
[1464022068] SERVICE FLAPPING ALERT: NETIDADPROXY;Uptime;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1464022069] SERVICE ALERT: NETIDAD;C:\ Drive Space;OK;HARD;3;c:\ - total: 79.66 Gb - used: 14.16 Gb (18%) - free 65.49 Gb (82%)
[1464022069] SERVICE FLAPPING ALERT: NETIDAD;C:\ Drive Space;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1464022070] HOST ALERT: NETIDAD;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1464022111] SERVICE ALERT: NETIDADPROXY;VMWare Tools;OK;HARD;3;VMTools: Started
[1464022111] SERVICE FLAPPING ALERT: NETIDADPROXY;VMWare Tools;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1464022160] HOST ALERT: NETIDAD;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1464022173] SERVICE ALERT: NETIDADPROXY;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1464022173] SERVICE FLAPPING ALERT: NETIDADPROXY;Service: User Profile;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1464022224] HOST ALERT: NETIDAD;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.40 ms
[1464023123] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 211.35 Gb (96%) - free 8.55 Gb (4%)
[1464024923] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 203.67 Gb (93%) - free 16.23 Gb (7%)
[1464027323] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 209.16 Gb (95%) - free 10.74 Gb (5%)
[1464029119] SERVICE ALERT: NETIDAD;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.99.247 and port 12489: Connection refused
[1464029123] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 203.57 Gb (93%) - free 16.33 Gb (7%)
[1464029239] SERVICE ALERT: NETIDAD;Service: RPC;OK;SOFT;2;RpcSs: Started
[1464029777] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 91% (15 min average)
[1464030054] SERVICE ALERT: NETIDADPROXY;Service: Windows Time;OK;HARD;3;W32Time: Started
[1464030377] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464030497] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464030617] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464030748] SERVICE FLAPPING ALERT: NETIDAD;Service: Firewall;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1464030750] SERVICE FLAPPING ALERT: NETIDADPROXY;NSClient++ Version;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1464030833] SERVICE FLAPPING ALERT: NETIDAD;Service: WMI;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1464030836] SERVICE FLAPPING ALERT: NETIDAD;Service: NetLogin;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1464030836] SERVICE FLAPPING ALERT: NETIDADPROXY;Service: DNS Client;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1464030863] SERVICE FLAPPING ALERT: NETIDAD;Service: Group Policy;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1464030881] SERVICE FLAPPING ALERT: NETIDADPROXY;Service: RPC;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1464030883] SERVICE FLAPPING ALERT: NETIDADPROXY;Service: Firewall;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1464030888] SERVICE FLAPPING ALERT: NETIDADPROXY;CPU Load;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1464030922] SERVICE FLAPPING ALERT: NETIDADPROXY;Service: Group Policy;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1464030923] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 209.47 Gb (95%) - free 10.43 Gb (5%)
[1464031039] SERVICE FLAPPING ALERT: NETIDAD;Service: RPC;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1464031065] SERVICE FLAPPING ALERT: NETIDAD;VMWare Tools;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1464031065] SERVICE FLAPPING ALERT: NETIDAD;C:\ Drive Space;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1464031067] SERVICE FLAPPING ALERT: NETIDADPROXY;Uptime;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1464031111] SERVICE FLAPPING ALERT: NETIDADPROXY;VMWare Tools;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1464031173] SERVICE FLAPPING ALERT: NETIDADPROXY;Service: User Profile;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1464031245] SERVICE FLAPPING ALERT: NETIDAD;Service: User Profile;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1464031254] SERVICE FLAPPING ALERT: NETIDADPROXY;C:\ Drive Space;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1464031255] SERVICE FLAPPING ALERT: NETIDAD;Uptime;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1464031263] SERVICE FLAPPING ALERT: NETIDAD;Memory Usage;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1464031263] SERVICE FLAPPING ALERT: NETIDADPROXY;Memory Usage;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1464031263] SERVICE FLAPPING ALERT: NETIDADPROXY;Service: WMI;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1464031264] SERVICE FLAPPING ALERT: NETIDAD;Service: DNS Client;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1464031280] SERVICE FLAPPING ALERT: NETIDAD;CPU Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1464031295] SERVICE FLAPPING ALERT: NETIDAD;Service: Windows Time;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1464031296] SERVICE FLAPPING ALERT: NETIDAD;NSClient++ Version;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1464032260] SERVICE ALERT: mckennypk;Device Status;OK;HARD;3;BAD DUPLEXER CONNECTION 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1464032417] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 83% (15 min average)
[1464032723] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 203.59 Gb (93%) - free 16.31 Gb (7%)
[1464034217] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464034337] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464034457] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464034523] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 210.38 Gb (96%) - free 9.52 Gb (4%)
[1464036323] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 203.60 Gb (93%) - free 16.30 Gb (7%)
[1464038169] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1464038169] SERVICE ALERT: hallefirstcolor;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1464038199] HOST ALERT: hallefirstcolor;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1464038229] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1464038229] SERVICE ALERT: hallefirstcolor;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1464038259] HOST ALERT: hallefirstcolor;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1464038278] SERVICE ALERT: hallefirstcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1464038289] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1464038289] SERVICE ALERT: hallefirstcolor;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1464038308] HOST ALERT: hallefirstcolor;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1464038333] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1464038334] SERVICE ALERT: hallefirstcolor;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1464038338] SERVICE ALERT: hallefirstcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1464038363] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1464038363] HOST ALERT: hallefirstcolor;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1464038393] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1464038394] SERVICE ALERT: hallefirstcolor;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1464038398] SERVICE ALERT: hallefirstcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1464038409] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1464038423] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1464038423] HOST ALERT: hallefirstcolor;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1464038453] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1464038454] SERVICE ALERT: hallefirstcolor;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1464038457] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1464038469] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1464038483] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1464038483] HOST ALERT: hallefirstcolor;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1464038517] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1464038529] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1464038539] SERVICE ALERT: hallefirstcolor;Toner Collection Unit HP CE265A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1464038547] HOST ALERT: hallefirstcolor;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1464038577] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1464038599] SERVICE ALERT: hallefirstcolor;Toner Collection Unit HP CE265A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1464038607] HOST ALERT: hallefirstcolor;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1464038659] SERVICE ALERT: hallefirstcolor;Toner Collection Unit HP CE265A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1464038689] HOST ALERT: hallefirstcolor;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1464038779] HOST ALERT: hallefirstcolor;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1464042257] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 84% (15 min average)
[1464042323] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 211.62 Gb (96%) - free 8.27 Gb (4%)
[1464042857] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464042977] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1464043097] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464044123] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 207.33 Gb (94%) - free 12.57 Gb (6%)
[1464045323] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 212.54 Gb (97%) - free 7.36 Gb (3%)
[1464045437] SERVICE ALERT: idm1.emich.edu;Oracle Driver;OK;HARD;3;OK: cn=oradap,cn=driverset1,o=system is running.
[1464047123] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 207.33 Gb (94%) - free 12.56 Gb (6%)
[1464047723] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 209.58 Gb (95%) - free 10.32 Gb (5%)
[1464048295] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;SOFT;1;WARNING: Replication with ADDC5-OLD failed!
[1464048414] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;SOFT;2;WARNING: Replication with ADDC5-OLD failed!
[1464048535] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;HARD;3;WARNING: Replication with ADDC5-OLD failed!
[1464048971] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464049522] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 207.33 Gb (94%) - free 12.57 Gb (6%)
[1464049571] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 228 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464049691] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464050291] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 449 found since last run. oneLevelSearchOps OK. searchOps is too high! 630 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464050411] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464051011] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 542 found since last run. oneLevelSearchOps OK. searchOps is too high! 572 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464051131] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 211 found since last run. oneLevelSearchOps OK. searchOps is too high! 253 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464051251] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464051322] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 210.28 Gb (96%) - free 9.62 Gb (4%)
[1464051851] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 582 found since last run. oneLevelSearchOps OK. searchOps is too high! 590 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464051971] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464052203] SERVICE ALERT: AD-RMAD;E:\ Drive Space;CRITICAL;HARD;3;e:\ - total: 599.87 Gb - used: 542.35 Gb (90%) - free 57.52 Gb (10%)
[1464052571] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 743 found since last run. oneLevelSearchOps OK. searchOps is too high! 812 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464052691] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 215 found since last run. oneLevelSearchOps OK. searchOps is too high! 217 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464052811] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464053411] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 391 found since last run. oneLevelSearchOps OK. searchOps is too high! 479 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464053531] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464053722] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 207.50 Gb (94%) - free 12.40 Gb (6%)
[1464054131] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 120 found since last run. wholeSubtreeSearchOps is too high! 606 found since last run. oneLevelSearchOps OK. searchOps is too high! 793 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 642 found since last run.
[1464054251] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464054322] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 211.27 Gb (96%) - free 8.63 Gb (4%)
[1464054497] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1464054603] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;HARD;3;e:\ - total: 599.87 Gb - used: 536.19 Gb (89%) - free 63.68 Gb (11%)
[1464054851] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 642 found since last run. oneLevelSearchOps OK. searchOps is too high! 831 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 510 found since last run.
[1464054971] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464055571] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 164 found since last run. wholeSubtreeSearchOps is too high! 951 found since last run. oneLevelSearchOps OK. searchOps is too high! 1191 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 919 found since last run.
[1464055691] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464055697] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 80% (15 min average)
[1464056122] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 207.36 Gb (94%) - free 12.54 Gb (6%)
[1464056291] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 137 found since last run. wholeSubtreeSearchOps is too high! 529 found since last run. oneLevelSearchOps OK. searchOps is too high! 693 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 743 found since last run.
[1464056297] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464056411] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464056417] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464056537] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464057011] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 165 found since last run. wholeSubtreeSearchOps is too high! 549 found since last run. oneLevelSearchOps OK. searchOps is too high! 655 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 829 found since last run.
[1464057131] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464057137] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1464057322] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 212.03 Gb (96%) - free 7.87 Gb (4%)
[1464057731] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 164 found since last run. wholeSubtreeSearchOps is too high! 474 found since last run. oneLevelSearchOps OK. searchOps is too high! 507 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 932 found since last run.
[1464057737] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464057737] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1464057851] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464058337] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1464058451] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 299 found since last run. oneLevelSearchOps OK. searchOps is too high! 376 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 574 found since last run.
[1464058571] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464058937] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1464059057] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464059122] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 207.44 Gb (94%) - free 12.46 Gb (6%)
[1464059171] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 123 found since last run. wholeSubtreeSearchOps is too high! 397 found since last run. oneLevelSearchOps OK. searchOps is too high! 459 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 690 found since last run.
[1464059177] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 93% (15 min average)
[1464059291] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464059777] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464059891] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 323 found since last run. oneLevelSearchOps OK. searchOps is too high! 446 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 590 found since last run.
[1464059897] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464060011] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464060017] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464060611] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 105 found since last run. wholeSubtreeSearchOps is too high! 361 found since last run. oneLevelSearchOps OK. searchOps is too high! 488 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 621 found since last run.
[1464060617] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 77% (15 min average)
[1464060731] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464061217] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464061331] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 303 found since last run. oneLevelSearchOps OK. searchOps is too high! 421 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 611 found since last run.
[1464061337] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464061451] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464061457] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464061522] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 213.06 Gb (97%) - free 6.84 Gb (3%)
[1464062051] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 224 found since last run. oneLevelSearchOps OK. searchOps is too high! 236 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 547 found since last run.
[1464062171] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.