[1478668097] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1478668217] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1478668337] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1478670737] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 85% (15 min average)
[1478671337] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478671457] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1478671577] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1478671984] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1478672009] SERVICE ALERT: ADWSUSCLIENTS;E:\ Drive Space;WARNING;SOFT;1;e:\ - total: 399.87 Gb - used: 324.68 Gb (81%) - free 75.19 Gb (19%)
[1478672103] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1478672129] SERVICE ALERT: ADWSUSCLIENTS;E:\ Drive Space;WARNING;SOFT;2;e:\ - total: 399.87 Gb - used: 324.68 Gb (81%) - free 75.19 Gb (19%)
[1478672177] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 85% (15 min average)
[1478672249] SERVICE ALERT: ADWSUSCLIENTS;E:\ Drive Space;WARNING;HARD;3;e:\ - total: 399.87 Gb - used: 324.68 Gb (81%) - free 75.19 Gb (19%)
[1478672777] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478672897] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1478673000] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1478673017] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1478673306] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1478673423] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1478673543] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1478674217] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 84% (15 min average)
[1478674817] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478674937] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1478675057] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1478678057] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1478678343] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1478678524] SERVICE ALERT: AXIOM-TEST-2012;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478678644] SERVICE ALERT: AXIOM-TEST-2012;CPU Load;OK;SOFT;2;CPU Load 1% (15 min average)
[1478678657] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1478678777] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1478678897] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1478678943] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1478679063] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;2;Warning:  idm1  are more than 120 seconds behind!
[1478679183] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1478679689] SERVICE ALERT: epprogsql2;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478679783] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1478679783] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1478679809] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1478679929] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1478680384] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1478680697] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1478680983] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1478681103] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1478681297] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478681417] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1478681537] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 76% (15 min average)
[1478681703] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1478681823] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;2;Warning:  idm1  are more than 120 seconds behind!
[1478681943] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1478682137] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1478682257] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 89% (15 min average)
[1478682543] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1478683529] SERVICE ALERT: epprogsql2;CPU Load;OK;HARD;3;CPU Load 39% (15 min average)
[1478683803] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (41.0% change >= 20.0% threshold)
[1478683803] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (37.6% change >= 20.0% threshold)
[1478683803] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1478683803] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1478683803] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1478683803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1478683803] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1478683803] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1478683803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1478683803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1478683803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1478683803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1478683803] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1478683803] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1478683803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1478683803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1478683803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1478683803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1478683803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1478685603] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (27.0% change >= 20.0% threshold)
[1478685603] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (34.8% change >= 20.0% threshold)
[1478685603] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1478685603] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1478685603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1478685603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1478685603] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1478685603] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1478685603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1478685603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1478685603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1478685603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1478685603] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1478685603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1478685603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1478685603] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1478685603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1478685603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1478685603] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1478690055] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478690175] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1478690295] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1478691000] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service has exited from a period of scheduled downtime
[1478692095] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 78% (15 min average)
[1478692695] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1478692815] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1478692935] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1478693549] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1541 found since last run.
[1478693576] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1525 found since last run.
[1478693669] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478693696] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478693943] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1478694269] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1529 found since last run.
[1478694296] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1547 found since last run.
[1478694389] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478694416] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478694456] SERVICE ALERT: halleground1;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1478694516] SERVICE ALERT: halleground1;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1478694576] SERVICE ALERT: halleground1;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1478694597] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1478694629] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1478694657] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1478694659] HOST ALERT: prayharroldpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1478694689] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1478694697] HOST ALERT: prayharroldpk;UP;SOFT;2;PING OK - Packet loss = 44%, RTA = 0.66 ms
[1478694711] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;OK;SOFT;3;Black Print Cartridge HP C8543X is at 7% - OK!
[1478694734] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1478694742] SERVICE ALERT: prayharroldpk;Paper Tray Status;OK;SOFT;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1478694779] SERVICE ALERT: halleground1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1478694839] SERVICE ALERT: halleground1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1478694899] SERVICE ALERT: sillpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1478694899] SERVICE ALERT: halleground1;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1478694959] SERVICE ALERT: sillpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1478694989] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1478694989] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1547 found since last run.
[1478695016] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1550 found since last run.
[1478695019] SERVICE ALERT: sillpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1478695049] HOST ALERT: sillpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1478695049] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1478695109] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478695109] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1478695113] HOST ALERT: sillpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1478695136] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478695139] HOST ALERT: halleground1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1478695180] HOST ALERT: halleground1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.56 ms
[1478695270] SERVICE ALERT: sillpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1478695330] SERVICE ALERT: sillpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1478695334] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478695343] SERVICE ALERT: prayharroldpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1478695390] SERVICE ALERT: sillpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1478695403] SERVICE ALERT: prayharroldpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1478695454] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1478695463] SERVICE ALERT: prayharroldpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1478695469] SERVICE ALERT: sillpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1478695529] SERVICE ALERT: sillpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1478695574] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1478695589] SERVICE ALERT: sillpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1478695709] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1627 found since last run.
[1478695736] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1589 found since last run.
[1478695829] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478695856] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478696063] SERVICE ALERT: prayharroldpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1478696174] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1478696429] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1526 found since last run.
[1478696456] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1520 found since last run.
[1478696549] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478696576] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478696704] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1478696729] HOST ALERT: halleground1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1478696764] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1478696774] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1478696793] HOST ALERT: halleground1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.54 ms
[1478696824] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1478696859] SERVICE ALERT: studentcenterpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1478696894] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1478696919] SERVICE ALERT: studentcenterpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1478696980] SERVICE ALERT: studentcenterpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1478697014] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1478697149] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7575 found since last run. oneLevelSearchOps OK. searchOps is too high! 7818 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478697269] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478697509] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 7% - OK!
[1478697869] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5703 found since last run. oneLevelSearchOps OK. searchOps is too high! 5964 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478697954] HOST ALERT: halleground2;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.53 ms
[1478697989] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478698062] SERVICE ALERT: halleground2;Device Status;OK;HARD;3;INPUT PAPER PATH OPEN 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1478698109] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1478698124] HOST ALERT: halleground2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1478698169] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1478698214] HOST ALERT: halleground2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1478698229] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1478698250] HOST ALERT: halleground2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1478698340] HOST ALERT: halleground2;DOWN;SOFT;4;(Host check timed out after 30.00 seconds)
[1478698430] HOST ALERT: halleground2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1478698520] HOST ALERT: halleground2;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1478698610] HOST ALERT: halleground2;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1478698623] SERVICE ALERT: studentcenterpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1478698668] SERVICE ALERT: halleground2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1478698698] HOST ALERT: halleground2;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1478698728] SERVICE ALERT: halleground2;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1478698758] HOST ALERT: halleground2;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1478698788] SERVICE ALERT: halleground2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1478698818] HOST ALERT: halleground2;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1478698976] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5445 found since last run. oneLevelSearchOps OK. searchOps is too high! 5828 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478699096] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478699379] SERVICE ALERT: studentcenterpk;Device Status;OK;HARD;3;Ready 164.76.150.252
[1478699554] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1478699614] SERVICE ALERT: hallefirst2;Device Status;OK;SOFT;2;Clearing paper path
[1478699789] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5032 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478699909] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478700003] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (36.2% change >= 20.0% threshold)
[1478700003] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1478700003] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1478700003] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1478700003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1478700003] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1478700003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1478700003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1478700003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1478700003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1478700003] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1478700003] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1478700003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1478700003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1478700003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1478700003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1478700003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1478700014] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 47% (15 min average)
[1478700296] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5189 found since last run. oneLevelSearchOps OK. searchOps is too high! 5625 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478700416] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478700496] HOST ALERT: hallefirst2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1478700509] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6342 found since last run. oneLevelSearchOps OK. searchOps is too high! 6558 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478700529] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1478700559] HOST ALERT: hallefirst2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1478700575] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1478700589] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1478700605] HOST ALERT: hallefirst2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1478700629] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478700635] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1478700649] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1478700665] HOST ALERT: hallefirst2;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1478700695] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1478700725] HOST ALERT: hallefirst2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1478700815] HOST ALERT: hallefirst2;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1478700820] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1478700850] HOST ALERT: hallefirst2;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1478700880] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1478700910] HOST ALERT: hallefirst2;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1478700940] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1478700970] HOST ALERT: hallefirst2;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1478701016] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6869 found since last run. oneLevelSearchOps OK. searchOps is too high! 7253 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478701019] HOST ALERT: sillpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1478701060] HOST ALERT: hallefirst2;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1478701083] HOST ALERT: sillpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.40 ms
[1478701136] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478701229] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7913 found since last run. oneLevelSearchOps OK. searchOps is too high! 8139 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478701349] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478701413] HOST ALERT: sillpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1478701503] HOST ALERT: sillpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1478701583] SERVICE ALERT: sillpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1478701587] HOST ALERT: sillpk;UP;SOFT;3;PING WARNING - Packet loss = 80%, RTA = 0.45 ms
[1478701587] HOST FLAPPING ALERT: sillpk;STARTED; Host appears to have started flapping (23.2% change > 20.0% threshold)
[1478701736] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7597 found since last run. oneLevelSearchOps OK. searchOps is too high! 8039 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478701802] HOST FLAPPING ALERT: sillpk;STARTED; Host appears to have started flapping (22.1% change > 20.0% threshold)
[1478701802] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (39.5% change >= 20.0% threshold)
[1478701802] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1478701802] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1478701802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1478701802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1478701802] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1478701802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1478701802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1478701802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1478701802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1478701802] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1478701802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1478701802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1478701802] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1478701802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1478701802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1478701802] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1478701814] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478701856] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478701864] SERVICE ALERT: ADFS;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1478701864] SERVICE ALERT: ADFS;C:\ Drive Space;OK;HARD;3;c:\ - total: 79.51 Gb - used: 18.69 Gb (24%) - free 60.81 Gb (76%)
[1478701868] HOST ALERT: ADFS;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1478701934] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1478701949] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 12029 found since last run. oneLevelSearchOps OK. searchOps is too high! 12246 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478701951] SERVICE ALERT: ADFS;Service: RPC;OK;HARD;3;RpcSs: Started
[1478701951] SERVICE ALERT: ADFS;CPU Load;OK;HARD;3;CPU Load 2% (15 min average)
[1478702041] SERVICE ALERT: ADFS;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1478702041] SERVICE ALERT: ADFS;Memory Usage;OK;HARD;3;Memory usage: total:10239.64 MB - used: 1702.83 MB (17%) - free: 8536.80 MB (83%)
[1478702054] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1478702069] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478702122] SERVICE ALERT: ADFS;Service: WMI;OK;HARD;3;Winmgmt: Started
[1478702122] SERVICE ALERT: ADFS;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1478702189] SERVICE ALERT: sillpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1478702197] SERVICE ALERT: ADFS;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1478702217] SERVICE ALERT: ADFS;Service: Windows Time;OK;HARD;3;W32Time: Started
[1478702217] SERVICE ALERT: ADFS;VMWare Tools;OK;HARD;3;VMTools: Started
[1478702217] SERVICE ALERT: ADFS;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1478702249] SERVICE ALERT: sillpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1478702300] SERVICE ALERT: ADFS;Uptime;OK;HARD;3;System Uptime - 4 day(s) 23 hour(s) 46 minute(s)
[1478702301] SERVICE ALERT: ADFS;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1478702309] SERVICE ALERT: sillpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1478702669] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10540 found since last run. oneLevelSearchOps OK. searchOps is too high! 10749 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478702789] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478702823] HOST FLAPPING ALERT: sillpk;STOPPED; Host appears to have stopped flapping (4.0% change < 5.0% threshold)
[1478703254] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 32% (15 min average)
[1478703389] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 12557 found since last run. oneLevelSearchOps OK. searchOps is too high! 12807 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478703509] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478704110] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9251 found since last run. oneLevelSearchOps OK. searchOps is too high! 9446 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478704229] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478704454] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1478704574] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 89% (15 min average)
[1478704828] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8143 found since last run. oneLevelSearchOps OK. searchOps is too high! 8341 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478704856] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6019 found since last run. oneLevelSearchOps OK. searchOps is too high! 6418 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478704948] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478704976] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478705548] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10069 found since last run. oneLevelSearchOps OK. searchOps is too high! 10306 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478705576] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5837 found since last run. oneLevelSearchOps OK. searchOps is too high! 6415 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478705668] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478705696] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478705774] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1478705894] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1478706014] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1478706268] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5625 found since last run. oneLevelSearchOps OK. searchOps is too high! 5826 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478706296] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 13630 found since last run. oneLevelSearchOps OK. searchOps is too high! 14350 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478706388] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478706416] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478706614] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1478706988] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7294 found since last run. oneLevelSearchOps OK. searchOps is too high! 7514 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478707016] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 19460 found since last run. oneLevelSearchOps OK. searchOps is too high! 20084 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478707108] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478707137] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478707214] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1478707708] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9908 found since last run. oneLevelSearchOps OK. searchOps is too high! 10150 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478707736] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 11968 found since last run. oneLevelSearchOps OK. searchOps is too high! 12329 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478707828] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478707855] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478707979] SERVICE ALERT: prayharroldpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1478708039] SERVICE ALERT: prayharroldpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1478708099] SERVICE ALERT: prayharroldpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1478708428] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9929 found since last run. oneLevelSearchOps OK. searchOps is too high! 10137 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478708455] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5425 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478708548] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478708575] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478708699] SERVICE ALERT: prayharroldpk;Device Status;OK;HARD;3;ORDER CARTRIDGE LESS THAN 1400 PAGES 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1478709148] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8795 found since last run. oneLevelSearchOps OK. searchOps is too high! 8862 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478709175] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7255 found since last run. oneLevelSearchOps OK. searchOps is too high! 7716 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478709268] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478709295] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478709868] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10871 found since last run. oneLevelSearchOps OK. searchOps is too high! 10937 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478709895] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5463 found since last run. oneLevelSearchOps OK. searchOps is too high! 5946 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478709988] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478710015] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478710588] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9801 found since last run. oneLevelSearchOps OK. searchOps is too high! 9880 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478710615] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8604 found since last run. oneLevelSearchOps OK. searchOps is too high! 9162 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478710708] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478710735] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478711308] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 12400 found since last run. oneLevelSearchOps OK. searchOps is too high! 12552 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478711335] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10751 found since last run. oneLevelSearchOps OK. searchOps is too high! 11415 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478711428] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478711455] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478712014] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1478712028] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 13514 found since last run. oneLevelSearchOps OK. searchOps is too high! 13673 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478712055] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 19696 found since last run. oneLevelSearchOps OK. searchOps is too high! 20163 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478712148] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478712175] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478712558] SERVICE ALERT: strongpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1478712603] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (33.8% change >= 20.0% threshold)
[1478712603] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1478712603] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1478712603] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1478712603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1478712603] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1478712603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1478712603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1478712603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1478712603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1478712603] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1478712603] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1478712603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1478712603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1478712603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1478712603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1478712603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1478712617] SERVICE ALERT: strongpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1478712667] SERVICE ALERT: ADFS;Service: NetLogin;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478712667] SERVICE ALERT: ADFS;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478712670] HOST ALERT: ADFS;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.57)
[1478712678] SERVICE ALERT: strongpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1478712733] HOST ALERT: ADFS;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.4.57)
[1478712748] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 11377 found since last run. oneLevelSearchOps OK. searchOps is too high! 11565 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478712754] SERVICE ALERT: ADFS;CPU Load;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478712754] SERVICE ALERT: ADFS;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478712757] HOST ALERT: ADFS;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.4.57)
[1478712775] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 17793 found since last run. oneLevelSearchOps OK. searchOps is too high! 18204 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478712787] SERVICE ALERT: ADFS;Service: NetLogin;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478712787] SERVICE ALERT: ADFS;C:\ Drive Space;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478712790] HOST ALERT: ADFS;DOWN;SOFT;4;CRITICAL - Host Unreachable (164.76.4.57)
[1478712844] SERVICE ALERT: ADFS;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478712844] SERVICE ALERT: ADFS;Service: User Profile;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478712847] HOST ALERT: ADFS;DOWN;SOFT;5;CRITICAL - Host Unreachable (164.76.4.57)
[1478712868] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478712874] SERVICE ALERT: ADFS;Service: RPC;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478712874] SERVICE ALERT: ADFS;CPU Load;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478712877] HOST ALERT: ADFS;DOWN;SOFT;6;CRITICAL - Host Unreachable (164.76.4.57)
[1478712895] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478712907] SERVICE ALERT: ADFS;C:\ Drive Space;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478712907] SERVICE ALERT: ADFS;Service: NetLogin;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478712910] HOST ALERT: ADFS;DOWN;SOFT;7;CRITICAL - Host Unreachable (164.76.4.57)
[1478712925] SERVICE ALERT: ADFS;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478712925] SERVICE ALERT: ADFS;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478712928] HOST ALERT: ADFS;DOWN;SOFT;8;CRITICAL - Host Unreachable (164.76.4.57)
[1478712964] SERVICE ALERT: ADFS;Service: User Profile;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478712964] SERVICE ALERT: ADFS;Memory Usage;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478712967] HOST ALERT: ADFS;DOWN;SOFT;9;CRITICAL - Host Unreachable (164.76.4.57)
[1478712994] SERVICE ALERT: ADFS;CPU Load;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478712994] SERVICE ALERT: ADFS;Service: RPC;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478712997] HOST ALERT: ADFS;DOWN;HARD;10;CRITICAL - Host Unreachable (164.76.4.57)
[1478713000] SERVICE ALERT: ADFS;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478713020] SERVICE ALERT: ADFS;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478713020] SERVICE ALERT: ADFS;VMWare Tools;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478713020] SERVICE ALERT: ADFS;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478713045] SERVICE ALERT: ADFS;Service: WMI;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478713045] SERVICE ALERT: ADFS;NSClient++ Version;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478713084] SERVICE ALERT: ADFS;Service: User Profile;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478713084] SERVICE ALERT: ADFS;Memory Usage;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478713103] SERVICE ALERT: ADFS;Uptime;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478713103] SERVICE ALERT: ADFS;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.4.57 and port 12489: No route to host
[1478713120] SERVICE ALERT: ADFS;Service: Group Policy;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478713140] SERVICE ALERT: ADFS;Service: DNS Client;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478713140] SERVICE ALERT: ADFS;Service: Windows Time;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478713140] SERVICE ALERT: ADFS;VMWare Tools;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478713165] SERVICE ALERT: ADFS;Service: WMI;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478713165] SERVICE ALERT: ADFS;NSClient++ Version;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478713214] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 65% (15 min average)
[1478713223] SERVICE ALERT: ADFS;Service: Firewall;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478713223] SERVICE ALERT: ADFS;Uptime;CRITICAL;SOFT;2;connect to address 164.76.4.57 and port 12489: No route to host
[1478713240] SERVICE ALERT: ADFS;Service: Group Policy;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478713260] SERVICE ALERT: ADFS;Service: Windows Time;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478713260] SERVICE ALERT: ADFS;VMWare Tools;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478713260] SERVICE ALERT: ADFS;Service: DNS Client;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478713343] SERVICE ALERT: ADFS;Uptime;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478713343] SERVICE ALERT: ADFS;Service: Firewall;CRITICAL;HARD;3;connect to address 164.76.4.57 and port 12489: No route to host
[1478713468] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9513 found since last run. oneLevelSearchOps OK. searchOps is too high! 9685 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478713495] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8188 found since last run. oneLevelSearchOps OK. searchOps is too high! 8727 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478713588] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478713615] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478713738] HOST ALERT: sillpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1478713802] HOST ALERT: sillpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.46 ms
[1478713814] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478713934] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 91% (15 min average)
[1478714188] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5831 found since last run. oneLevelSearchOps OK. searchOps is too high! 5984 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478714215] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6503 found since last run. oneLevelSearchOps OK. searchOps is too high! 6897 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478714249] HOST ALERT: sillpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1478714308] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478714335] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478714338] HOST ALERT: sillpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1478714428] HOST ALERT: sillpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1478714448] SERVICE ALERT: porterpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1478714475] SERVICE ALERT: porterpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1478714508] SERVICE ALERT: porterpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1478714518] HOST ALERT: sillpk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1478714534] SERVICE ALERT: porterpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1478714534] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1478714568] SERVICE ALERT: porterpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1478714581] SERVICE ALERT: porterpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1478714594] SERVICE ALERT: porterpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1478714608] HOST ALERT: sillpk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1478714641] SERVICE ALERT: porterpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1478714654] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1478714698] HOST ALERT: sillpk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1478714701] SERVICE ALERT: porterpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1478714728] HOST ALERT: halleground1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1478714774] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1478714788] HOST ALERT: sillpk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1478714818] HOST ALERT: halleground1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1478714849] HOST ALERT: sillpk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1478714908] HOST ALERT: halleground1;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1478714908] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8926 found since last run. oneLevelSearchOps OK. searchOps is too high! 9112 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478714935] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9002 found since last run. oneLevelSearchOps OK. searchOps is too high! 9422 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478714938] HOST ALERT: sillpk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1478714998] HOST ALERT: halleground1;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1478715028] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478715028] HOST ALERT: sillpk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1478715055] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478715058] HOST ALERT: halleground1;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1478715148] HOST ALERT: halleground1;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1478715238] HOST ALERT: halleground1;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1478715328] HOST ALERT: halleground1;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1478715418] HOST ALERT: halleground1;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1478715508] HOST ALERT: halleground1;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1478715628] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9196 found since last run. oneLevelSearchOps OK. searchOps is too high! 9345 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478715694] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1478715748] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478715754] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1478715813] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1478715974] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 79% (15 min average)
[1478716255] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8960 found since last run. oneLevelSearchOps OK. searchOps is too high! 9236 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478716348] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9956 found since last run. oneLevelSearchOps OK. searchOps is too high! 10122 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478716375] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478716468] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478716554] HOST ALERT: halleground2;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1478716574] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478716688] SERVICE ALERT: halleground2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1478716693] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1478716782] SERVICE ALERT: halleground2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1478716813] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 73% (15 min average)
[1478716814] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 74% - OK!
[1478716975] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10244 found since last run. oneLevelSearchOps OK. searchOps is too high! 10489 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478717068] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 14051 found since last run. oneLevelSearchOps OK. searchOps is too high! 14226 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478717095] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478717188] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478717413] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478717477] SERVICE ALERT: strongpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1478717533] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1478717652] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1478717695] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 18652 found since last run. oneLevelSearchOps OK. searchOps is too high! 18971 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478717788] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 17167 found since last run. oneLevelSearchOps OK. searchOps is too high! 17347 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478717815] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478717908] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478718162] SERVICE ALERT: porterpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 3% - OK!
[1478718188] SERVICE ALERT: porterpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1478718296] SERVICE ALERT: porterpk;Device Status;OK;HARD;3;ORDER CARTRIDGE LESS THAN 700 PAGES 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1478718415] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8826 found since last run. oneLevelSearchOps OK. searchOps is too high! 9127 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478718508] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 13815 found since last run. oneLevelSearchOps OK. searchOps is too high! 14134 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478718535] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478718629] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478718784] SERVICE ALERT: sillpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 10% - OK!
[1478718788] HOST ALERT: sillpk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.53 ms
[1478718851] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 91% (15 min average)
[1478718923] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1478719013] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Ready
[1478719043] SERVICE ALERT: BOEAPPTEST;CPU Load;OK;SOFT;2;CPU Load 1% (15 min average)
[1478719102] SERVICE ALERT: sillpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1478719135] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6288 found since last run. oneLevelSearchOps OK. searchOps is too high! 6669 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478719228] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 11469 found since last run. oneLevelSearchOps OK. searchOps is too high! 11811 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478719255] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478719347] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478719782] SERVICE ALERT: halleground2;Device Status;OK;HARD;3;INPUT PAPER PATH OPEN 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1478719855] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5900 found since last run. oneLevelSearchOps OK. searchOps is too high! 6143 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478719947] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6473 found since last run. oneLevelSearchOps OK. searchOps is too high! 6787 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478719975] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478720048] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1478720067] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478720168] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 80% (15 min average)
[1478720575] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5046 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478720614] SERVICE ALERT: marshallpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1478720667] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8354 found since last run. oneLevelSearchOps OK. searchOps is too high! 8651 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478720695] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478720768] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478720787] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478720888] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1478721008] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1478721295] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8509 found since last run. oneLevelSearchOps OK. searchOps is too high! 8765 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478721387] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8679 found since last run. oneLevelSearchOps OK. searchOps is too high! 9029 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478721415] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478721507] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478721710] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1478721770] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 7% - OK!
[1478722015] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5743 found since last run. oneLevelSearchOps OK. searchOps is too high! 6233 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478722107] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9783 found since last run. oneLevelSearchOps OK. searchOps is too high! 10060 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478722135] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478722193] SERVICE ALERT: hallefirst3;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1478722209] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1478722227] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478722253] SERVICE ALERT: hallefirst3;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1478722313] SERVICE ALERT: hallefirst3;Device Status;OK;SOFT;3;CLOSE RIGHT DOOR  46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1478722735] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5319 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478722827] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 15434 found since last run. oneLevelSearchOps OK. searchOps is too high! 15649 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478722855] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478722947] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478723547] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 14259 found since last run. oneLevelSearchOps OK. searchOps is too high! 14450 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478723667] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478724267] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 11626 found since last run. oneLevelSearchOps OK. searchOps is too high! 11810 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478724387] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478724655] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6279 found since last run. oneLevelSearchOps OK. searchOps is too high! 6621 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478724775] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478724987] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8513 found since last run. oneLevelSearchOps OK. searchOps is too high! 8794 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478725107] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478725375] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8736 found since last run. oneLevelSearchOps OK. searchOps is too high! 9073 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478725495] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478725806] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1478725926] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1478726046] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 51% (15 min average)
[1478726095] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7965 found since last run. oneLevelSearchOps OK. searchOps is too high! 8342 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478726215] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478726307] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5315 found since last run. oneLevelSearchOps OK. searchOps is too high! 5499 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478726427] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478726646] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1478726698] SERVICE ALERT: prayharroldpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1478726758] SERVICE ALERT: prayharroldpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1478726766] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1478726815] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6396 found since last run. oneLevelSearchOps OK. searchOps is too high! 6770 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478726818] SERVICE ALERT: prayharroldpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1478726886] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 89% (15 min average)
[1478726935] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478727027] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5471 found since last run. oneLevelSearchOps OK. searchOps is too high! 5551 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478727147] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478727419] SERVICE ALERT: prayharroldpk;Device Status;OK;HARD;3;ORDER CARTRIDGE LESS THAN 1400 PAGES 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1478727486] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1478727535] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10759 found since last run. oneLevelSearchOps OK. searchOps is too high! 11064 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478727606] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1478727655] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478727726] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1478727747] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8517 found since last run. oneLevelSearchOps OK. searchOps is too high! 8580 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478727748] SERVICE ALERT: hallesouthcommons;Black Cartridge HP CC364X;CRITICAL;HARD;3;Black Cartridge HP CC364X is at 1% - CRITICAL!
[1478727867] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478728255] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5161 found since last run. oneLevelSearchOps OK. searchOps is too high! 5539 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478728375] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478728467] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9137 found since last run. oneLevelSearchOps OK. searchOps is too high! 9237 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478728587] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478728975] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5068 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478729095] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478729187] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8623 found since last run. oneLevelSearchOps OK. searchOps is too high! 8691 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478729307] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478729695] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 11540 found since last run. oneLevelSearchOps OK. searchOps is too high! 11817 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478729784] SERVICE ALERT: hallefirst3;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1478729815] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478729833] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1478729844] SERVICE ALERT: hallefirst3;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1478729893] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1478729904] SERVICE ALERT: hallefirst3;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1478729954] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1478730415] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8782 found since last run. oneLevelSearchOps OK. searchOps is too high! 9016 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478730507] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8564 found since last run. oneLevelSearchOps OK. searchOps is too high! 8625 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478730535] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478730627] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478731135] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9104 found since last run. oneLevelSearchOps OK. searchOps is too high! 9504 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478731227] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5509 found since last run. oneLevelSearchOps OK. searchOps is too high! 5570 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478731255] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478731347] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478731704] SERVICE ALERT: hallefirst3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1478731855] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5523 found since last run. oneLevelSearchOps OK. searchOps is too high! 5775 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478731925] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 91% (15 min average)
[1478731975] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478732525] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1478732547] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6512 found since last run. oneLevelSearchOps OK. searchOps is too high! 6560 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478732576] SERVICE ALERT: cava2;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1478732645] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1478732667] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478732696] SERVICE ALERT: cava2;CPU Load;OK;SOFT;2;CPU Load 70% (15 min average)
[1478732765] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 92% (15 min average)
[1478733175] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5248 found since last run. oneLevelSearchOps OK. searchOps is too high! 5634 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478733267] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6586 found since last run. oneLevelSearchOps OK. searchOps is too high! 6632 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478733295] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478733387] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478733553] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;CRITICAL;HARD;3;Black Print Cartridge HP C8543X is at 1% - CRITICAL!
[1478733895] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5023 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478733965] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478733987] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9420 found since last run. oneLevelSearchOps OK. searchOps is too high! 9468 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478734015] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478734085] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1478734107] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478734205] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1478734615] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5063 found since last run. oneLevelSearchOps OK. searchOps is too high! 5484 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478734707] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7908 found since last run. oneLevelSearchOps OK. searchOps is too high! 7954 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478734736] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478734827] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478735062] SERVICE ALERT: prayharroldpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1478735122] SERVICE ALERT: prayharroldpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1478735182] SERVICE ALERT: prayharroldpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1478735335] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6292 found since last run. oneLevelSearchOps OK. searchOps is too high! 6645 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478735427] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8088 found since last run. oneLevelSearchOps OK. searchOps is too high! 8134 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478735454] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478735547] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478735953] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1478736054] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8633 found since last run. oneLevelSearchOps OK. searchOps is too high! 9086 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478736113] SERVICE ALERT: hallefirst3;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1478736147] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7798 found since last run. oneLevelSearchOps OK. searchOps is too high! 7850 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478736161] SERVICE ALERT: porterpk;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1478736173] SERVICE ALERT: hallefirst3;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1478736174] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478736221] SERVICE ALERT: porterpk;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1478736233] SERVICE ALERT: hallefirst3;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1478736267] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478736281] SERVICE ALERT: porterpk;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1478736381] SERVICE ALERT: prayharroldpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1478736559] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1478736589] HOST ALERT: hallefirst3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1478736619] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1478736648] SERVICE ALERT: owen101lab3;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1478736649] HOST ALERT: hallefirst3;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1478736679] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1478736679] SERVICE FLAPPING ALERT: hallefirst3;Black Print Cartridge HP C8543X;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1478736708] SERVICE ALERT: owen101lab3;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1478736709] HOST ALERT: hallefirst3;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1478736768] SERVICE ALERT: owen101lab3;Device Status;OK;SOFT;3;Clearing paper path
[1478736774] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6669 found since last run. oneLevelSearchOps OK. searchOps is too high! 7030 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478736799] HOST ALERT: hallefirst3;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1478736839] SERVICE ALERT: hallefirst3;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1478736867] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6865 found since last run. oneLevelSearchOps OK. searchOps is too high! 6911 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478736869] HOST ALERT: hallefirst3;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1478736894] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478736959] HOST ALERT: hallefirst3;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1478736987] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478737049] HOST ALERT: hallefirst3;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1478737109] SERVICE ALERT: hallefirst3;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1478737139] HOST ALERT: hallefirst3;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1478737169] SERVICE ALERT: hallefirst3;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1478737199] HOST ALERT: hallefirst3;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1478737229] SERVICE ALERT: hallefirst3;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1478737259] HOST ALERT: hallefirst3;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1478737369] SERVICE ALERT: owen101lab3;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1478737384] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.138)
[1478737384] SERVICE ALERT: SNAP;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1478737384] SERVICE ALERT: SNAP;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1478737394] HOST ALERT: SNAP;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1478737421] SERVICE ALERT: SNAP;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1478737428] SERVICE ALERT: owen101lab3;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1478737433] SERVICE ALERT: SNAP;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1478737488] SERVICE ALERT: owen101lab3;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1478737494] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5339 found since last run. oneLevelSearchOps OK. searchOps is too high! 5896 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478737515] SERVICE ALERT: SNAP;NSClient++ Version;CRITICAL;SOFT;2;connect to address 164.76.4.138 and port 12489: No route to host
[1478737515] SERVICE ALERT: SNAP;Memory Usage;CRITICAL;SOFT;2;connect to address 164.76.4.138 and port 12489: No route to host
[1478737518] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.138)
[1478737524] SERVICE ALERT: SNAP;Service: User Profile;CRITICAL;SOFT;2;connect to address 164.76.4.138 and port 12489: No route to host
[1478737536] SERVICE ALERT: SNAP;Service: Firewall;CRITICAL;SOFT;2;connect to address 164.76.4.138 and port 12489: No route to host
[1478737539] HOST ALERT: SNAP;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.4.138)
[1478737587] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6469 found since last run. oneLevelSearchOps OK. searchOps is too high! 6516 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478737602] HOST ALERT: SNAP;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.4.138)
[1478737614] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478737620] SERVICE ALERT: SNAP;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1478737623] HOST ALERT: SNAP;DOWN;SOFT;4;CRITICAL - Host Unreachable (164.76.4.138)
[1478737626] SERVICE ALERT: SNAP;NSClient++ Version;CRITICAL;HARD;3;connect to address 164.76.4.138 and port 12489: No route to host
[1478737626] SERVICE ALERT: SNAP;Memory Usage;CRITICAL;HARD;3;connect to address 164.76.4.138 and port 12489: No route to host
[1478737641] SERVICE ALERT: SNAP;Service: User Profile;OK;SOFT;3;ProfSvc: Started
[1478737645] HOST ALERT: SNAP;UP;SOFT;5;PING OK - Packet loss = 0%, RTA = 0.36 ms
[1478737673] SERVICE ALERT: SNAP;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1478737676] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.138)
[1478737707] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478737737] SERVICE ALERT: SNAP;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1478737740] HOST ALERT: SNAP;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.53 ms
[1478737740] HOST FLAPPING ALERT: SNAP;STARTED; Host appears to have started flapping (22.4% change > 20.0% threshold)
[1478737995] SERVICE ALERT: SNAP;Service: NetLogin;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1478737995] SERVICE ALERT: SNAP;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1478737995] SERVICE ALERT: SNAP;Uptime;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1478737998] SERVICE ALERT: SNAP;CPU Load;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1478737998] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.138)
[1478738034] HOST ALERT: SNAP;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.40 ms
[1478738115] SERVICE ALERT: SNAP;Service: NetLogin;CRITICAL;SOFT;2;connect to address 164.76.4.138 and port 12489: No route to host
[1478738115] SERVICE ALERT: SNAP;Uptime;CRITICAL;SOFT;2;connect to address 164.76.4.138 and port 12489: No route to host
[1478738115] SERVICE ALERT: SNAP;CPU Load;CRITICAL;SOFT;2;connect to address 164.76.4.138 and port 12489: No route to host
[1478738118] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.138)
[1478738123] SERVICE ALERT: SNAP;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1478738214] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7187 found since last run. oneLevelSearchOps OK. searchOps is too high! 7605 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478738223] SERVICE ALERT: SNAP;Service: RPC;OK;SOFT;3;RpcSs: Started
[1478738223] SERVICE ALERT: SNAP;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1478738223] SERVICE ALERT: SNAP;Memory Usage;OK;HARD;3;Memory usage: total:16362.75 MB - used: 2201.82 MB (13%) - free: 14160.93 MB (87%)
[1478738250] SERVICE ALERT: SNAP;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1478738250] SERVICE ALERT: SNAP;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1478738253] HOST ALERT: SNAP;UP;SOFT;2;PING WARNING - Packet loss = 86%, RTA = 0.39 ms
[1478738254] SERVICE ALERT: SNAP;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1478738261] SERVICE ALERT: SNAP;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1478738262] SERVICE ALERT: SNAP;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1478738334] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478738349] SERVICE ALERT: SNAP;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1478738352] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.138)
[1478738364] SERVICE ALERT: SNAP;Service: User Profile;CRITICAL;SOFT;2;connect to address 164.76.4.138 and port 12489: No route to host
[1478738367] HOST ALERT: SNAP;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.4.138)
[1478738405] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1478738430] HOST ALERT: SNAP;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.4.138)
[1478738460] SERVICE ALERT: SNAP;Service: Group Policy;CRITICAL;SOFT;2;connect to address 164.76.4.138 and port 12489: No route to host
[1478738463] HOST ALERT: SNAP;DOWN;SOFT;4;CRITICAL - Host Unreachable (164.76.4.138)
[1478738484] SERVICE ALERT: SNAP;Service: User Profile;CRITICAL;HARD;3;connect to address 164.76.4.138 and port 12489: No route to host
[1478738487] HOST ALERT: SNAP;DOWN;SOFT;5;CRITICAL - Host Unreachable (164.76.4.138)
[1478738551] HOST ALERT: SNAP;UP;SOFT;6;PING OK - Packet loss = 0%, RTA = 0.45 ms
[1478738577] SERVICE ALERT: SNAP;Service: Group Policy;OK;SOFT;3;gpsvc: Started
[1478738838] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.138)
[1478738843] SERVICE ALERT: SNAP;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1478738843] SERVICE ALERT: SNAP;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1478738843] SERVICE ALERT: SNAP;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1478738856] SERVICE ALERT: SNAP;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1478738859] HOST ALERT: SNAP;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.4.138)
[1478738922] HOST ALERT: SNAP;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.45 ms
[1478738943] SERVICE ALERT: SNAP;Service: RPC;OK;SOFT;2;RpcSs: Started
[1478738943] SERVICE ALERT: SNAP;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1478738943] SERVICE ALERT: SNAP;Memory Usage;OK;SOFT;2;Memory usage: total:16362.75 MB - used: 2212.55 MB (14%) - free: 14150.19 MB (86%)
[1478738973] SERVICE ALERT: SNAP;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1478739005] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1478739081] SERVICE ALERT: SNAP;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1478739125] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 81% (15 min average)
[1478739429] SERVICE ALERT: SNAP;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1478739432] SERVICE ALERT: SNAP;Uptime;OK;HARD;3;System Uptime - 19 day(s) 2 hour(s) 4 minute(s)
[1478739433] SERVICE ALERT: SNAP;CPU Load;OK;HARD;3;CPU Load 0% (15 min average)
[1478739433] HOST FLAPPING ALERT: SNAP;STOPPED; Host appears to have stopped flapping (4.4% change < 5.0% threshold)
[1478739725] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1478739845] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1478739965] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1478740107] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6696 found since last run. oneLevelSearchOps OK. searchOps is too high! 6743 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478740227] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478740471] SERVICE ALERT: SNAP;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1478740474] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.138)
[1478740480] SERVICE ALERT: SNAP;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1478740565] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 90% (15 min average)
[1478740571] SERVICE ALERT: SNAP;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1478740575] HOST ALERT: SNAP;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1478740580] SERVICE ALERT: SNAP;C:\ Drive Space;OK;SOFT;2;c:\ - total: 1116.65 Gb - used: 36.91 Gb (3%) - free 1079.74 Gb (97%)
[1478740827] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6579 found since last run. oneLevelSearchOps OK. searchOps is too high! 6625 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478740947] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478741547] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7677 found since last run. oneLevelSearchOps OK. searchOps is too high! 7728 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478741668] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478741765] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1478741885] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1478742534] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6018 found since last run. oneLevelSearchOps OK. searchOps is too high! 6368 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478742654] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478743085] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1478743205] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1478743254] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5246 found since last run. oneLevelSearchOps OK. searchOps is too high! 5638 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478743325] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1478743374] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478743974] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5136 found since last run. oneLevelSearchOps OK. searchOps is too high! 5412 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478744094] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478744664] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5155 found since last run. oneLevelSearchOps OK. searchOps is too high! 5199 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478744783] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478745382] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7090 found since last run. oneLevelSearchOps OK. searchOps is too high! 7143 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478745502] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478745840] SERVICE ALERT: SNAP;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1478745843] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.138)
[1478745907] HOST ALERT: SNAP;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.40 ms
[1478745954] SERVICE ALERT: SNAP;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1478746040] SERVICE ALERT: SNAP;Uptime;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1478746040] SERVICE ALERT: SNAP;Service: NetLogin;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1478746043] SERVICE ALERT: SNAP;CPU Load;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1478746043] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.138)
[1478746101] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5013 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1478746147] HOST ALERT: SNAP;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1478746147] HOST FLAPPING ALERT: SNAP;STARTED; Host appears to have started flapping (21.4% change > 20.0% threshold)
[1478746148] SERVICE ALERT: SNAP;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1478746148] SERVICE ALERT: SNAP;Uptime;OK;SOFT;2;System Uptime - 19 day(s) 3 hour(s) 56 minute(s)
[1478746153] SERVICE ALERT: SNAP;CPU Load;OK;SOFT;2;CPU Load 0% (15 min average)
[1478746221] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1478746325] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1478746748] HOST FLAPPING ALERT: SNAP;STOPPED; Host appears to have stopped flapping (4.1% change < 5.0% threshold)
[1478746925] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1478747045] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1478747165] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 59% (15 min average)
[1478747765] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478747885] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1478748005] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1478748079] SERVICE FLAPPING ALERT: hallefirst3;Black Print Cartridge HP C8543X;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1478748600] HOST DOWNTIME ALERT: EMUSQLEXPRESS;STARTED; Host has entered a period of scheduled downtime
[1478748600] HOST DOWNTIME ALERT: FTK;STARTED; Host has entered a period of scheduled downtime
[1478748600] HOST DOWNTIME ALERT: 4WINDSMANAGER;STARTED; Host has entered a period of scheduled downtime
[1478748600] HOST DOWNTIME ALERT: HDBACKUPS;STARTED; Host has entered a period of scheduled downtime
[1478748600] HOST DOWNTIME ALERT: IMAGES;STARTED; Host has entered a period of scheduled downtime
[1478748600] HOST DOWNTIME ALERT: AD-RMAD;STARTED; Host has entered a period of scheduled downtime
[1478748600] HOST DOWNTIME ALERT: LANSWEEPER;STARTED; Host has entered a period of scheduled downtime
[1478748600] HOST DOWNTIME ALERT: MBAM;STARTED; Host has entered a period of scheduled downtime
[1478748600] HOST DOWNTIME ALERT: ADAUTH-RMAD;STARTED; Host has entered a period of scheduled downtime
[1478748600] HOST DOWNTIME ALERT: MBAMDB;STARTED; Host has entered a period of scheduled downtime
[1478748600] HOST DOWNTIME ALERT: NESSUS;STARTED; Host has entered a period of scheduled downtime
[1478748600] HOST DOWNTIME ALERT: ADCA1;STARTED; Host has entered a period of scheduled downtime
[1478748600] HOST DOWNTIME ALERT: NETENG-PA;STARTED; Host has entered a period of scheduled downtime
[1478748600] HOST DOWNTIME ALERT: OFFICEKMS;STARTED; Host has entered a period of scheduled downtime
[1478748600] HOST DOWNTIME ALERT: XDM;STARTED; Host has entered a period of scheduled downtime
[1478748600] HOST DOWNTIME ALERT: ADCA2;STARTED; Host has entered a period of scheduled downtime
[1478748600] HOST DOWNTIME ALERT: OPIE;STARTED; Host has entered a period of scheduled downtime
[1478748600] HOST DOWNTIME ALERT: PNCAPP-2012;STARTED; Host has entered a period of scheduled downtime
[1478748600] HOST DOWNTIME ALERT: ADFS;STARTED; Host has entered a period of scheduled downtime
[1478748600] HOST DOWNTIME ALERT: PNCDB-2012;STARTED; Host has entered a period of scheduled downtime
[1478748600] HOST DOWNTIME ALERT: PNCWEB-2012;STARTED; Host has entered a period of scheduled downtime
[1478748600] HOST DOWNTIME ALERT: ADWSUSCLIENTS;STARTED; Host has entered a period of scheduled downtime
[1478748600] HOST DOWNTIME ALERT: PROCARE;STARTED; Host has entered a period of scheduled downtime
[1478748600] HOST DOWNTIME ALERT: R25-ALPHA2012;STARTED; Host has entered a period of scheduled downtime
[1478748600] HOST DOWNTIME ALERT: ADWU1;STARTED; Host has entered a period of scheduled downtime
[1478748600] HOST DOWNTIME ALERT: R25-DELTA2012;STARTED; Host has entered a period of scheduled downtime
[1478748600] HOST DOWNTIME ALERT: R25-INTERFACE;STARTED; Host has entered a period of scheduled downtime
[1478748600] HOST DOWNTIME ALERT: CBORD;STARTED; Host has entered a period of scheduled downtime
[1478748600] HOST DOWNTIME ALERT: R25-WEBSERVICES;STARTED; Host has entered a period of scheduled downtime
[1478748600] HOST DOWNTIME ALERT: R25-WEBVIEWER;STARTED; Host has entered a period of scheduled downtime
[1478748600] HOST DOWNTIME ALERT: CEDARON;STARTED; Host has entered a period of scheduled downtime
[1478748600] HOST DOWNTIME ALERT: R25-ZETA2012;STARTED; Host has entered a period of scheduled downtime
[1478748600] HOST DOWNTIME ALERT: REFWIN;STARTED; Host has entered a period of scheduled downtime
[1478748600] HOST DOWNTIME ALERT: CLIENTLANSWEEP;STARTED; Host has entered a period of scheduled downtime
[1478748600] HOST DOWNTIME ALERT: SANDBOX-DB;STARTED; Host has entered a period of scheduled downtime
[1478748600] HOST DOWNTIME ALERT: SCCM-PRIM01;STARTED; Host has entered a period of scheduled downtime
[1478748600] HOST DOWNTIME ALERT: COE-WEBDAV;STARTED; Host has entered a period of scheduled downtime
[1478748600] HOST DOWNTIME ALERT: SP2010;STARTED; Host has entered a period of scheduled downtime
[1478748600] HOST DOWNTIME ALERT: SP2010DB;STARTED; Host has entered a period of scheduled downtime
[1478748600] HOST DOWNTIME ALERT: DIRECTACCESS;STARTED; Host has entered a period of scheduled downtime
[1478748600] HOST DOWNTIME ALERT: TABLEAU;STARTED; Host has entered a period of scheduled downtime
[1478748600] HOST DOWNTIME ALERT: WINKMS;STARTED; Host has entered a period of scheduled downtime
[1478748600] HOST DOWNTIME ALERT: TI22;STARTED; Host has entered a period of scheduled downtime
[1478748600] HOST DOWNTIME ALERT: DPS-SECURE02;STARTED; Host has entered a period of scheduled downtime
[1478748600] HOST DOWNTIME ALERT: TI22-COE;STARTED; Host has entered a period of scheduled downtime
[1478748600] HOST DOWNTIME ALERT: TI22-PSYCLINIC;STARTED; Host has entered a period of scheduled downtime
[1478748600] HOST DOWNTIME ALERT: EMU-VALUE01;STARTED; Host has entered a period of scheduled downtime
[1478748600] HOST DOWNTIME ALERT: UHSPHARM-2012;STARTED; Host has entered a period of scheduled downtime
[1478748600] HOST DOWNTIME ALERT: EMUSCRIPT;STARTED; Host has entered a period of scheduled downtime
[1478751605] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1478752205] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1478752325] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1478752445] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1478753045] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 61% (15 min average)
[1478753645] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1478753765] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 80% (15 min average)