[1486361834] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1486361903] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1486361954] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1486362013] SERVICE FLAPPING ALERT: hallefirst1;Device Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1486362024] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1486362074] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1486362143] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1486362600] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1486363943] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1486364474] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1486365074] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1486365143] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1486365527] SERVICE ALERT: FITNESSGRAM;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1486365647] SERVICE ALERT: FITNESSGRAM;CPU Load;OK;SOFT;2;CPU Load 7% (15 min average)
[1486365675] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1486365675] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1486365743] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1486365743] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1486366274] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1486366943] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1486368047] SERVICE ALERT: FITNESSGRAM;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1486368074] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1486368143] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1486368167] SERVICE ALERT: FITNESSGRAM;CPU Load;OK;SOFT;2;CPU Load 16% (15 min average)
[1486368674] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1486368794] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1486368914] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1486369343] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1486370114] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1486372514] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1486372943] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1486373402] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (44.5% change >= 20.0% threshold)
[1486373402] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (39.2% change >= 20.0% threshold)
[1486373402] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1486373402] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1486373402] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1486373402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1486373402] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1486373402] HOST DOWNTIME ALERT: SP2010;STARTED; Host has entered a period of scheduled downtime
[1486373402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1486373402] HOST DOWNTIME ALERT: SP2010DB;STARTED; Host has entered a period of scheduled downtime
[1486373402] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1486373402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1486373402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1486373402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1486373402] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1486373402] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1486373402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1486373402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1486373402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1486373402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1486373402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1486373714] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1486373834] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 16% (15 min average)
[1486375202] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (37.5% change >= 20.0% threshold)
[1486375202] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (32.6% change >= 20.0% threshold)
[1486375202] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1486375202] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1486375202] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1486375202] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1486375202] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1486375202] HOST DOWNTIME ALERT: SP2010;STARTED; Host has entered a period of scheduled downtime
[1486375202] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1486375202] HOST DOWNTIME ALERT: SP2010DB;STARTED; Host has entered a period of scheduled downtime
[1486375202] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1486375202] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1486375202] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1486375202] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1486375202] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1486375202] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1486375202] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1486375202] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1486375202] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1486375202] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1486375202] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1486375943] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1486376063] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1486379234] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1486379354] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1486379474] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1486380074] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1486380600] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service has exited from a period of scheduled downtime
[1486381274] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1486381463] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1486381874] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1486382273] SERVICE ALERT: ADRODC-DET;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1486382373] SERVICE ALERT: ADRODC-DET;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1486382474] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 76% (15 min average)
[1486383674] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1486383794] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 85% (15 min average)
[1486384394] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1486384514] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1486384634] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1486385473] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1486385533] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 8% - OK!
[1486386434] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 47% (15 min average)
[1486386972] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1486387032] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1486387092] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1486387692] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Checking printer
[1486388317] SERVICE ALERT: strongpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1486388377] SERVICE ALERT: strongpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1486388437] SERVICE ALERT: strongpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1486388872] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1486388933] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1486388992] SERVICE ALERT: halleground2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1486389365] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5248 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486389485] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486389613] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1486389673] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1486389734] SERVICE ALERT: hallefirst1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1486390085] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6902 found since last run. oneLevelSearchOps OK. searchOps is too high! 6960 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486390205] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486390333] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;CLOSE RIGHT DOOR  46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1486390382] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1486390442] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 100% - OK!
[1486390792] SERVICE ALERT: halleground2;Device Status;OK;HARD;3;Ready 164.76.106.52
[1486390805] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5702 found since last run. oneLevelSearchOps OK. searchOps is too high! 5904 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486390837] SERVICE ALERT: strongpk;Device Status;OK;HARD;3;TRAY 3 OPEN  46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1486390925] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486390932] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1486390992] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1486391052] SERVICE ALERT: hallefirst1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1486391234] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1486391354] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 79% (15 min average)
[1486391386] HOST ALERT: hallefirst1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1486391403] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (30.5% change >= 20.0% threshold)
[1486391403] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1486391403] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1486391403] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1486391403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1486391403] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1486391403] HOST DOWNTIME ALERT: SP2010;STARTED; Host has entered a period of scheduled downtime
[1486391403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1486391403] HOST DOWNTIME ALERT: SP2010DB;STARTED; Host has entered a period of scheduled downtime
[1486391403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1486391403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1486391403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1486391403] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1486391403] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1486391403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1486391403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1486391403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1486391403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1486391403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1486391448] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1486391476] HOST ALERT: hallefirst1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1486391508] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1486391525] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8916 found since last run. oneLevelSearchOps OK. searchOps is too high! 9211 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486391538] HOST ALERT: hallefirst1;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1486391568] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1486391598] HOST ALERT: hallefirst1;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1486391645] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486391648] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1486391658] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1486391658] SERVICE FLAPPING ALERT: hallefirst1;Device Status;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1486391678] HOST ALERT: hallefirst1;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1486391708] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1486391738] HOST ALERT: hallefirst1;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1486391768] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1486391798] HOST ALERT: hallefirst1;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1486391888] HOST ALERT: hallefirst1;DOWN;SOFT;8;(Host check timed out after 30.22 seconds)
[1486391978] HOST ALERT: hallefirst1;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1486392068] HOST ALERT: hallefirst1;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1486392783] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6558 found since last run. oneLevelSearchOps OK. searchOps is too high! 6564 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486392845] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7145 found since last run. oneLevelSearchOps OK. searchOps is too high! 7302 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486392903] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486392965] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486393240] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1486393299] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1486393360] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1486393503] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7540 found since last run. oneLevelSearchOps OK. searchOps is too high! 7558 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486393565] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5206 found since last run. oneLevelSearchOps OK. searchOps is too high! 5429 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486393623] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486393685] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486393754] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1486393874] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1486393959] SERVICE ALERT: crossroadspk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1486394168] SERVICE ALERT: halleground3;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1486394199] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1486394228] SERVICE ALERT: halleground3;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1486394259] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1486394285] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5573 found since last run. oneLevelSearchOps OK. searchOps is too high! 6103 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486394288] SERVICE ALERT: halleground3;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1486394319] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1486394405] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486394698] SERVICE ALERT: halleground3;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1486394757] SERVICE ALERT: halleground3;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1486394809] SERVICE ALERT: halleground1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1486394816] SERVICE ALERT: halleground3;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1486394869] SERVICE ALERT: halleground1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1486394929] SERVICE ALERT: halleground1;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1486395005] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8368 found since last run. oneLevelSearchOps OK. searchOps is too high! 8802 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486395125] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486395199] SERVICE ALERT: halleground1;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1486395227] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1486395259] SERVICE ALERT: halleground1;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1486395287] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1486395319] SERVICE ALERT: halleground1;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1486395347] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1486395423] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5661 found since last run. oneLevelSearchOps OK. searchOps is too high! 5684 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486395543] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486395725] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 36767 found since last run. oneLevelSearchOps OK. searchOps is too high! 37076 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486395845] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486396143] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9273 found since last run. oneLevelSearchOps OK. searchOps is too high! 9313 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486396263] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486396445] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 14629 found since last run. oneLevelSearchOps OK. searchOps is too high! 15006 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486396462] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1486396463] HOST ALERT: halleground1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1486396523] HOST ALERT: halleground1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.54 ms
[1486396523] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1486396565] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486396583] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1486396603] SERVICE ALERT: sillpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1486396638] SERVICE ALERT: owen101lab1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1486396662] SERVICE ALERT: sillpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1486396692] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1486396698] SERVICE ALERT: owen101lab1;Device Status;OK;SOFT;2;Open left door. Fuser is HOT.
[1486396723] SERVICE ALERT: sillpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1486396752] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1486396759] HOST ALERT: halleground1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1486396792] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1486396812] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1486396823] HOST ALERT: halleground1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1486396823] HOST FLAPPING ALERT: halleground1;STARTED; Host appears to have started flapping (23.3% change > 20.0% threshold)
[1486396852] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1486396863] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9341 found since last run. oneLevelSearchOps OK. searchOps is too high! 9495 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486396873] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1486396912] SERVICE ALERT: halleground2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1486396983] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486397165] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 13321 found since last run. oneLevelSearchOps OK. searchOps is too high! 13875 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486397285] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486397298] SERVICE ALERT: owen101lab1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1486397358] SERVICE ALERT: owen101lab1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1486397412] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Processing job from tray 4
[1486397412] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1486397419] SERVICE ALERT: owen101lab1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1486397473] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1486397512] SERVICE ALERT: halleground2;Device Status;OK;HARD;3;Open tray indicated.
[1486397512] SERVICE FLAPPING ALERT: halleground2;Device Status;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1486397593] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1486397713] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1486397885] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5882 found since last run. oneLevelSearchOps OK. searchOps is too high! 6216 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486398005] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486398018] SERVICE ALERT: owen101lab1;Device Status;OK;HARD;3;Open left door. Fuser is HOT.
[1486398113] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1486398172] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1486398232] SERVICE ALERT: halleground2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1486398267] SERVICE ALERT: hallegroundcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1486398313] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1486398321] SERVICE ALERT: hallegroundcolor;Image Fuser Kit HP 110V-CE246A;OK;SOFT;2;Image Fuser Kit HP 110V-CE246A, 220V-CE247A is at 49% - OK!
[1486398529] HOST FLAPPING ALERT: halleground1;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1486398605] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5991 found since last run. oneLevelSearchOps OK. searchOps is too high! 6326 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 1741 found since last run.
[1486398617] SERVICE ALERT: owen101lab1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1486398677] SERVICE ALERT: owen101lab1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1486398725] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486398737] SERVICE ALERT: owen101lab1;Device Status;OK;SOFT;3;Move left accessory, if attached.
[1486398783] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10222 found since last run. oneLevelSearchOps OK. searchOps is too high! 10345 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 2395 found since last run.
[1486398832] SERVICE ALERT: halleground2;Device Status;OK;HARD;3;Tray 3 PLAIN A4 To change size or 74 79 70 65 20 70 72 65 73 73 20 1E  Tray 3 PLAIN A4 To accept settings 70 72 65 73 73 20 18
[1486398903] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486398913] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1486399337] SERVICE ALERT: owen101lab1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1486399397] SERVICE ALERT: owen101lab1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1486399457] SERVICE ALERT: owen101lab1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1486399503] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8925 found since last run. oneLevelSearchOps OK. searchOps is too high! 9046 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486399623] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486399722] SERVICE ALERT: sillpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1486399925] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8688 found since last run. oneLevelSearchOps OK. searchOps is too high! 9096 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486400045] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486400053] SERVICE ALERT: eateriespk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1486400113] SERVICE ALERT: eateriespk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1486400113] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1486400173] SERVICE ALERT: eateriespk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1486400233] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 68% (15 min average)
[1486400347] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1486400377] HOST ALERT: owen101lab1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1486400406] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1486400436] HOST ALERT: owen101lab1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1486400466] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1486400483] SERVICE ALERT: owen101lab1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1486400496] HOST ALERT: owen101lab1;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1486400543] SERVICE ALERT: owen101lab1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1486400573] HOST ALERT: owen101lab1;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1486400603] SERVICE ALERT: owen101lab1;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1486400633] HOST ALERT: owen101lab1;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1486400645] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9652 found since last run. oneLevelSearchOps OK. searchOps is too high! 9759 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486400661] SERVICE ALERT: hallefirst3;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1486400663] SERVICE ALERT: owen101lab1;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1486400663] SERVICE FLAPPING ALERT: owen101lab1;Device Status;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1486400693] HOST ALERT: owen101lab1;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1486400721] SERVICE ALERT: hallefirst3;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1486400765] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486400773] SERVICE ALERT: eateriespk;Device Status;OK;HARD;3;TRAY 3 EMPTY PLAIN LETTER   Ready
[1486400781] SERVICE ALERT: hallefirst3;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1486400782] SERVICE ALERT: marshallpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1486400783] HOST ALERT: owen101lab1;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1486400823] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9165 found since last run. oneLevelSearchOps OK. searchOps is too high! 9367 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486400873] HOST ALERT: owen101lab1;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1486400943] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486400963] HOST ALERT: owen101lab1;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1486401053] HOST ALERT: owen101lab1;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1486401341] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 6% - OK!
[1486401365] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 14806 found since last run. oneLevelSearchOps OK. searchOps is too high! 14952 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486401381] SERVICE ALERT: hallefirst3;Device Status;OK;HARD;3;Open front door.
[1486401485] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486401524] SERVICE ALERT: halleground1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1486401543] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 11422 found since last run. oneLevelSearchOps OK. searchOps is too high! 11669 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486401663] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486401911] SERVICE ALERT: halleground1;Device Status;OK;HARD;3;ORDER CARTRIDGE LESS THAN 400 PAGES 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1486402085] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 19363 found since last run. oneLevelSearchOps OK. searchOps is too high! 19740 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486402202] SERVICE FLAPPING ALERT: halleground2;Device Status;STARTED; Service appears to have started flapping (25.5% change >= 20.0% threshold)
[1486402202] SERVICE FLAPPING ALERT: owen101lab1;Device Status;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1486402202] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1486402202] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1486402202] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1486402202] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1486402202] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1486402202] HOST DOWNTIME ALERT: SP2010;STARTED; Host has entered a period of scheduled downtime
[1486402202] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1486402202] HOST DOWNTIME ALERT: SP2010DB;STARTED; Host has entered a period of scheduled downtime
[1486402202] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1486402202] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1486402202] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1486402202] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1486402202] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1486402202] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1486402202] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1486402202] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1486402202] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1486402202] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1486402205] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486402263] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 11539 found since last run. oneLevelSearchOps OK. searchOps is too high! 11886 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486402383] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486402633] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1486402753] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1486402805] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8289 found since last run. oneLevelSearchOps OK. searchOps is too high! 8575 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486402873] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1486402873] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1486402889] SERVICE ALERT: hallegroundcolor;Yellow Cartridge HP CE262A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1486402925] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486402938] SERVICE ALERT: hallegroundcolor;Yellow Cartridge HP CE262A;OK;SOFT;2;Yellow Cartridge HP CE262A is at 90% - OK!
[1486402983] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5768 found since last run. oneLevelSearchOps OK. searchOps is too high! 5999 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486403031] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1486403090] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1486403103] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486403150] SERVICE ALERT: halleground2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1486403473] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1486403525] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7069 found since last run. oneLevelSearchOps OK. searchOps is too high! 7444 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486403645] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486403703] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5963 found since last run. oneLevelSearchOps OK. searchOps is too high! 6237 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486403823] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486404073] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1486404245] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7326 found since last run. oneLevelSearchOps OK. searchOps is too high! 7747 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486404365] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486404381] SERVICE ALERT: hallefirst3;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1486404423] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6363 found since last run. oneLevelSearchOps OK. searchOps is too high! 6598 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486404441] SERVICE ALERT: hallefirst3;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1486404501] SERVICE ALERT: hallefirst3;Device Status;OK;SOFT;3;Lower green handle.
[1486404544] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486404673] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1486404857] SERVICE ALERT: owen101lab1;Device Status;OK;HARD;3;CLOSE FRONT OR LEFT DOOR To continue turn off then on CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1486404861] HOST ALERT: owen101lab1;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1486404965] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6670 found since last run. oneLevelSearchOps OK. searchOps is too high! 6756 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486405085] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486405101] SERVICE ALERT: hallefirst3;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1486405143] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8634 found since last run. oneLevelSearchOps OK. searchOps is too high! 8850 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486405162] SERVICE ALERT: hallefirst3;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1486405221] SERVICE ALERT: hallefirst3;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1486405261] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 7% - OK!
[1486405262] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486405273] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1486405393] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1486405397] SERVICE ALERT: owen101lab1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1486405513] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1486405820] SERVICE ALERT: hallefirst3;Device Status;OK;HARD;3;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1486405820] SERVICE FLAPPING ALERT: hallefirst3;Device Status;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1486405862] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 11011 found since last run. oneLevelSearchOps OK. searchOps is too high! 11249 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486405982] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486406113] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1486406151] SERVICE ALERT: halleground2;Device Status;OK;HARD;3;Ready 164.76.106.52
[1486406227] SERVICE ALERT: ADDC5;GPO Replication;UNKNOWN;SOFT;1;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1486406244] SERVICE ALERT: ADDC5;GPO Replication;OK;SOFT;2;Export-CSV : The process cannot access the file 'C:\Program
[1486406285] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9860 found since last run. oneLevelSearchOps OK. searchOps is too high! 10147 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486406329] SERVICE ALERT: halleground1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1486406389] SERVICE ALERT: halleground1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1486406405] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486406426] SERVICE ALERT: hallefirst3;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1486406449] SERVICE ALERT: halleground1;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1486406456] HOST ALERT: hallefirst3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1486406486] SERVICE ALERT: hallefirst3;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1486406516] HOST ALERT: hallefirst3;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1486406539] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1486406546] SERVICE ALERT: hallefirst3;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1486406569] HOST ALERT: hallefirst3;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1486406582] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6734 found since last run. oneLevelSearchOps OK. searchOps is too high! 7007 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486406599] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1486406629] HOST ALERT: hallefirst3;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1486406659] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1486406689] HOST ALERT: hallefirst3;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1486406702] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486406713] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1486406717] SERVICE ALERT: halleground1;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1486406739] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1486406747] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1486406756] SERVICE ALERT: halleground2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1486406777] SERVICE ALERT: halleground1;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1486406779] HOST ALERT: hallefirst3;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1486406799] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1486406807] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1486406816] SERVICE ALERT: halleground2;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1486406827] SERVICE ALERT: hallefirst3;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1486406837] SERVICE ALERT: halleground1;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1486406857] HOST ALERT: hallefirst3;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1486406859] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1486406867] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1486406876] SERVICE ALERT: halleground2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1486406886] SERVICE ALERT: hallefirst3;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1486406916] HOST ALERT: hallefirst3;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1486406946] SERVICE ALERT: hallefirst3;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1486406976] HOST ALERT: hallefirst3;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1486407005] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 13965 found since last run. oneLevelSearchOps OK. searchOps is too high! 14132 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486407066] HOST ALERT: hallefirst3;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1486407125] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486407292] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1486407302] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 11917 found since last run. oneLevelSearchOps OK. searchOps is too high! 12090 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486407313] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 85% (15 min average)
[1486407352] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1486407412] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1486407422] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486407725] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 11346 found since last run. oneLevelSearchOps OK. searchOps is too high! 11619 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486407749] SERVICE ALERT: porterpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1486407764] SERVICE ALERT: FITNESSGRAM;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1486407789] SERVICE ALERT: porterpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1486407809] SERVICE ALERT: porterpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1486407845] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486407848] SERVICE ALERT: porterpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1486407869] SERVICE ALERT: porterpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1486407884] SERVICE ALERT: FITNESSGRAM;CPU Load;OK;SOFT;2;CPU Load 0% (15 min average)
[1486407908] SERVICE ALERT: porterpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1486407938] HOST ALERT: porterpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1486408002] HOST ALERT: porterpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.59 ms
[1486408022] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7013 found since last run. oneLevelSearchOps OK. searchOps is too high! 7151 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486408142] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486408445] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5105 found since last run. oneLevelSearchOps OK. searchOps is too high! 5401 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486408463] SERVICE ALERT: porterpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 37% - OK!
[1486408502] SERVICE ALERT: porterpk;Device Status;OK;HARD;3;Ready
[1486408565] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486408742] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6377 found since last run. oneLevelSearchOps OK. searchOps is too high! 6501 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486408862] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486409166] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9329 found since last run. oneLevelSearchOps OK. searchOps is too high! 9501 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486409285] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486409444] SERVICE ALERT: halleground1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1486409652] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;Ready
[1486409656] HOST ALERT: hallefirst1;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1486409762] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 51% - OK!
[1486409811] SERVICE ALERT: halleground3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1486409831] SERVICE ALERT: halleground1;Device Status;OK;HARD;3;ORDER CARTRIDGE LESS THAN 400 PAGES 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1486409842] HOST ALERT: halleground2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1486409861] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 6% - OK!
[1486409870] HOST ALERT: halleground2;UP;SOFT;2;PING OK - Packet loss = 58%, RTA = 0.47 ms
[1486409870] SERVICE ALERT: halleground2;Device Status;OK;HARD;3;Ready
[1486409880] SERVICE ALERT: halleground3;Device Status;OK;HARD;3;Ready
[1486409884] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5815 found since last run. oneLevelSearchOps OK. searchOps is too high! 6262 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486409913] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 48% - OK!
[1486409940] SERVICE ALERT: hallefirst3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1486409944] HOST ALERT: hallefirst3;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.44 ms
[1486410004] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486410140] SERVICE ALERT: hallefirst3;Device Status;OK;HARD;3;ORDER CARTRIDGE LESS THAN 1700 PAGES 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1486410160] SERVICE ALERT: hallefirst1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1486410253] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 8% - OK!
[1486410313] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1486410407] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 44% - OK!
[1486410433] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1486410454] SERVICE ALERT: halleground2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1486410553] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 83% (15 min average)
[1486410604] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7522 found since last run. oneLevelSearchOps OK. searchOps is too high! 7739 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486410662] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6906 found since last run. oneLevelSearchOps OK. searchOps is too high! 7024 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486410724] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486410740] SERVICE ALERT: hallefirst3;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1486410782] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486410800] SERVICE ALERT: hallefirst3;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1486410861] SERVICE ALERT: hallefirst3;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1486411323] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6146 found since last run. oneLevelSearchOps OK. searchOps is too high! 6312 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486411382] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8449 found since last run. oneLevelSearchOps OK. searchOps is too high! 8473 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486411443] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486411502] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486411753] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1486411873] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1486411993] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1486412043] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7118 found since last run. oneLevelSearchOps OK. searchOps is too high! 7267 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486412102] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7570 found since last run. oneLevelSearchOps OK. searchOps is too high! 7595 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486412163] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486412222] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486412657] SERVICE FLAPPING ALERT: owen101lab1;Device Status;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1486412660] SERVICE ALERT: hallefirst3;Device Status;OK;HARD;3;Genuine HP cartridge installed
[1486412763] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5902 found since last run. oneLevelSearchOps OK. searchOps is too high! 6437 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486412822] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 13913 found since last run. oneLevelSearchOps OK. searchOps is too high! 13927 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486412883] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486412942] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486413193] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 88% (15 min average)
[1486413483] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5541 found since last run. oneLevelSearchOps OK. searchOps is too high! 5868 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486413542] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10006 found since last run. oneLevelSearchOps OK. searchOps is too high! 10125 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486413603] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486413662] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486413684] SERVICE ALERT: MAPS01TEST;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1486413793] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1486413795] SERVICE ALERT: MAPS01TEST;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1486413913] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1486413940] SERVICE ALERT: hallefirst2;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1486414001] SERVICE ALERT: hallefirst2;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1486414033] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1486414061] SERVICE ALERT: hallefirst2;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1486414070] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1486414131] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1486414190] SERVICE ALERT: halleground2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1486414203] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5085 found since last run. oneLevelSearchOps OK. searchOps is too high! 5417 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486414212] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1486414262] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5538 found since last run. oneLevelSearchOps OK. searchOps is too high! 5655 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486414272] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1486414323] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486414332] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1486414382] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486414633] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 61% (15 min average)
[1486414923] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5914 found since last run. oneLevelSearchOps OK. searchOps is too high! 6224 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486414982] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6021 found since last run. oneLevelSearchOps OK. searchOps is too high! 6142 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486415043] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486415102] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486415437] SERVICE ALERT: DPS-SECURE02;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1486415461] SERVICE ALERT: DPS-SECURE02;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1486415467] HOST ALERT: DPS-SECURE02;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1486415507] SERVICE ALERT: DPS-SECURE02;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1486415515] SERVICE ALERT: DPS-SECURE02;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1486415520] SERVICE ALERT: DPS-SECURE02;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1486415520] SERVICE ALERT: DPS-SECURE02;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1486415537] HOST ALERT: DPS-SECURE02;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1486415543] SERVICE ALERT: DPS-SECURE02;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1486415543] SERVICE ALERT: DPS-SECURE02;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1486415543] SERVICE ALERT: DPS-SECURE02;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1486415547] SERVICE ALERT: DPS-SECURE02;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1486415557] SERVICE ALERT: DPS-SECURE02;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1486415573] HOST ALERT: DPS-SECURE02;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1486415581] SERVICE ALERT: DPS-SECURE02;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1486415611] HOST ALERT: DPS-SECURE02;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1486415611] SERVICE ALERT: DPS-SECURE02;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1486415627] SERVICE ALERT: DPS-SECURE02;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1486415635] SERVICE ALERT: DPS-SECURE02;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1486415640] SERVICE ALERT: DPS-SECURE02;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1486415640] SERVICE ALERT: DPS-SECURE02;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1486415641] HOST ALERT: DPS-SECURE02;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1486415663] SERVICE ALERT: DPS-SECURE02;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1486415663] SERVICE ALERT: DPS-SECURE02;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1486415663] SERVICE ALERT: DPS-SECURE02;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1486415667] SERVICE ALERT: DPS-SECURE02;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1486415677] SERVICE ALERT: DPS-SECURE02;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1486415693] HOST ALERT: DPS-SECURE02;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1486415701] SERVICE ALERT: DPS-SECURE02;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1486415702] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7944 found since last run. oneLevelSearchOps OK. searchOps is too high! 8056 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486415731] HOST ALERT: DPS-SECURE02;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1486415731] SERVICE ALERT: DPS-SECURE02;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1486415747] SERVICE ALERT: DPS-SECURE02;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1486415755] SERVICE ALERT: DPS-SECURE02;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1486415760] SERVICE ALERT: DPS-SECURE02;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1486415760] SERVICE ALERT: DPS-SECURE02;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1486415761] HOST ALERT: DPS-SECURE02;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1486415781] SERVICE ALERT: DPS-SECURE02;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1486415783] SERVICE ALERT: DPS-SECURE02;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1486415783] SERVICE ALERT: DPS-SECURE02;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1486415783] SERVICE ALERT: DPS-SECURE02;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1486415787] SERVICE ALERT: DPS-SECURE02;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1486415811] HOST ALERT: DPS-SECURE02;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1486415822] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486415851] SERVICE ALERT: DPS-SECURE02;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1486415856] SERVICE ALERT: DPS-SECURE02;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1486415861] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1486415881] HOST ALERT: DPS-SECURE02;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1486415901] SERVICE ALERT: DPS-SECURE02;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1486415976] SERVICE ALERT: DPS-SECURE02;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1486416021] SERVICE ALERT: DPS-SECURE02;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1486416096] SERVICE ALERT: DPS-SECURE02;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1486416243] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6316 found since last run. oneLevelSearchOps OK. searchOps is too high! 6644 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486416257] SERVICE ALERT: DPS-SECURE02;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1486416363] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486416422] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5483 found since last run. oneLevelSearchOps OK. searchOps is too high! 5598 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486416433] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1486416461] SERVICE ALERT: hallefirst2;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1486416461] HOST ALERT: DPS-SECURE02;UP;HARD;10;PING OK - Packet loss = 54%, RTA = 0.45 ms
[1486416520] SERVICE ALERT: hallefirst2;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1486416542] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486416553] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1486416580] SERVICE ALERT: hallefirst2;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1486416601] SERVICE ALERT: DPS-SECURE02;CPU Load;OK;HARD;3;CPU Load 2% (15 min average)
[1486416673] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1486416676] SERVICE ALERT: DPS-SECURE02;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1486416880] SERVICE ALERT: DPS-SECURE02;Memory Usage;OK;HARD;3;Memory usage: total:4790.25 MB - used: 874.81 MB (18%) - free: 3915.44 MB (82%)
[1486416927] SERVICE ALERT: DPS-SECURE02;C:\ Drive Space;OK;HARD;3;c:\ - total: 278.53 Gb - used: 14.60 Gb (5%) - free 263.93 Gb (95%)
[1486416935] SERVICE ALERT: DPS-SECURE02;Service: RPC;OK;HARD;3;RpcSs: Started
[1486416940] SERVICE ALERT: DPS-SECURE02;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 9 minute(s)
[1486416940] SERVICE ALERT: DPS-SECURE02;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1486416962] SERVICE ALERT: DPS-SECURE02;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1486416962] SERVICE ALERT: DPS-SECURE02;Service: Windows Time;OK;HARD;3;W32Time: Started
[1486416963] SERVICE ALERT: DPS-SECURE02;Service: WMI;OK;HARD;3;Winmgmt: Started
[1486416963] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7506 found since last run. oneLevelSearchOps OK. searchOps is too high! 7921 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486416967] SERVICE ALERT: DPS-SECURE02;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1486417031] SERVICE ALERT: DPS-SECURE02;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1486417066] HOST ALERT: hallefirst2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1486417083] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486417130] HOST ALERT: hallefirst2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1486417142] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7554 found since last run. oneLevelSearchOps OK. searchOps is too high! 7701 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486417185] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1486417185] SERVICE FLAPPING ALERT: hallefirst2;Paper Tray Status;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1486417215] HOST ALERT: hallefirst2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1486417262] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486417305] HOST ALERT: hallefirst2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1486417338] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1486417348] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1486417368] HOST ALERT: hallefirst2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1486417408] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1486417438] HOST ALERT: hallefirst2;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1486417468] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1486417498] HOST ALERT: hallefirst2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1486417588] HOST ALERT: hallefirst2;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1486417678] HOST ALERT: hallefirst2;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1486417683] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9197 found since last run. oneLevelSearchOps OK. searchOps is too high! 9394 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486417768] HOST ALERT: hallefirst2;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1486417803] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486417815] HOST ALERT: hallefirst2;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1486417873] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 78% (15 min average)
[1486417905] HOST ALERT: hallefirst2;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1486418403] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9102 found since last run. oneLevelSearchOps OK. searchOps is too high! 9446 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486418462] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7538 found since last run. oneLevelSearchOps OK. searchOps is too high! 7629 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486418473] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1486418523] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486418582] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486418593] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 91% (15 min average)
[1486418990] SERVICE ALERT: halleground2;Device Status;OK;HARD;3;Processing job from tray 3
[1486419123] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7239 found since last run. oneLevelSearchOps OK. searchOps is too high! 7678 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486419182] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10140 found since last run. oneLevelSearchOps OK. searchOps is too high! 10229 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486419193] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1486419243] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486419302] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486419313] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1486419433] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1486419590] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1486419650] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1486419710] SERVICE ALERT: halleground2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1486419843] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7497 found since last run. oneLevelSearchOps OK. searchOps is too high! 7786 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486419902] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5677 found since last run. oneLevelSearchOps OK. searchOps is too high! 5767 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486419963] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486420022] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486420033] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1486420310] SERVICE ALERT: halleground2;Device Status;OK;HARD;3;Ready 164.76.106.52
[1486420563] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7880 found since last run. oneLevelSearchOps OK. searchOps is too high! 8188 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486420622] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7533 found since last run. oneLevelSearchOps OK. searchOps is too high! 7544 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486420633] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1486420683] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486420742] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486420883] SERVICE ALERT: rooseveltpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1486420944] SERVICE ALERT: rooseveltpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1486421003] SERVICE ALERT: rooseveltpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1486421283] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7762 found since last run. oneLevelSearchOps OK. searchOps is too high! 8117 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486421403] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486421942] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7084 found since last run. oneLevelSearchOps OK. searchOps is too high! 7108 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486421979] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1486421983] HOST ALERT: hallefirst2;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.47 ms
[1486422062] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486422132] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Processing job from tray 4
[1486422176] SERVICE ALERT: cava2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1486422262] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 26% - OK!
[1486422296] SERVICE ALERT: cava2;CPU Load;OK;SOFT;2;CPU Load 4% (15 min average)
[1486422662] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7986 found since last run. oneLevelSearchOps OK. searchOps is too high! 8001 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486422782] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486422859] SERVICE FLAPPING ALERT: hallefirst3;Device Status;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1486422980] SERVICE ALERT: rooseveltpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1486423040] SERVICE ALERT: rooseveltpk;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 100% - OK!
[1486423203] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7276 found since last run. oneLevelSearchOps OK. searchOps is too high! 7681 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486423323] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486423382] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5498 found since last run. oneLevelSearchOps OK. searchOps is too high! 5507 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486423403] SERVICE ALERT: rooseveltpk;Device Status;OK;HARD;3;SIZE MISMATCH IN TRAY 4 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1486423456] SERVICE ALERT: owen101lab1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1486423502] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486423516] SERVICE ALERT: owen101lab1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1486423576] SERVICE ALERT: owen101lab1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1486424233] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1486424702] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5790 found since last run. oneLevelSearchOps OK. searchOps is too high! 5802 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486424776] SERVICE ALERT: owen101lab1;Device Status;OK;HARD;3;Ready
[1486424822] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486424833] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1486424953] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1486425073] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 94% (15 min average)
[1486425262] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1486425322] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 100% - OK!
[1486425422] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5794 found since last run. oneLevelSearchOps OK. searchOps is too high! 5811 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486425542] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486425673] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1486425793] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1486426142] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8891 found since last run. oneLevelSearchOps OK. searchOps is too high! 8973 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486426262] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486426393] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1486426513] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1486426633] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1486426862] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8425 found since last run. oneLevelSearchOps OK. searchOps is too high! 8522 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486426982] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486427582] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7072 found since last run. oneLevelSearchOps OK. searchOps is too high! 7158 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486427702] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486428302] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5130 found since last run. oneLevelSearchOps OK. searchOps is too high! 5200 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486428422] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486429022] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5246 found since last run. oneLevelSearchOps OK. searchOps is too high! 5336 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486429033] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 87% (15 min average)
[1486429142] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486429178] SERVICE FLAPPING ALERT: hallefirst2;Paper Tray Status;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1486429323] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7025 found since last run. oneLevelSearchOps OK. searchOps is too high! 7211 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486429443] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486429633] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1486429742] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7554 found since last run. oneLevelSearchOps OK. searchOps is too high! 7562 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486429753] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1486429862] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486429873] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 91% (15 min average)
[1486430643] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6900 found since last run. oneLevelSearchOps OK. searchOps is too high! 7098 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486430652] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1486430712] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1486430763] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486430773] SERVICE ALERT: hallefirst1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1486431363] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5591 found since last run. oneLevelSearchOps OK. searchOps is too high! 6051 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486431372] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;Ready
[1486431483] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486431662] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5554 found since last run. oneLevelSearchOps OK. searchOps is too high! 5563 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486431709] SERVICE FLAPPING ALERT: halleground2;Device Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1486431782] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486432083] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5527 found since last run. oneLevelSearchOps OK. searchOps is too high! 5835 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1486432101] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1486432161] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1486432203] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1486432222] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1486432821] SERVICE ALERT: hallesecondsouth;Device Status;OK;HARD;3;Processing job from tray 4
[1486433473] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1486433593] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1486433713] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1486434913] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1486435631] HOST ALERT: IMAGES;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1486435700] HOST ALERT: IMAGES;UP;SOFT;2;PING OK - Packet loss = 50%, RTA = 0.26 ms
[1486438513] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1486438633] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1486438753] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 93% (15 min average)
[1486439058] SERVICE ALERT: hallefirst3;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1486439118] SERVICE ALERT: hallefirst3;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1486439178] SERVICE ALERT: hallefirst3;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1486439953] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1486440073] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1486440193] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 90% (15 min average)
[1486440794] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1486440913] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1486441033] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1486442178] SERVICE ALERT: hallefirst3;Device Status;OK;HARD;3;Checking printer
[1486442233] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 90% (15 min average)
[1486442778] SERVICE ALERT: hallefirst3;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1486442833] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1486442838] SERVICE ALERT: hallefirst3;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1486442898] SERVICE ALERT: hallefirst3;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1486442953] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)