[1461816716] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1461816836] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 76% (15 min average)
[1461817800] HOST DOWNTIME ALERT: REFWIN;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: EMU-VALUE01;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: DPS-SECURE02;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: FTK;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: Echo-01;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: R25-ZETA2012;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: EMUSCRIPT;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: PNCAPP-2012;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: TI22-COE;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: R25-DELTA2012;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: ADDC55;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: TI22;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: R25-ALPHA2012;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: COE-WEBDAV;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: DIRECTACCESS;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: PROCARE;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: XDM;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: CEDARON;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: EMUSQLEXPRESS;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: CBORD;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: ADWU1;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: PNCWEB-2012;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: ADWSUSCLIENTS;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: ADDC66;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: PNCDB-2012;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: TABLEAU;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: ADDC44;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: ADDC5;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: TI22-PSYCLINIC;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: ADDC11;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: SP2010DB;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: WINKMS;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: ADDC4;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: OPIE;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: ADDC33;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: ADDC3;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: OFFICEKMS;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: ADDC2;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: HDBACKUPS;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: ADDC22;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: NETENG-PA;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: SP2010;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: ADAUTHDC2;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: ADCA2;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: NESSUS;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: ADCA1;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: MBAMDB;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: UHSPHARM-2012;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: SCCM-PRIM01;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: IMAGES;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: SANDBOX-DB;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: MBAM;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: ADAUTHDC1;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: LANSWEEPER;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: ADAUTH-RMAD;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: ADARC;STOPPED; Host has exited from a period of scheduled downtime
[1461817800] HOST DOWNTIME ALERT: AD-RMAD;STOPPED; Host has exited from a period of scheduled downtime
[1461818635] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 91% (15 min average)
[1461819235] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1461819355] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461819475] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461820675] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 75% (15 min average)
[1461821275] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461821395] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461821515] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1461822115] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461822715] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461822835] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1461822954] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1461823073] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1461823258] SERVICE ALERT: ADCA1;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1461823315] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461823378] SERVICE ALERT: ADCA1;CPU Load;OK;SOFT;2;CPU Load 24% (15 min average)
[1461824515] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 85% (15 min average)
[1461825115] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1461825235] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1461825355] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461825955] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461826673] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1461827273] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1461827393] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1461827513] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461827753] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 87% (15 min average)
[1461828353] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461828473] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461828593] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461828713] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461828713] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1461829913] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461830513] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461831113] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461833513] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1461833991] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 72% (15 min average)
[1461834113] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1461834233] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1461834353] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 93% (15 min average)
[1461835553] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1461835673] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 88% (15 min average)
[1461836273] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1461836391] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1461836393] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1461836511] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461836513] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461836631] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1461837831] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1461838431] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461838551] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461838671] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461838913] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1461839271] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461839513] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1461839634] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1461839752] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 91% (15 min average)
[1461839871] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461845788] SERVICE ALERT: R25-ALPHA2012;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.99.204 and port 12489: Connection refused
[1461845905] SERVICE ALERT: R25-ALPHA2012;Memory Usage;OK;SOFT;2;Memory usage: total:10943.64 MB - used: 720.42 MB (7%) - free: 10223.22 MB (93%)
[1461845941] SERVICE ALERT: ADRODC-LIV;Service: Group Policy;CRITICAL;SOFT;1;gpsvc: Stopped
[1461846081] SERVICE ALERT: ADRODC-LIV;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1461846091] SERVICE ALERT: ADRODC-LIV;Service: DNS Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461846091] SERVICE ALERT: ADRODC-LIV;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461846100] SERVICE ALERT: ADRODC-LIV;Service: MDT Monitor;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461846111] HOST ALERT: ADRODC-LIV;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1461846131] SERVICE ALERT: ADRODC-LIV;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461846161] HOST ALERT: ADRODC-LIV;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1461846161] SERVICE ALERT: ADRODC-LIV;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461846191] HOST ALERT: ADRODC-LIV;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1461846200] SERVICE ALERT: ADRODC-LIV;Service: DNS Service;CRITICAL;SOFT;2;connect to address 207.74.163.140 and port 12489: Connection refused
[1461846200] SERVICE ALERT: ADRODC-LIV;Service: MDT Monitor;CRITICAL;SOFT;2;connect to address 207.74.163.140 and port 12489: Connection refused
[1461846200] SERVICE ALERT: ADRODC-LIV;CPU Load;CRITICAL;SOFT;2;connect to address 207.74.163.140 and port 12489: Connection refused
[1461846201] SERVICE ALERT: ADRODC-LIV;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461846201] SERVICE ALERT: ADRODC-LIV;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1461846217] HOST ALERT: ADRODC-LIV;UP;SOFT;4;PING OK - Packet loss = 70%, RTA = 6.63 ms
[1461846231] SERVICE ALERT: ADRODC-LIV;Memory Usage;OK;SOFT;2;Memory usage: total:37586.86 MB - used: 1265.25 MB (3%) - free: 36321.62 MB (97%)
[1461846261] SERVICE ALERT: ADRODC-LIV;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1461846301] SERVICE ALERT: ADRODC-LIV;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1461846311] SERVICE ALERT: ADRODC-LIV;Service: DNS Service;OK;SOFT;3;DNS: Started
[1461846311] SERVICE ALERT: ADRODC-LIV;CPU Load;OK;SOFT;3;CPU Load 6% (15 min average)
[1461846320] SERVICE ALERT: ADRODC-LIV;Service: MDT Monitor;OK;SOFT;3;MDT_Monitor: Started
[1461846360] SERVICE ALERT: idmrbstest.emich.edu;Banner Driver;WARNING;SOFT;1;0
[1461846480] SERVICE ALERT: idmrbstest.emich.edu;Banner Driver;WARNING;SOFT;2;0
[1461846595] SERVICE ALERT: idmrbstest.emich.edu;Banner Driver;OK;SOFT;3;0
[1461846671] SERVICE ALERT: c137.emich.edu;Service: HAProxy Status;CRITICAL;SOFT;1;OK: Frontend Online. CRITICAL: Backend Offline!
[1461846781] SERVICE ALERT: ADRODC-LIV;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1461846791] SERVICE ALERT: c137.emich.edu;Service: HAProxy Status;CRITICAL;HARD;2;OK: Frontend Online. CRITICAL: Backend Offline!
[1461847071] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 35% (15 min average)
[1461847671] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461847691] SERVICE ALERT: c137.emich.edu;Service: HAProxy Status;OK;HARD;2;OK: Frontend Online. Backend Online.
[1461847791] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1461848152] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1461848272] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 86% (15 min average)
[1461848291] SERVICE ALERT: c137.emich.edu;Service: HAProxy Status;CRITICAL;SOFT;1;OK: Frontend Online. CRITICAL: Backend Offline!
[1461848391] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1461848411] SERVICE ALERT: c137.emich.edu;Service: HAProxy Status;CRITICAL;HARD;2;OK: Frontend Online. CRITICAL: Backend Offline!
[1461848511] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461848631] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461848872] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1461849611] SERVICE ALERT: c137.emich.edu;Service: HAProxy Status;OK;HARD;2;OK: Frontend Online. Backend Online.
[1461849611] SERVICE FLAPPING ALERT: c137.emich.edu;Service: HAProxy Status;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1461851031] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 82% (15 min average)
[1461851631] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1461851751] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461851871] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461854411] SERVICE FLAPPING ALERT: c137.emich.edu;Service: HAProxy Status;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1461856072] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1461856192] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1461857271] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 74% (15 min average)
[1461857871] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461857991] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461858111] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 64% (15 min average)
[1461858592] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1461858711] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461858712] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 90% (15 min average)
[1461858831] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1461858951] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 90% (15 min average)
[1461859312] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461859432] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 90% (15 min average)
[1461859551] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461859671] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461859791] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1461861591] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (15 min average)
[1461862191] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1461862311] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461862431] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1461863631] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1461864231] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461864351] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461864471] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461864848] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461864909] SERVICE ALERT: hallefirst2;Device Status;OK;SOFT;2;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1461865071] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461865432] SERVICE ALERT: ADRODC-DET;Service: AD Web Service;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1461865432] SERVICE ALERT: ADRODC-DET;Service: Firewall;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1461865432] SERVICE ALERT: ADRODC-DET;Service: Key Distribution Center;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1461865459] SERVICE ALERT: ADRODC-DET;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461865488] SERVICE ALERT: ADRODC-DET;AD Replication;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 30 seconds.
[1461865489] HOST ALERT: ADRODC-DET;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1461865518] SERVICE ALERT: ADRODC-DET;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461865518] SERVICE ALERT: ADRODC-DET;Service: DNS Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461865518] SERVICE ALERT: ADRODC-DET;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461865548] HOST ALERT: ADRODC-DET;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1461865572] SERVICE ALERT: ADRODC-DET;Service: AD Web Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1461865572] SERVICE ALERT: ADRODC-DET;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1461865572] SERVICE ALERT: ADRODC-DET;Service: Key Distribution Center;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1461865579] SERVICE ALERT: ADRODC-DET;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1461865599] SERVICE ALERT: ADRODC-DET;AD Replication;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1461865602] HOST ALERT: ADRODC-DET;DOWN;SOFT;3;(Host check timed out after 30.00 seconds)
[1461865638] SERVICE ALERT: ADRODC-DET;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1461865638] SERVICE ALERT: ADRODC-DET;Service: DNS Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1461865638] SERVICE ALERT: ADRODC-DET;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1461865642] HOST ALERT: ADRODC-DET;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 465.27 ms
[1461865673] SERVICE ALERT: ADRODC-DET;Service: AD Web Service;OK;SOFT;3;ADWS: Started
[1461865673] SERVICE ALERT: ADRODC-DET;Service: Firewall;OK;SOFT;3;MpsSvc: Started
[1461865673] SERVICE ALERT: ADRODC-DET;Service: Key Distribution Center;OK;SOFT;3;KDC: Started
[1461865680] SERVICE ALERT: ADRODC-DET;C:\ Drive Space;OK;SOFT;3;c:\ - total: 1115.16 Gb - used: 279.49 Gb (25%) - free 835.67 Gb (75%)
[1461865708] SERVICE ALERT: ADRODC-DET;AD Replication;OK;SOFT;3;OK: replication is up and running
[1461865739] SERVICE ALERT: ADRODC-DET;Service: Windows Time;OK;SOFT;3;W32Time: Started
[1461865739] SERVICE ALERT: ADRODC-DET;Service: DNS Service;OK;SOFT;3;DNS: Started
[1461865739] SERVICE ALERT: ADRODC-DET;Service: Group Policy;OK;SOFT;3;gpsvc: Started
[1461866032] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1461866152] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 90% (15 min average)
[1461868671] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1461869271] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461869353] SERVICE ALERT: ADWSUSSERVERS;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1461869391] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1461869482] SERVICE ALERT: ADWSUSSERVERS;Service: Windows Time;CRITICAL;SOFT;2;W32Time: Stopped
[1461869511] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461869593] SERVICE ALERT: ADWSUSSERVERS;Service: Windows Time;OK;SOFT;3;W32Time: Started
[1461869708] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461869768] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461869828] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461870428] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Ready
[1461871311] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461871911] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461875511] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (15 min average)
[1461876111] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461876231] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461876351] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461877203] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (35.3% change >= 20.0% threshold)
[1461877203] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1461877203] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1461877203] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1461877203] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1461878095] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;SOFT;1;WARNING: Replication with ADDC5 failed!
[1461878214] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;SOFT;2;WARNING: Replication with ADDC5 failed!
[1461878335] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;HARD;3;WARNING: Replication with ADDC5 failed!
[1461879003] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (29.2% change >= 20.0% threshold)
[1461879003] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1461879003] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1461879003] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1461879003] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1461879350] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461879950] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461880803] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (38.9% change >= 20.0% threshold)
[1461880803] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1461880803] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1461880803] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1461880803] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1461881149] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1461881269] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 76% (15 min average)
[1461882350] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 70% (15 min average)
[1461882950] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461883070] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1461883190] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1461884990] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 90% (15 min average)
[1461885590] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461885710] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1461885830] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 74% (15 min average)
[1461886430] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461886550] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1461886670] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461890270] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461890870] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461891470] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461892070] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1461892583] SERVICE ALERT: IT-2012;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1461892670] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461892703] SERVICE ALERT: IT-2012;CPU Load;OK;SOFT;2;CPU Load 40% (15 min average)
[1461893870] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461898216] SERVICE ALERT: CBORD;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461898316] SERVICE ALERT: CBORD;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1461899268] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1461899339] SERVICE ALERT: ADRODC-LIV;AD Replication;OK;HARD;3;OK: replication is up and running
[1461899388] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1461899508] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 93% (15 min average)
[1461900109] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1461900228] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461900348] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1461900948] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1461902270] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)