[1463373773] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463374373] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463376773] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 69% (15 min average)
[1463377373] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463377493] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463377613] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 86% (15 min average)
[1463378213] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463378333] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1463378453] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463379053] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 77% (15 min average)
[1463379653] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1463379773] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463379893] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463380493] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 77% (15 min average)
[1463381093] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463381213] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463381333] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 81% (15 min average)
[1463381933] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1463382053] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463382173] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 93% (15 min average)
[1463382773] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463382893] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463383013] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 80% (15 min average)
[1463383611] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463383731] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 88% (15 min average)
[1463384331] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1463384451] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1463384571] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1463385771] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 80% (15 min average)
[1463386371] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1463386491] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1463386611] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1463387211] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 82% (15 min average)
[1463391411] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463391531] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463391651] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463391791] SERVICE ALERT: ADWU1;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1463391911] SERVICE ALERT: ADWU1;CPU Load;OK;SOFT;2;CPU Load 0% (15 min average)
[1463392251] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 29% (15 min average)
[1463393451] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463393571] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463393691] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 89% (15 min average)
[1463394891] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463395011] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463395131] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463396331] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1463396931] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463397051] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463397171] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463397771] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 84% (15 min average)
[1463398104] SERVICE ALERT: hallefirst3;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1463398163] SERVICE ALERT: hallefirst3;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! 2 of 3 Trays Still Full
[1463398371] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1463398491] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463398611] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 75% (15 min average)
[1463400411] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463400531] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463400651] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1463401321] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (36.5% change >= 20.0% threshold)
[1463401321] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1463401321] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1463401321] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1463401321] SERVICE DOWNTIME ALERT: NETIDADPROXY;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1463401321] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1463401387] HOST ALERT: rooseveltpk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.59 ms
[1463401450] SERVICE ALERT: rooseveltpk;Device Status;OK;HARD;3;Ready
[1463401453] SERVICE ALERT: rooseveltpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1463401487] HOST DOWNTIME ALERT: GOVGIS;STARTED; Host has entered a period of scheduled downtime
[1463401523] HOST ALERT: sillpk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1463401663] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Ready
[1463401783] SERVICE ALERT: sillpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1463401804] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (36.5% change >= 20.0% threshold)
[1463401804] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1463401804] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1463401804] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1463401804] SERVICE DOWNTIME ALERT: NETIDADPROXY;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1463401804] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1463401804] HOST DOWNTIME ALERT: GOVGIS;STARTED; Host has entered a period of scheduled downtime
[1463402369] SERVICE ALERT: porterpk;Device Status;OK;HARD;3;Warming up
[1463402373] HOST ALERT: porterpk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.54 ms
[1463402417] SERVICE ALERT: porterpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1463402451] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463402995] SERVICE ALERT: BOEAPPTEST;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1463403051] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 87% (15 min average)
[1463403114] SERVICE ALERT: BOEAPPTEST;CPU Load;OK;SOFT;2;CPU Load 92% (15 min average)
[1463404851] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1463404971] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463405091] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1463405691] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 82% (15 min average)
[1463406291] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1463406411] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1463406531] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1463408329] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 86% (15 min average)
[1463409529] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463409649] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 40% (15 min average)
[1463410249] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463410369] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463410489] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463410803] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (45.6% change >= 20.0% threshold)
[1463410803] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1463410803] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1463410803] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1463410803] SERVICE DOWNTIME ALERT: NETIDADPROXY;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1463410803] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1463410803] HOST DOWNTIME ALERT: GOVGIS;STARTED; Host has entered a period of scheduled downtime
[1463410832] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463410836] HOST ALERT: authtestdb;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1463410952] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 87% (15 min average)
[1463411089] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 77% (15 min average)
[1463412232] SERVICE ALERT: honorspk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1463412292] SERVICE ALERT: honorspk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1463412352] SERVICE ALERT: honorspk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1463412934] SERVICE ALERT: MBAM;Service: RPC;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1463412969] SERVICE ALERT: MBAM;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463412982] SERVICE ALERT: ADWSUSCLIENTS;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463412982] SERVICE ALERT: ADWSUSCLIENTS;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463412982] SERVICE ALERT: ADWU1;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463412983] SERVICE ALERT: ADWSUSCLIENTS;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463412984] SERVICE ALERT: ADWSUSCLIENTS;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463412984] SERVICE ALERT: ADWSUSCLIENTS;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463412984] SERVICE ALERT: ADWSUSCLIENTS;E:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463412984] SERVICE ALERT: ADWSUSCLIENTS;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463412984] SERVICE ALERT: ADWSUSCLIENTS;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463412984] SERVICE ALERT: ADWSUSCLIENTS;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463412984] SERVICE ALERT: ADWSUSCLIENTS;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463412986] SERVICE ALERT: MBAMDB;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463412987] SERVICE ALERT: ADWSUSCLIENTS;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463412990] SERVICE ALERT: ADWSUSCLIENTS;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463412990] SERVICE ALERT: OFFICEKMS;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463412994] SERVICE ALERT: ADWSUSCLIENTS;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463412994] SERVICE ALERT: ADWSUSCLIENTS;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463412995] SERVICE ALERT: MBAM;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463412995] SERVICE ALERT: ADWSUSCLIENTS;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463412997] SERVICE ALERT: OPIE;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463412998] SERVICE ALERT: ADWU1;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463412999] HOST ALERT: MBAM;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1463413001] SERVICE ALERT: MBAMDB;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463413011] HOST ALERT: OPIE;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1463413012] HOST ALERT: ADWSUSCLIENTS;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1463413012] HOST ALERT: ADWU1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1463413016] HOST ALERT: MBAMDB;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1463413020] HOST ALERT: OFFICEKMS;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1463413030] SERVICE ALERT: MBAMDB;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463413031] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 12072;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 30 seconds.
[1463413031] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10042;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 30 seconds.
[1463413032] SERVICE ALERT: MBAM;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463413033] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 30 seconds.
[1463413042] SERVICE ALERT: ADWU1;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463413044] SERVICE ALERT: ADWU1;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463413047] SERVICE ALERT: ADWU1;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463413048] SERVICE ALERT: OFFICEKMS;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463413050] SERVICE ALERT: OPIE;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463413060] HOST ALERT: MBAMDB;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1463413061] HOST ALERT: ADWSUSCLIENTS;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1463413062] HOST ALERT: MBAM;DOWN;SOFT;2;(Host check timed out after 30.00 seconds)
[1463413072] HOST ALERT: ADWU1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1463413074] SERVICE ALERT: MBAM;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463413078] HOST ALERT: OFFICEKMS;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1463413080] HOST ALERT: OPIE;DOWN;SOFT;2;(Host check timed out after 30.00 seconds)
[1463413089] SERVICE ALERT: MBAM;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463413102] SERVICE ALERT: ADWSUSCLIENTS;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463413102] SERVICE ALERT: ADWSUSCLIENTS;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463413102] SERVICE ALERT: ADWU1;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463413103] SERVICE ALERT: ADWSUSCLIENTS;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463413104] SERVICE ALERT: ADWSUSCLIENTS;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463413104] SERVICE ALERT: ADWSUSCLIENTS;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463413104] SERVICE ALERT: ADWSUSCLIENTS;E:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463413104] SERVICE ALERT: ADWSUSCLIENTS;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463413104] SERVICE ALERT: ADWSUSCLIENTS;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463413104] SERVICE ALERT: ADWSUSCLIENTS;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463413104] SERVICE ALERT: ADWSUSCLIENTS;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463413104] HOST ALERT: MBAM;DOWN;SOFT;3;(Host check timed out after 30.00 seconds)
[1463413106] SERVICE ALERT: MBAMDB;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463413107] SERVICE ALERT: ADWSUSCLIENTS;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463413110] SERVICE ALERT: ADWSUSCLIENTS;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463413110] SERVICE ALERT: OFFICEKMS;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463413114] SERVICE ALERT: ADWSUSCLIENTS;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463413114] SERVICE ALERT: ADWSUSCLIENTS;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463413115] SERVICE ALERT: MBAM;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463413115] SERVICE ALERT: ADWSUSCLIENTS;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463413117] SERVICE ALERT: OPIE;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463413118] SERVICE ALERT: ADWU1;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463413121] SERVICE ALERT: MBAMDB;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463413124] SERVICE ALERT: ADWU1;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463413132] HOST ALERT: ADWSUSCLIENTS;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1463413132] HOST ALERT: ADWU1;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1463413136] HOST ALERT: MBAMDB;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1463413140] HOST ALERT: OFFICEKMS;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1463413145] HOST ALERT: MBAM;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1463413147] HOST ALERT: OPIE;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1463413149] SERVICE ALERT: ADWU1;Windows Update Services Event 10042;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 30 seconds.
[1463413150] SERVICE ALERT: MBAMDB;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463413151] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10042;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 30 seconds.
[1463413151] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 12072;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 30 seconds.
[1463413152] SERVICE ALERT: MBAM;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463413153] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 30 seconds.
[1463413162] SERVICE ALERT: ADWU1;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463413162] SERVICE ALERT: OFFICEKMS;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463413162] SERVICE ALERT: MBAMDB;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463413164] SERVICE ALERT: ADWU1;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463413167] SERVICE ALERT: ADWU1;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463413168] SERVICE ALERT: OFFICEKMS;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463413170] SERVICE ALERT: OPIE;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463413175] SERVICE ALERT: MBAMDB;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463413179] HOST ALERT: ADWU1;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1463413180] HOST ALERT: MBAMDB;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1463413181] HOST ALERT: ADWSUSCLIENTS;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1463413182] HOST ALERT: MBAM;DOWN;SOFT;5;(Host check timed out after 30.00 seconds)
[1463413192] HOST ALERT: OFFICEKMS;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1463413194] SERVICE ALERT: MBAM;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1463413200] HOST ALERT: OPIE;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1463413202] SERVICE ALERT: ADWU1;Service: WMI;OK;SOFT;3;Winmgmt: Started
[1463413204] SERVICE ALERT: ADWSUSCLIENTS;Service: DNS Client;CRITICAL;HARD;3;connect to address 164.76.4.77 and port 12489: Connection refused
[1463413204] SERVICE ALERT: ADWSUSCLIENTS;Service: Firewall;CRITICAL;HARD;3;connect to address 164.76.4.77 and port 12489: Connection refused
[1463413204] SERVICE ALERT: ADWSUSCLIENTS;E:\ Drive Space;CRITICAL;HARD;3;connect to address 164.76.4.77 and port 12489: Connection refused
[1463413204] SERVICE ALERT: ADWSUSCLIENTS;Memory Usage;CRITICAL;HARD;3;connect to address 164.76.4.77 and port 12489: Connection refused
[1463413204] SERVICE ALERT: ADWSUSCLIENTS;VMWare Tools;CRITICAL;HARD;3;connect to address 164.76.4.77 and port 12489: Connection refused
[1463413204] SERVICE ALERT: ADWSUSCLIENTS;Service: User Profile;CRITICAL;HARD;3;connect to address 164.76.4.77 and port 12489: Connection refused
[1463413204] SERVICE ALERT: ADWSUSCLIENTS;Service: Windows Time;CRITICAL;HARD;3;connect to address 164.76.4.77 and port 12489: Connection refused
[1463413206] HOST ALERT: ADWU1;UP;SOFT;5;PING OK - Packet loss = 0%, RTA = 0.36 ms
[1463413209] SERVICE ALERT: MBAM;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1463413211] SERVICE ALERT: ADWSUSCLIENTS;Uptime;OK;SOFT;3;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1463413211] SERVICE ALERT: ADWSUSCLIENTS;C:\ Drive Space;OK;SOFT;3;c:\ - total: 79.66 Gb - used: 39.34 Gb (49%) - free 40.31 Gb (51%)
[1463413212] SERVICE ALERT: ADWSUSCLIENTS;CPU Load;OK;SOFT;3;CPU Load 0% (15 min average)
[1463413213] SERVICE ALERT: ADWSUSCLIENTS;Service: Group Policy;OK;SOFT;3;gpsvc: Started
[1463413214] SERVICE ALERT: ADWSUSCLIENTS;Service: NetLogin;OK;SOFT;3;Netlogon: Started
[1463413214] SERVICE ALERT: ADWSUSCLIENTS;Service: WMI;OK;SOFT;3;Winmgmt: Started
[1463413215] SERVICE ALERT: ADWSUSCLIENTS;Service: RPC;OK;SOFT;3;RpcSs: Started
[1463413216] SERVICE ALERT: ADWSUSCLIENTS;NSClient++ Version;OK;SOFT;3;NSClient++ 0,4,1,73 2012-12-17
[1463413217] HOST ALERT: ADWSUSCLIENTS;UP;SOFT;5;PING OK - Packet loss = 64%, RTA = 0.32 ms
[1463413218] SERVICE ALERT: ADWU1;VMWare Tools;OK;SOFT;3;VMTools: Started
[1463413224] SERVICE ALERT: ADWU1;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1463413224] HOST ALERT: MBAM;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1463413224] SERVICE ALERT: MBAMDB;Service: Firewall;OK;SOFT;3;MpsSvc: Started
[1463413226] SERVICE ALERT: MBAMDB;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1463413230] SERVICE ALERT: OFFICEKMS;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1463413232] SERVICE ALERT: MBAMDB;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463413232] SERVICE ALERT: MBAM;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463413235] SERVICE ALERT: MBAM;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1463413237] SERVICE ALERT: OPIE;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1463413238] HOST ALERT: MBAMDB;UP;SOFT;5;PING OK - Packet loss = 66%, RTA = 0.31 ms
[1463413239] SERVICE ALERT: ADWU1;Windows Update Services Event 10042;OK;SOFT;2;OK: No entries found
[1463413241] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 12072;OK;SOFT;3;Eventlog check ok
[1463413241] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10042;OK;SOFT;3;Eventlog check ok
[1463413243] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10032;OK;SOFT;3;Eventlog check ok
[1463413249] HOST ALERT: OFFICEKMS;UP;SOFT;5;PING OK - Packet loss = 75%, RTA = 0.34 ms
[1463413250] SERVICE ALERT: MBAMDB;CPU Load;OK;SOFT;3;CPU Load 17% (15 min average)
[1463413252] SERVICE ALERT: MBAM;Service: Firewall;OK;SOFT;3;MpsSvc: Started
[1463413262] SERVICE ALERT: ADWU1;Service: Group Policy;OK;SOFT;3;gpsvc: Started
[1463413262] SERVICE ALERT: MBAMDB;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1463413262] SERVICE ALERT: OFFICEKMS;CPU Load;OK;SOFT;2;CPU Load 73% (15 min average)
[1463413262] HOST ALERT: MBAM;UP;SOFT;7;PING WARNING - Packet loss = 90%, RTA = 0.38 ms
[1463413264] SERVICE ALERT: ADWU1;Service: RPC;OK;SOFT;3;RpcSs: Started
[1463413267] SERVICE ALERT: ADWU1;Service: Windows Time;OK;SOFT;3;W32Time: Started
[1463413267] HOST ALERT: OPIE;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1463413268] SERVICE ALERT: OFFICEKMS;Service: NetLogin;OK;SOFT;3;Netlogon: Started
[1463413273] SERVICE ALERT: OPIE;Service: Windows Time;CRITICAL;HARD;3;W32Time: Stopped
[1463413275] SERVICE ALERT: MBAMDB;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1463413284] HOST ALERT: OPIE;UP;SOFT;6;PING OK - Packet loss = 58%, RTA = 0.36 ms
[1463413332] SERVICE ALERT: MBAMDB;Service: RPC;OK;SOFT;2;RpcSs: Started
[1463413332] SERVICE ALERT: MBAM;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1463413489] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463413609] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463413729] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1463413774] SERVICE ALERT: MBAM;Service: RPC;OK;HARD;3;RpcSs: Started
[1463413789] SERVICE ALERT: MBAM;CPU Load;OK;HARD;3;CPU Load 0% (15 min average)
[1463413804] SERVICE ALERT: ADWSUSCLIENTS;E:\ Drive Space;OK;HARD;3;e:\ - total: 299.87 Gb - used: 228.93 Gb (76%) - free 70.94 Gb (24%)
[1463413804] SERVICE ALERT: ADWSUSCLIENTS;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1463413804] SERVICE ALERT: ADWSUSCLIENTS;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1463413804] SERVICE ALERT: ADWSUSCLIENTS;VMWare Tools;OK;HARD;3;VMTools: Started
[1463413804] SERVICE ALERT: ADWSUSCLIENTS;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1463413804] SERVICE ALERT: ADWSUSCLIENTS;Service: Windows Time;OK;HARD;3;W32Time: Started
[1463413804] SERVICE ALERT: ADWSUSCLIENTS;Memory Usage;OK;HARD;3;Memory usage: total:14335.67 MB - used: 3505.88 MB (24%) - free: 10829.79 MB (76%)
[1463413806] SERVICE ALERT: MBAMDB;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1463413810] SERVICE ALERT: OFFICEKMS;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1463413815] SERVICE ALERT: MBAM;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1463413817] SERVICE ALERT: OPIE;Service: WMI;OK;HARD;3;Winmgmt: Started
[1463413870] SERVICE ALERT: OPIE;Service: Windows Time;OK;HARD;3;W32Time: Started
[1463413952] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463414072] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 83% (15 min average)
[1463414223] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1463415084] SERVICE ALERT: eateriespk;Device Status;OK;HARD;3;TRAY 3 TYPE=  ANY TYPE *PLAIN  PREPRINTED
[1463415088] HOST ALERT: eateriespk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.49 ms
[1463415294] SERVICE ALERT: eateriespk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1463415443] SERVICE ALERT: eateriespk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 34% - OK!
[1463417329] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463417929] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463418263] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 1 Trays Empty
[1463418323] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 1 Trays Empty
[1463418383] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 1 Trays Empty
[1463419129] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1463421602] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (28.9% change >= 20.0% threshold)
[1463421602] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1463421602] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1463421602] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1463421602] SERVICE DOWNTIME ALERT: NETIDADPROXY;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1463421602] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1463421602] HOST DOWNTIME ALERT: GOVGIS;STARTED; Host has entered a period of scheduled downtime
[1463422129] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1463422249] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1463422369] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463424769] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 91% (15 min average)
[1463425203] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (30.4% change >= 20.0% threshold)
[1463425203] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1463425203] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1463425203] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1463425203] SERVICE DOWNTIME ALERT: NETIDADPROXY;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1463425203] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1463425203] HOST DOWNTIME ALERT: GOVGIS;STARTED; Host has entered a period of scheduled downtime
[1463425369] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463425489] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463425609] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463426809] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463427383] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1463427409] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 56% (15 min average)
[1463428009] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1463428129] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463428249] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463428849] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463429449] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1463433382] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 1 Trays Empty
[1463433441] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 1 Trays Empty
[1463433501] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 1 Trays Empty
[1463440849] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1463442049] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463442649] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463446849] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463447691] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;SOFT;1;e:\ - total: 599.87 Gb - used: 486.76 Gb (81%) - free 113.11 Gb (19%)
[1463447811] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;SOFT;2;e:\ - total: 599.87 Gb - used: 486.76 Gb (81%) - free 113.11 Gb (19%)
[1463447931] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;HARD;3;e:\ - total: 599.87 Gb - used: 486.76 Gb (81%) - free 113.11 Gb (19%)
[1463448049] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463448049] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1463448647] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463449131] SERVICE ALERT: AD-RMAD;E:\ Drive Space;OK;HARD;3;e:\ - total: 599.87 Gb - used: 474.58 Gb (79%) - free 125.30 Gb (21%)
[1463449247] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1463451603] SERVICE ALERT: BBTS_DB;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463451646] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463451703] SERVICE ALERT: BBTS_DB;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1463452246] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1463452846] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463453070] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1463453190] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1463453446] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463453790] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1463453910] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1463454510] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463454630] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1463454750] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1463455951] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463456550] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 73% (15 min average)