[1453718065] SERVICE ALERT: halleground3;Paper Status Tray 2;OK;HARD;3;Tray 2 is at 40% - OK!
[1453724772] SERVICE ALERT: studentcenterpk;Paper Status Tray 2;UNKNOWN;HARD;3;Tray 2 status is UNKNOWN!
[1453727108] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1453727168] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 100% - OK!
[1453727208] SERVICE ALERT: hallefirst2;Paper Status Tray 3;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1453727238] HOST ALERT: hallefirst2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1453727242] SERVICE ALERT: studentcenterpk;Paper Status Tray 5;UNKNOWN;HARD;3;Tray 5 status is UNKNOWN!
[1453727266] SERVICE ALERT: hallefirst2;Paper Status Tray 3;CRITICAL;SOFT;2;Tray 3 is at CRITICAL level - please refill with more 11.0 x 8.5 paper.
[1453727270] HOST ALERT: hallefirst2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.53 ms
[1453727322] SERVICE ALERT: hallefirst2;Paper Status Tray 3;OK;SOFT;3;Tray 3 is at 40% - OK!
[1453727338] SERVICE ALERT: hallefirst2;Paper Status Tray 2;OK;HARD;3;Tray 2 is at 40% - OK!
[1453727602] SERVICE ALERT: markjeffersonpk;Paper Status Tray 3;CRITICAL;SOFT;1;Tray 3 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453727662] SERVICE ALERT: markjeffersonpk;Paper Status Tray 3;CRITICAL;SOFT;2;Tray 3 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453727723] SERVICE ALERT: markjeffersonpk;Paper Status Tray 3;CRITICAL;HARD;3;Tray 3 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453727922] SERVICE ALERT: hallefirst2;Paper Status Tray 3;CRITICAL;SOFT;1;Tray 3 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453727983] SERVICE ALERT: hallefirst2;Paper Status Tray 3;CRITICAL;SOFT;2;Tray 3 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453728042] SERVICE ALERT: owen101lab3;Paper Status Tray 3;CRITICAL;SOFT;1;Tray 3 is at CRITICAL level - please refill with more 11.0 x 8.5 paper.
[1453728043] SERVICE ALERT: hallefirst2;Paper Status Tray 3;CRITICAL;HARD;3;Tray 3 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453728102] SERVICE ALERT: owen101lab3;Paper Status Tray 3;OK;SOFT;2;Tray 3 is at 40% - OK!
[1453728245] SERVICE ALERT: owen101lab3;Maintenance Kit HP 110V-C9152A;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1453728299] SERVICE ALERT: owen101lab3;Maintenance Kit HP 110V-C9152A;OK;SOFT;2;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 73% - OK!
[1453728311] SERVICE ALERT: owen101lab3;Paper Status Tray 2;CRITICAL;SOFT;1;Tray 2 is at CRITICAL level - please refill with more 11.0 x 8.5 paper.
[1453728371] SERVICE ALERT: owen101lab3;Paper Status Tray 2;OK;SOFT;2;Tray 2 is at 40% - OK!
[1453728442] SERVICE ALERT: studentcenterpk;Paper Status Tray 5;CRITICAL;HARD;3;Tray 5 is at 5% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453728538] SERVICE ALERT: hallefirst2;Paper Status Tray 2;CRITICAL;SOFT;1;Tray 2 is at CRITICAL level - please refill with more 11.0 x 8.5 paper.
[1453728598] SERVICE ALERT: hallefirst2;Paper Status Tray 2;CRITICAL;SOFT;2;Tray 2 is at CRITICAL level - please refill with more 11.0 x 8.5 paper.
[1453728642] SERVICE ALERT: hallefirst2;Paper Status Tray 3;OK;HARD;3;Tray 3 is at 40% - OK!
[1453728659] SERVICE ALERT: hallefirst2;Paper Status Tray 2;CRITICAL;HARD;3;Tray 2 is at CRITICAL level - please refill with more 11.0 x 8.5 paper.
[1453728807] HOST ALERT: hallesecondsouth;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.45 ms
[1453728954] HOST ALERT: halleground1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1453729044] HOST ALERT: halleground1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1453729125] HOST ALERT: halleground1;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1453729215] HOST ALERT: halleground1;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1453729258] SERVICE ALERT: hallefirst2;Paper Status Tray 2;OK;HARD;3;Tray 2 is at 40% - OK!
[1453729305] HOST ALERT: halleground1;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1453729395] HOST ALERT: halleground1;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1453729442] SERVICE ALERT: mckennypk;Black Cartridge HP CC364X;OK;HARD;3;Black Cartridge HP CC364X is at 31% - OK!
[1453729446] HOST ALERT: mckennypk;UP;HARD;10;PING WARNING - Packet loss = 80%, RTA = 0.49 ms
[1453729468] SERVICE ALERT: mckennypk;Paper Status Tray 3;OK;HARD;3;Tray 3 is OK!
[1453729485] HOST ALERT: halleground1;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1453729548] HOST ALERT: halleground1;UP;SOFT;8;PING WARNING - Packet loss = 80%, RTA = 0.55 ms
[1453729689] SERVICE ALERT: halleground1;Paper Status Tray 5;OK;HARD;3;Tray 5 is at 40% - OK!
[1453729692] SERVICE ALERT: halleground1;Paper Status Tray 3;OK;HARD;3;Tray 3 is at 40% - OK!
[1453729940] SERVICE ALERT: ADAUTHDC1;AD Replication;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_nrpe, ...) failed. errno is 2: No such file or directory
[1453729958] SERVICE ALERT: mckennypk;Maintenance Kit HP 110V-CB388A;OK;HARD;3;Maintenance Kit HP 110V-CB388A, 220V-CB389A is at 10% - OK!
[1453730015] SERVICE ALERT: mckennypk;Paper Status Tray 2;OK;HARD;3;Tray 2 is OK!
[1453730060] SERVICE ALERT: ADAUTHDC1;AD Replication;CRITICAL;SOFT;2;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_nrpe, ...) failed. errno is 2: No such file or directory
[1453730118] SERVICE ALERT: halleground1;Paper Status Tray 2;OK;HARD;3;Tray 2 is at 40% - OK!
[1453730123] SERVICE ALERT: halleground1;Maintenance Kit HP 110V-C9152A;OK;HARD;3;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 51% - OK!
[1453730130] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 31% - OK!
[1453730180] SERVICE ALERT: ADAUTHDC1;AD Replication;CRITICAL;HARD;3;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_nrpe, ...) failed. errno is 2: No such file or directory
[1453730443] SERVICE ALERT: hallefirst2;Paper Status Tray 3;CRITICAL;SOFT;1;Tray 3 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453730502] SERVICE ALERT: hallefirst2;Paper Status Tray 3;CRITICAL;SOFT;2;Tray 3 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453730562] SERVICE ALERT: hallefirst2;Paper Status Tray 3;CRITICAL;HARD;3;Tray 3 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453731373] SERVICE ALERT: studentcenterpk;Paper Status Tray 2;CRITICAL;HARD;3;Tray 2 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453732254] SERVICE ALERT: ADAUTHDC1;AD Replication;UNKNOWN;HARD;3;Unknown command(s): addcreplicationcheck
[1453732416] SERVICE ALERT: ADAUTHDC1;AD Replication;OK;HARD;3;OK: replication is up and running
[1453732573] SERVICE ALERT: studentcenterpk;Paper Status Tray 2;UNKNOWN;HARD;3;Tray 2 status is UNKNOWN!
[1453732635] SERVICE ALERT: studentcenterpk;Paper Status Tray 3;UNKNOWN;HARD;3;Tray 3 status is UNKNOWN!
[1453733172] SERVICE ALERT: studentcenterpk;Paper Status Tray 2;OK;HARD;3;Tray 2 is at 40% - OK!
[1453733172] SERVICE FLAPPING ALERT: studentcenterpk;Paper Status Tray 2;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1453733769] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at 5% - WARNING!
[1453733828] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at 5% - WARNING!
[1453733887] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at 5% - WARNING!
[1453734052] SERVICE FLAPPING ALERT: studentcenterpk;Paper Status Tray 2;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1453734209] SERVICE FLAPPING ALERT: studentcenterpk;Paper Status Tray 2;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1453734322] SERVICE ALERT: markjeffersonpk;Paper Status Tray 3;OK;HARD;3;Tray 3 is at 10% - OK!
[1453736495] SERVICE ALERT: strongpk;Paper Status Tray 3;CRITICAL;SOFT;1;Tray 3 is at CRITICAL level - please refill with more 8.5 x 11.0 paper.
[1453736556] SERVICE ALERT: strongpk;Paper Status Tray 3;CRITICAL;SOFT;2;Tray 3 is at CRITICAL level - please refill with more 8.5 x 11.0 paper.
[1453736615] SERVICE ALERT: strongpk;Paper Status Tray 3;CRITICAL;HARD;3;Tray 3 is at CRITICAL level - please refill with more 8.5 x 11.0 paper.
[1453737044] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.81.90! Make sure host is up and SNMP is configured properly.
[1453737087] SERVICE ALERT: alexanderpk;Paper Status Tray 3;CRITICAL;SOFT;1;Tray 3 is at CRITICAL level - please refill with more 8.5 x 11.0 paper.
[1453737097] SERVICE ALERT: markjeffersonpk;Maintenance Kit HP 110V-C9152A;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.81.90! Make sure host is up and SNMP is configured properly.
[1453737104] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.81.90! Make sure host is up and SNMP is configured properly.
[1453737146] SERVICE ALERT: alexanderpk;Paper Status Tray 3;CRITICAL;SOFT;2;Tray 3 is at CRITICAL level - please refill with more 8.5 x 11.0 paper.
[1453737148] SERVICE ALERT: markjeffersonpk;Paper Status Tray 2;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.81.90! Make sure host is up and SNMP is configured properly.
[1453737157] SERVICE ALERT: markjeffersonpk;Maintenance Kit HP 110V-C9152A;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.81.90! Make sure host is up and SNMP is configured properly.
[1453737164] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;WARNING;HARD;3;WARNING: No SNMP response from 164.76.81.90! Make sure host is up and SNMP is configured properly.
[1453737206] SERVICE ALERT: alexanderpk;Paper Status Tray 3;CRITICAL;HARD;3;Tray 3 is at CRITICAL level - please refill with more 8.5 x 11.0 paper.
[1453737208] SERVICE ALERT: markjeffersonpk;Paper Status Tray 2;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.81.90! Make sure host is up and SNMP is configured properly.
[1453737215] SERVICE ALERT: strongpk;Paper Status Tray 3;OK;HARD;3;Tray 3 is OK!
[1453737217] SERVICE ALERT: markjeffersonpk;Maintenance Kit HP 110V-C9152A;WARNING;HARD;3;WARNING: No SNMP response from 164.76.81.90! Make sure host is up and SNMP is configured properly.
[1453737268] SERVICE ALERT: markjeffersonpk;Paper Status Tray 2;WARNING;HARD;3;WARNING: No SNMP response from 164.76.81.90! Make sure host is up and SNMP is configured properly.
[1453737298] HOST ALERT: markjeffersonpk;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1453737325] HOST ALERT: markjeffersonpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.56 ms
[1453737656] SERVICE ALERT: hallefirst2;Paper Status Tray 2;CRITICAL;SOFT;1;Tray 2 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453737716] SERVICE ALERT: hallefirst2;Paper Status Tray 2;CRITICAL;SOFT;2;Tray 2 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453737758] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 9% - OK!
[1453737777] SERVICE ALERT: hallefirst2;Paper Status Tray 2;CRITICAL;HARD;3;Tray 2 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453737811] SERVICE ALERT: markjeffersonpk;Maintenance Kit HP 110V-C9152A;OK;HARD;3;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 76% - OK!
[1453737862] SERVICE ALERT: markjeffersonpk;Paper Status Tray 2;OK;HARD;3;Tray 2 is at 40% - OK!
[1453738257] SERVICE ALERT: HDBACKUPS;D:\ Drive Space;UNKNOWN;SOFT;1;Free disk space : Invalid drive
[1453738377] SERVICE ALERT: HDBACKUPS;D:\ Drive Space;UNKNOWN;SOFT;2;Free disk space : Invalid drive
[1453738497] SERVICE ALERT: HDBACKUPS;D:\ Drive Space;UNKNOWN;HARD;3;Free disk space : Invalid drive
[1453738952] SERVICE ALERT: HDBACKUPS;F:\ Drive Space;UNKNOWN;SOFT;1;Free disk space : Invalid drive
[1453739072] SERVICE ALERT: HDBACKUPS;F:\ Drive Space;UNKNOWN;SOFT;2;Free disk space : Invalid drive
[1453739192] SERVICE ALERT: HDBACKUPS;F:\ Drive Space;UNKNOWN;HARD;3;Free disk space : Invalid drive
[1453739213] SERVICE ALERT: owen101lab2;Paper Status Tray 5;CRITICAL;SOFT;1;Tray 5 is at 5% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453739274] SERVICE ALERT: owen101lab2;Paper Status Tray 5;CRITICAL;SOFT;2;Tray 5 is at 5% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453739333] SERVICE ALERT: owen101lab2;Paper Status Tray 5;CRITICAL;HARD;3;Tray 5 is at 5% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453739562] SERVICE ALERT: hallefirst2;Paper Status Tray 3;OK;HARD;3;Tray 3 is at 40% - OK!
[1453739577] SERVICE ALERT: hallefirst2;Paper Status Tray 2;OK;HARD;3;Tray 2 is at 40% - OK!
[1453739577] SERVICE FLAPPING ALERT: hallefirst2;Paper Status Tray 2;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1453739824] SERVICE FLAPPING ALERT: hallefirst2;Paper Status Tray 2;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1453739883] SERVICE ALERT: ADDC1;AD Replication;UNKNOWN;SOFT;1;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1453740003] SERVICE ALERT: ADDC1;AD Replication;UNKNOWN;SOFT;2;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1453740122] SERVICE ALERT: ADDC1;AD Replication;UNKNOWN;HARD;3;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1453740392] SERVICE ALERT: HDBACKUPS;F:\ Drive Space;OK;HARD;3;f:\ - total: 3315.22 Gb - used: 1227.61 Gb (37%) - free 2087.61 Gb (63%)
[1453741143] SERVICE ALERT: prayharroldpk;Paper Status Tray 5;CRITICAL;SOFT;1;Tray 5 is at 5% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453741204] SERVICE ALERT: prayharroldpk;Paper Status Tray 5;CRITICAL;SOFT;2;Tray 5 is at 5% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453741263] SERVICE ALERT: prayharroldpk;Paper Status Tray 5;CRITICAL;HARD;3;Tray 5 is at 5% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453741881] SERVICE ALERT: hallesecondsouth;Paper Status Tray 3;CRITICAL;SOFT;1;Tray 3 is at CRITICAL level - please refill with more 11.0 x 8.5 paper.
[1453741936] SERVICE ALERT: hallesecondsouth;Paper Status Tray 5;UNKNOWN;SOFT;1;Tray 5 status is UNKNOWN!
[1453741940] SERVICE ALERT: hallesecondsouth;Paper Status Tray 3;CRITICAL;SOFT;2;Tray 3 is at CRITICAL level - please refill with more 11.0 x 8.5 paper.
[1453741995] SERVICE ALERT: hallesecondsouth;Paper Status Tray 5;UNKNOWN;SOFT;2;Tray 5 status is UNKNOWN!
[1453742001] SERVICE ALERT: hallesecondsouth;Paper Status Tray 3;CRITICAL;HARD;3;Tray 3 is at CRITICAL level - please refill with more 11.0 x 8.5 paper.
[1453742055] SERVICE ALERT: hallesecondsouth;Paper Status Tray 5;UNKNOWN;HARD;3;Tray 5 status is UNKNOWN!
[1453742338] SERVICE ALERT: owen101lab2;Paper Status Tray 5;WARNING;HARD;3;WARNING: No SNMP response from 164.76.203.102! Make sure host is up and SNMP is configured properly.
[1453742368] HOST ALERT: owen101lab2;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1453742458] HOST ALERT: owen101lab2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1453742459] HOST ALERT: owen101lab1;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.61 ms
[1453742548] HOST ALERT: owen101lab2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1453742607] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.203.102! Make sure host is up and SNMP is configured properly.
[1453742637] HOST ALERT: owen101lab2;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1453742667] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.203.102! Make sure host is up and SNMP is configured properly.
[1453742697] HOST ALERT: owen101lab2;DOWN;SOFT;5;(Host check timed out after 30.00 seconds)
[1453742700] SERVICE ALERT: crossroadspk;Paper Status Tray 3;CRITICAL;SOFT;1;Tray 3 is at CRITICAL level - please refill with more 8.5 x 11.0 paper.
[1453742719] SERVICE ALERT: owen101lab2;Paper Status Tray 2;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.203.102! Make sure host is up and SNMP is configured properly.
[1453742727] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;WARNING;HARD;3;WARNING: No SNMP response from 164.76.203.102! Make sure host is up and SNMP is configured properly.
[1453742749] HOST ALERT: owen101lab2;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1453742761] SERVICE ALERT: crossroadspk;Paper Status Tray 3;CRITICAL;SOFT;2;Tray 3 is at CRITICAL level - please refill with more 8.5 x 11.0 paper.
[1453742765] SERVICE ALERT: owen101lab2;Paper Status Tray 3;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.203.102! Make sure host is up and SNMP is configured properly.
[1453742779] SERVICE ALERT: owen101lab2;Paper Status Tray 2;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.203.102! Make sure host is up and SNMP is configured properly.
[1453742782] SERVICE ALERT: owen101lab2;Maintenance Kit HP 110V-C9152A;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.203.102! Make sure host is up and SNMP is configured properly.
[1453742795] HOST ALERT: owen101lab2;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1453742820] SERVICE ALERT: crossroadspk;Paper Status Tray 3;CRITICAL;HARD;3;Tray 3 is at CRITICAL level - please refill with more 8.5 x 11.0 paper.
[1453742825] SERVICE ALERT: owen101lab2;Paper Status Tray 3;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.203.102! Make sure host is up and SNMP is configured properly.
[1453742839] SERVICE ALERT: owen101lab2;Paper Status Tray 2;WARNING;HARD;3;WARNING: No SNMP response from 164.76.203.102! Make sure host is up and SNMP is configured properly.
[1453742842] SERVICE ALERT: owen101lab2;Maintenance Kit HP 110V-C9152A;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.203.102! Make sure host is up and SNMP is configured properly.
[1453742855] HOST ALERT: owen101lab2;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1453742885] SERVICE ALERT: owen101lab2;Paper Status Tray 3;WARNING;HARD;3;WARNING: No SNMP response from 164.76.203.102! Make sure host is up and SNMP is configured properly.
[1453742902] SERVICE ALERT: owen101lab2;Maintenance Kit HP 110V-C9152A;WARNING;HARD;3;WARNING: No SNMP response from 164.76.203.102! Make sure host is up and SNMP is configured properly.
[1453742915] HOST ALERT: owen101lab2;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1453742968] HOST ALERT: owen101lab2;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1453743370] SERVICE ALERT: owen101lab3;Paper Status Tray 2;CRITICAL;SOFT;1;Tray 2 is at CRITICAL level - please refill with more 11.0 x 8.5 paper.
[1453743435] SERVICE ALERT: owen101lab3;Paper Status Tray 2;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1453743465] HOST ALERT: owen101lab3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1453743495] SERVICE ALERT: owen101lab3;Paper Status Tray 2;WARNING;HARD;3;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1453743525] HOST ALERT: owen101lab3;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1453743589] HOST ALERT: owen101lab3;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.56 ms
[1453744089] SERVICE ALERT: owen101lab3;Paper Status Tray 2;OK;HARD;3;Tray 2 is at 40% - OK!
[1453744288] SERVICE ALERT: owen101lab3;Paper Status Tray 5;UNKNOWN;SOFT;1;UNKNOWN - Tray 5 not found. Use the TRAY TEST option to determine which trays may be monitored.
[1453744348] SERVICE ALERT: owen101lab3;Paper Status Tray 5;UNKNOWN;SOFT;2;UNKNOWN - Tray 5 not found. Use the TRAY TEST option to determine which trays may be monitored.
[1453744408] SERVICE ALERT: owen101lab3;Paper Status Tray 5;UNKNOWN;HARD;3;UNKNOWN - Tray 5 not found. Use the TRAY TEST option to determine which trays may be monitored.
[1453744472] HOST ALERT: owen101lab2;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.56 ms
[1453744521] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 23% - OK!
[1453744633] SERVICE ALERT: owen101lab2;Paper Status Tray 2;OK;HARD;3;Tray 2 is at 20% - OK!
[1453744679] SERVICE ALERT: owen101lab2;Paper Status Tray 3;OK;HARD;3;Tray 3 is at 40% - OK!
[1453744696] SERVICE ALERT: owen101lab2;Maintenance Kit HP 110V-C9152A;OK;HARD;3;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 69% - OK!
[1453744732] SERVICE ALERT: owen101lab2;Paper Status Tray 5;CRITICAL;HARD;3;Tray 5 is at 5% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453744842] SERVICE ALERT: hallefirst2;Paper Status Tray 5;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1453744902] SERVICE ALERT: hallefirst2;Paper Status Tray 5;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1453744962] SERVICE ALERT: hallefirst2;Paper Status Tray 5;WARNING;HARD;3;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1453744965] SERVICE ALERT: hallefirst2;Paper Status Tray 3;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1453744981] SERVICE ALERT: hallefirst2;Paper Status Tray 2;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1453745025] SERVICE ALERT: hallefirst2;Paper Status Tray 3;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1453745041] SERVICE ALERT: hallefirst2;Paper Status Tray 2;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1453745085] SERVICE ALERT: hallefirst2;Paper Status Tray 3;WARNING;HARD;3;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1453745101] SERVICE ALERT: hallefirst2;Paper Status Tray 2;WARNING;HARD;3;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1453745167] SERVICE ALERT: hallefirst2;Maintenance Kit HP 110V-C9152A;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1453745227] SERVICE ALERT: hallefirst2;Maintenance Kit HP 110V-C9152A;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1453745281] SERVICE ALERT: hallegroundcolor;Paper Status Tray 2;UNKNOWN;SOFT;1;Tray 2 status is UNKNOWN!
[1453745287] SERVICE ALERT: hallefirst2;Maintenance Kit HP 110V-C9152A;WARNING;HARD;3;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1453745341] SERVICE ALERT: hallegroundcolor;Paper Status Tray 2;OK;SOFT;2;Tray 2 is OK!
[1453745792] SERVICE ALERT: ADDC1;AD Replication;UNKNOWN;SOFT;1;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1453745847] SERVICE ALERT: ADDC2;AD Replication;UNKNOWN;SOFT;1;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1453745901] SERVICE ALERT: ADDC3;AD Replication;UNKNOWN;SOFT;1;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1453745912] SERVICE ALERT: ADDC1;AD Replication;UNKNOWN;SOFT;2;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1453745956] SERVICE ALERT: ADDC4;AD Replication;UNKNOWN;SOFT;1;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1453745967] SERVICE ALERT: ADDC2;AD Replication;UNKNOWN;SOFT;2;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1453746011] SERVICE ALERT: ADDC5;AD Replication;UNKNOWN;SOFT;1;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1453746021] SERVICE ALERT: ADDC3;AD Replication;UNKNOWN;SOFT;2;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1453746032] SERVICE ALERT: ADDC1;AD Replication;UNKNOWN;HARD;3;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1453746065] SERVICE ALERT: ADDC6;AD Replication;UNKNOWN;SOFT;1;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1453746076] SERVICE ALERT: ADDC4;AD Replication;UNKNOWN;SOFT;2;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1453746087] SERVICE ALERT: ADDC2;AD Replication;UNKNOWN;HARD;3;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1453746131] SERVICE ALERT: ADDC5;AD Replication;UNKNOWN;SOFT;2;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1453746141] SERVICE ALERT: ADDC3;AD Replication;UNKNOWN;HARD;3;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1453746185] SERVICE ALERT: ADDC6;AD Replication;UNKNOWN;SOFT;2;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1453746196] SERVICE ALERT: ADDC4;AD Replication;UNKNOWN;HARD;3;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1453746198] SERVICE ALERT: hallefirst1;Paper Status Tray 3;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.104.232! Make sure host is up and SNMP is configured properly.
[1453746251] SERVICE ALERT: ADDC5;AD Replication;UNKNOWN;HARD;3;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1453746253] SERVICE ALERT: halleground3;Paper Status Tray 5;CRITICAL;SOFT;1;Tray 5 is at CRITICAL level - please refill with more 11.0 x 8.5 paper.
[1453746257] SERVICE ALERT: hallefirst3;Paper Status Tray 5;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.104.180! Make sure host is up and SNMP is configured properly.
[1453746258] SERVICE ALERT: hallefirst1;Paper Status Tray 3;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.104.232! Make sure host is up and SNMP is configured properly.
[1453746267] SERVICE ALERT: hallefirst1;Paper Status Tray 2;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.104.232! Make sure host is up and SNMP is configured properly.
[1453746288] HOST ALERT: hallefirst1;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1453746305] SERVICE ALERT: ADDC6;AD Replication;UNKNOWN;HARD;3;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1453746314] SERVICE ALERT: halleground3;Paper Status Tray 5;CRITICAL;SOFT;2;Tray 5 is at CRITICAL level - please refill with more 11.0 x 8.5 paper.
[1453746317] SERVICE ALERT: hallefirst3;Paper Status Tray 5;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.104.180! Make sure host is up and SNMP is configured properly.
[1453746318] SERVICE ALERT: hallefirst1;Paper Status Tray 3;WARNING;HARD;3;WARNING: No SNMP response from 164.76.104.232! Make sure host is up and SNMP is configured properly.
[1453746322] HOST ALERT: hallefirst1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.54 ms
[1453746324] SERVICE ALERT: hallefirst1;Paper Status Tray 2;OK;SOFT;2;Tray 2 is at 40% - OK!
[1453746356] SERVICE ALERT: ADDC1;AD Replication;OK;HARD;3;OK: replication is up and running
[1453746371] SERVICE ALERT: hallefirst3;Paper Status Tray 5;OK;SOFT;3;Tray 5 is at 40% - OK!
[1453746373] SERVICE ALERT: halleground3;Paper Status Tray 5;CRITICAL;HARD;3;Tray 5 is at CRITICAL level - please refill with more 11.0 x 8.5 paper.
[1453746481] SERVICE ALERT: hallefirst2;Maintenance Kit HP 110V-C9152A;OK;HARD;3;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 14% - OK!
[1453746493] SERVICE ALERT: ADDC3;Service: Remote Management;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1453746498] SERVICE ALERT: ADDC2;AD Replication;OK;HARD;3;OK: replication is up and running
[1453746510] SERVICE ALERT: ADDC4;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1453746511] SERVICE ALERT: ADDC6;AD Replication;OK;HARD;3;OK: replication is up and running
[1453746522] SERVICE ALERT: ADDC5;AD Replication;OK;HARD;3;OK: replication is up and running
[1453746556] SERVICE ALERT: ADDC3;Service: VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1453746613] SERVICE ALERT: ADDC3;Service: Remote Management;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1453746619] SERVICE ALERT: ADDC4;Service: DNS Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1453746630] SERVICE ALERT: ADDC4;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1453746644] SERVICE ALERT: ADDC4;AD Replication;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 30 seconds.
[1453746665] SERVICE ALERT: ADDC3;AD Replication;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 30 seconds.
[1453746671] SERVICE ALERT: ADDC3;Service: Distributed File System Replication;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1453746671] SERVICE ALERT: ADDC3;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1453746675] SERVICE ALERT: ADDC4;Service: Distributed File System Replication;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1453746676] SERVICE ALERT: ADDC4;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1453746676] SERVICE ALERT: ADDC3;Service: VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1453746677] SERVICE ALERT: ADDC3;Service: DFS Namespace;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1453746728] SERVICE ALERT: ADDC3;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1453746729] SERVICE ALERT: ADDC3;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1453746730] SERVICE ALERT: ADDC4;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1453746730] SERVICE ALERT: ADDC4;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1453746731] SERVICE ALERT: ADDC3;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1453746733] SERVICE ALERT: ADDC3;Service: Remote Management;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1453746739] SERVICE ALERT: ADDC4;Service: DNS Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1453746739] SERVICE ALERT: ADDC3;Service: DNS Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1453746750] SERVICE ALERT: ADDC4;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1453746752] SERVICE ALERT: ADDC3;Service: AD Web Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1453746756] SERVICE ALERT: hallefirst2;Paper Status Tray 5;OK;HARD;3;Tray 5 is at 40% - OK!
[1453746777] SERVICE ALERT: ADDC3;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1453746780] SERVICE ALERT: ADDC3;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1453746783] SERVICE ALERT: ADDC4;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1453746785] SERVICE ALERT: ADDC4;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1453746785] SERVICE ALERT: ADDC4;Service: DFS Namespace;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1453746791] SERVICE ALERT: ADDC3;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1453746791] SERVICE ALERT: ADDC3;Service: Distributed File System Replication;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1453746795] SERVICE ALERT: ADDC4;Service: Distributed File System Replication;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1453746796] SERVICE ALERT: ADDC4;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1453746796] SERVICE ALERT: ADDC3;Service: VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1453746797] SERVICE ALERT: ADDC3;Service: DFS Namespace;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1453746799] SERVICE ALERT: ADDC3;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1453746813] SERVICE ALERT: ADDC4;Service: AD Web Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1453746821] SERVICE ALERT: ADDC3;AD Replication;OK;HARD;3;OK: replication is up and running
[1453746828] SERVICE ALERT: ADDC3;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1453746829] SERVICE ALERT: ADDC3;Uptime;OK;SOFT;2;System Uptime - 10 day(s) 0 hour(s) 10 minute(s)
[1453746830] SERVICE ALERT: ADDC4;Uptime;OK;SOFT;2;System Uptime - 10 day(s) 0 hour(s) 8 minute(s)
[1453746830] SERVICE ALERT: ADDC4;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1453746831] SERVICE ALERT: ADDC3;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1453746839] SERVICE ALERT: ADDC4;Service: DNS Service;OK;SOFT;3;DNS: Started
[1453746839] SERVICE ALERT: ADDC3;Service: DNS Service;OK;SOFT;2;DNS: Started
[1453746852] SERVICE ALERT: ADDC3;Service: AD Web Service;OK;SOFT;2;ADWS: Started
[1453746877] SERVICE ALERT: ADDC3;C:\ Drive Space;OK;SOFT;2;c:\ - total: 79.66 Gb - used: 10.78 Gb (14%) - free 68.88 Gb (86%)
[1453746880] SERVICE ALERT: ADDC3;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1453746883] SERVICE ALERT: ADDC4;C:\ Drive Space;OK;SOFT;2;c:\ - total: 79.66 Gb - used: 10.84 Gb (14%) - free 68.81 Gb (86%)
[1453746885] SERVICE ALERT: ADDC4;Service: DFS Namespace;OK;SOFT;2;Dfs: Started
[1453746885] SERVICE ALERT: ADDC4;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1453746891] SERVICE ALERT: ADDC3;Service: Distributed File System Replication;OK;SOFT;3;DFSR: Started
[1453746891] SERVICE ALERT: ADDC3;Service: Windows Time;OK;SOFT;3;W32Time: Started
[1453746895] SERVICE ALERT: ADDC4;Service: Distributed File System Replication;OK;SOFT;3;DFSR: Started
[1453746896] SERVICE ALERT: ADDC4;Service: Windows Time;OK;SOFT;3;W32Time: Started
[1453746897] SERVICE ALERT: ADDC3;Service: DFS Namespace;OK;SOFT;3;Dfs: Started
[1453746899] SERVICE ALERT: ADDC3;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1453746912] SERVICE ALERT: hallefirst1;Paper Status Tray 3;OK;HARD;3;Tray 3 is at 40% - OK!
[1453746913] SERVICE ALERT: ADDC4;Service: AD Web Service;OK;SOFT;2;ADWS: Started
[1453746972] SERVICE ALERT: prayharroldpk;Paper Status Tray 3;CRITICAL;SOFT;1;Tray 3 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453746977] SERVICE ALERT: hallefirst3;Paper Status Tray 5;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.104.180! Make sure host is up and SNMP is configured properly.
[1453747031] SERVICE ALERT: hallefirst3;Paper Status Tray 5;OK;SOFT;2;Tray 5 is at 40% - OK!
[1453747033] SERVICE ALERT: prayharroldpk;Paper Status Tray 3;CRITICAL;SOFT;2;Tray 3 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453747092] SERVICE ALERT: prayharroldpk;Paper Status Tray 3;CRITICAL;HARD;3;Tray 3 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453747219] SERVICE ALERT: ADDC4;AD Replication;OK;HARD;3;OK: replication is up and running
[1453747224] SERVICE ALERT: hallesecondsouth;Paper Status Tray 2;CRITICAL;SOFT;1;Tray 2 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453747283] SERVICE ALERT: hallesecondsouth;Paper Status Tray 2;CRITICAL;SOFT;2;Tray 2 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453747313] SERVICE ALERT: ADDC3;Service: Remote Management;OK;HARD;3;WinRM: Started
[1453747330] SERVICE ALERT: ADDC4;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1453747343] SERVICE ALERT: hallesecondsouth;Paper Status Tray 2;CRITICAL;HARD;3;Tray 2 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453747376] SERVICE ALERT: ADDC3;Service: VMWare Tools;OK;HARD;3;VMTools: Started
[1453747479] SERVICE ALERT: hallefirst2;Paper Status Tray 3;OK;HARD;3;Tray 3 is at 40% - OK!
[1453747495] SERVICE ALERT: hallefirst2;Paper Status Tray 2;OK;HARD;3;Tray 2 is at 40% - OK!
[1453747495] SERVICE FLAPPING ALERT: hallefirst2;Paper Status Tray 2;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1453748122] SERVICE ALERT: markjeffersonpk;Paper Status Tray 3;CRITICAL;SOFT;1;Tray 3 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453748181] SERVICE ALERT: markjeffersonpk;Paper Status Tray 3;CRITICAL;SOFT;2;Tray 3 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453748221] SERVICE ALERT: marshallpk;Paper Status Tray 2;CRITICAL;SOFT;1;Tray 2 is at CRITICAL level - please refill with more 8.5 x 11.0 paper.
[1453748242] SERVICE ALERT: markjeffersonpk;Paper Status Tray 3;CRITICAL;HARD;3;Tray 3 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453748280] SERVICE ALERT: marshallpk;Paper Status Tray 2;CRITICAL;SOFT;2;Tray 2 is at CRITICAL level - please refill with more 8.5 x 11.0 paper.
[1453748340] SERVICE ALERT: marshallpk;Paper Status Tray 2;CRITICAL;HARD;3;Tray 2 is at CRITICAL level - please refill with more 8.5 x 11.0 paper.
[1453748833] SERVICE ALERT: studentcenterpk;Paper Status Tray 3;CRITICAL;HARD;3;Tray 3 is at CRITICAL level - please refill with more 11.0 x 8.5 paper.
[1453749030] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at 5% - WARNING!
[1453749090] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at 5% - WARNING!
[1453749150] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at 5% - WARNING!
[1453749316] SERVICE ALERT: prayharroldpk;Paper Status Tray 2;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.88.118! Make sure host is up and SNMP is configured properly.
[1453749346] HOST ALERT: prayharroldpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1453749376] SERVICE ALERT: prayharroldpk;Paper Status Tray 2;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.88.118! Make sure host is up and SNMP is configured properly.
[1453749406] HOST ALERT: prayharroldpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1453749436] SERVICE ALERT: prayharroldpk;Paper Status Tray 2;WARNING;HARD;3;WARNING: No SNMP response from 164.76.88.118! Make sure host is up and SNMP is configured properly.
[1453749466] HOST ALERT: prayharroldpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1453749489] SERVICE ALERT: owen101lab3;Paper Status Tray 2;CRITICAL;SOFT;1;Tray 2 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453749498] SERVICE ALERT: prayharroldpk;Paper Status Tray 3;WARNING;HARD;3;WARNING: No SNMP response from 164.76.88.118! Make sure host is up and SNMP is configured properly.
[1453749528] HOST ALERT: prayharroldpk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1453749550] SERVICE ALERT: owen101lab3;Paper Status Tray 2;CRITICAL;SOFT;2;Tray 2 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453749609] SERVICE ALERT: owen101lab3;Paper Status Tray 2;CRITICAL;HARD;3;Tray 2 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453749618] HOST ALERT: prayharroldpk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1453749668] SERVICE ALERT: prayharroldpk;Paper Status Tray 5;WARNING;HARD;3;WARNING: No SNMP response from 164.76.88.118! Make sure host is up and SNMP is configured properly.
[1453749698] HOST ALERT: prayharroldpk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1453749788] HOST ALERT: prayharroldpk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1453749868] SERVICE ALERT: prayharroldpk;Maintenance Kit HP 110V-C9152A;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.88.118! Make sure host is up and SNMP is configured properly.
[1453749874] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.88.118! Make sure host is up and SNMP is configured properly.
[1453749878] HOST ALERT: prayharroldpk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1453749928] SERVICE ALERT: prayharroldpk;Maintenance Kit HP 110V-C9152A;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.88.118! Make sure host is up and SNMP is configured properly.
[1453749934] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.88.118! Make sure host is up and SNMP is configured properly.
[1453749958] HOST ALERT: prayharroldpk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1453749988] SERVICE ALERT: prayharroldpk;Maintenance Kit HP 110V-C9152A;WARNING;HARD;3;WARNING: No SNMP response from 164.76.88.118! Make sure host is up and SNMP is configured properly.
[1453749994] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;WARNING;HARD;3;WARNING: No SNMP response from 164.76.88.118! Make sure host is up and SNMP is configured properly.
[1453750018] HOST ALERT: prayharroldpk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1453750086] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;CRITICAL;HARD;3;Black Print Cartridge HP C8543X is at 1% - CRITICAL!
[1453751375] SERVICE ALERT: hallesecondsouth;Maintenance Kit HP 110V-C9152A;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.107.53! Make sure host is up and SNMP is configured properly.
[1453751390] HOST ALERT: hallesecondsouth;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1453751409] SERVICE ALERT: owen101lab3;Paper Status Tray 2;OK;HARD;3;Tray 2 is at 10% - OK!
[1453751409] SERVICE FLAPPING ALERT: owen101lab3;Paper Status Tray 2;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1453751435] SERVICE ALERT: hallesecondsouth;Maintenance Kit HP 110V-C9152A;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.107.53! Make sure host is up and SNMP is configured properly.
[1453751462] SERVICE ALERT: prayharroldpk;Paper Status Tray 5;OK;HARD;3;Tray 5 is at 40% - OK!
[1453751465] HOST ALERT: hallesecondsouth;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1453751466] HOST ALERT: prayharroldpk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 10.04 ms
[1453751495] SERVICE ALERT: hallesecondsouth;Maintenance Kit HP 110V-C9152A;WARNING;HARD;3;WARNING: No SNMP response from 164.76.107.53! Make sure host is up and SNMP is configured properly.
[1453751525] HOST ALERT: hallesecondsouth;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1453751549] SERVICE ALERT: hallesecondsouth;Paper Status Tray 2;WARNING;HARD;3;WARNING: No SNMP response from 164.76.107.53! Make sure host is up and SNMP is configured properly.
[1453751579] HOST ALERT: hallesecondsouth;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1453751605] SERVICE ALERT: hallesecondsouth;Paper Status Tray 3;WARNING;HARD;3;WARNING: No SNMP response from 164.76.107.53! Make sure host is up and SNMP is configured properly.
[1453751635] HOST ALERT: hallesecondsouth;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1453751661] SERVICE ALERT: hallesecondsouth;Paper Status Tray 5;WARNING;HARD;3;WARNING: No SNMP response from 164.76.107.53! Make sure host is up and SNMP is configured properly.
[1453751691] HOST ALERT: hallesecondsouth;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1453751781] HOST ALERT: hallesecondsouth;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1453751782] SERVICE ALERT: prayharroldpk;Maintenance Kit HP 110V-C9152A;OK;HARD;3;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 81% - OK!
[1453751788] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 53% - OK!
[1453751830] SERVICE ALERT: prayharroldpk;Paper Status Tray 2;OK;HARD;3;Tray 2 is at 40% - OK!
[1453751871] HOST ALERT: hallesecondsouth;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1453751892] SERVICE ALERT: prayharroldpk;Paper Status Tray 3;OK;HARD;3;Tray 3 is at 40% - OK!
[1453751920] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.107.53! Make sure host is up and SNMP is configured properly.
[1453751950] HOST ALERT: hallesecondsouth;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1453751980] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.107.53! Make sure host is up and SNMP is configured properly.
[1453752010] HOST ALERT: hallesecondsouth;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1453752040] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;WARNING;HARD;3;WARNING: No SNMP response from 164.76.107.53! Make sure host is up and SNMP is configured properly.
[1453752533] SERVICE ALERT: owen101lab2;Paper Status Tray 5;OK;HARD;3;Tray 5 is at 40% - OK!
[1453752847] SERVICE ALERT: MBAM;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1453752860] SERVICE ALERT: hallesecondsouth;Paper Status Tray 5;UNKNOWN;HARD;3;Tray 5 status is UNKNOWN!
[1453752864] HOST ALERT: hallesecondsouth;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.45 ms
[1453752967] SERVICE ALERT: MBAM;CPU Load;OK;SOFT;2;CPU Load 15% (5 min average)
[1453753235] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 7% - OK!
[1453753288] SERVICE ALERT: hallesecondsouth;Maintenance Kit HP 110V-C9152A;OK;HARD;3;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 10% - OK!
[1453753343] SERVICE ALERT: hallesecondsouth;Paper Status Tray 2;OK;HARD;3;Tray 2 is at 40% - OK!
[1453753363] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.81.90! Make sure host is up and SNMP is configured properly.
[1453753375] HOST ALERT: markjeffersonpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1453753399] SERVICE ALERT: hallesecondsouth;Paper Status Tray 3;OK;HARD;3;Tray 3 is at 40% - OK!
[1453753416] SERVICE ALERT: markjeffersonpk;Maintenance Kit HP 110V-C9152A;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.81.90! Make sure host is up and SNMP is configured properly.
[1453753423] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.81.90! Make sure host is up and SNMP is configured properly.
[1453753446] HOST ALERT: markjeffersonpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1453753467] SERVICE ALERT: markjeffersonpk;Paper Status Tray 2;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.81.90! Make sure host is up and SNMP is configured properly.
[1453753476] SERVICE ALERT: markjeffersonpk;Maintenance Kit HP 110V-C9152A;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.81.90! Make sure host is up and SNMP is configured properly.
[1453753483] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;WARNING;HARD;3;WARNING: No SNMP response from 164.76.81.90! Make sure host is up and SNMP is configured properly.
[1453753497] HOST ALERT: markjeffersonpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1453753523] SERVICE ALERT: markjeffersonpk;Paper Status Tray 5;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.81.90! Make sure host is up and SNMP is configured properly.
[1453753527] SERVICE ALERT: markjeffersonpk;Paper Status Tray 2;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.81.90! Make sure host is up and SNMP is configured properly.
[1453753536] SERVICE ALERT: markjeffersonpk;Maintenance Kit HP 110V-C9152A;WARNING;HARD;3;WARNING: No SNMP response from 164.76.81.90! Make sure host is up and SNMP is configured properly.
[1453753553] HOST ALERT: markjeffersonpk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1453753583] SERVICE ALERT: markjeffersonpk;Paper Status Tray 5;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.81.90! Make sure host is up and SNMP is configured properly.
[1453753587] SERVICE ALERT: markjeffersonpk;Paper Status Tray 2;WARNING;HARD;3;WARNING: No SNMP response from 164.76.81.90! Make sure host is up and SNMP is configured properly.
[1453753613] HOST ALERT: markjeffersonpk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1453753643] SERVICE ALERT: markjeffersonpk;Paper Status Tray 5;WARNING;HARD;3;WARNING: No SNMP response from 164.76.81.90! Make sure host is up and SNMP is configured properly.
[1453753647] SERVICE ALERT: markjeffersonpk;Paper Status Tray 3;WARNING;HARD;3;WARNING: No SNMP response from 164.76.81.90! Make sure host is up and SNMP is configured properly.
[1453753673] HOST ALERT: markjeffersonpk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1453753763] HOST ALERT: markjeffersonpk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1453753809] SERVICE ALERT: OFFICEKMS;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1453753853] HOST ALERT: markjeffersonpk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1453753929] SERVICE ALERT: OFFICEKMS;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1453753943] HOST ALERT: markjeffersonpk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1453754033] HOST ALERT: markjeffersonpk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1453755119] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at 5% - WARNING!
[1453755179] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at 5% - WARNING!
[1453755239] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at 5% - WARNING!
[1453755492] SERVICE ALERT: halleground1;Paper Status Tray 3;CRITICAL;SOFT;1;Tray 3 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453755551] SERVICE ALERT: halleground1;Paper Status Tray 3;CRITICAL;SOFT;2;Tray 3 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453755612] SERVICE ALERT: halleground1;Paper Status Tray 3;CRITICAL;HARD;3;Tray 3 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453755750] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;CRITICAL;HARD;3;Black Print Cartridge HP C8543X is at 1% - CRITICAL!
[1453756351] HOST ALERT: markjeffersonpk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.63 ms
[1453756513] HOST ALERT: markjeffersonpk;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1453756566] HOST ALERT: markjeffersonpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1453756570] SERVICE ALERT: studentcenterpk;Paper Status Tray 2;CRITICAL;SOFT;1;Tray 2 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453756591] HOST ALERT: markjeffersonpk;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.52 ms
[1453756631] SERVICE ALERT: studentcenterpk;Paper Status Tray 2;CRITICAL;SOFT;2;Tray 2 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453756673] HOST ALERT: markjeffersonpk;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1453756673] HOST FLAPPING ALERT: markjeffersonpk;STARTED; Host appears to have started flapping (22.5% change > 20.0% threshold)
[1453756686] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1453756690] SERVICE ALERT: studentcenterpk;Paper Status Tray 2;CRITICAL;HARD;3;Tray 2 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453756763] HOST ALERT: markjeffersonpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1453756853] HOST ALERT: markjeffersonpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1453756917] HOST ALERT: markjeffersonpk;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.53 ms
[1453756999] SERVICE ALERT: hallesecondsouth;Paper Status Tray 3;CRITICAL;SOFT;1;Tray 3 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453757060] SERVICE ALERT: hallesecondsouth;Paper Status Tray 3;CRITICAL;SOFT;2;Tray 3 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453757077] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 9% - OK!
[1453757119] SERVICE ALERT: hallesecondsouth;Paper Status Tray 3;CRITICAL;HARD;3;Tray 3 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453757130] SERVICE ALERT: markjeffersonpk;Maintenance Kit HP 110V-C9152A;OK;HARD;3;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 76% - OK!
[1453757181] SERVICE ALERT: markjeffersonpk;Paper Status Tray 2;OK;HARD;3;Tray 2 is at 40% - OK!
[1453757237] SERVICE ALERT: markjeffersonpk;Paper Status Tray 5;OK;HARD;3;Tray 5 is at 40% - OK!
[1453757241] SERVICE ALERT: markjeffersonpk;Paper Status Tray 3;CRITICAL;HARD;3;Tray 3 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453758330] HOST FLAPPING ALERT: markjeffersonpk;STOPPED; Host appears to have stopped flapping (4.0% change < 5.0% threshold)
[1453759537] SERVICE ALERT: markjeffersonpk;Maintenance Kit HP 110V-C9152A;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.81.90! Make sure host is up and SNMP is configured properly.
[1453759587] SERVICE ALERT: markjeffersonpk;Paper Status Tray 2;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.81.90! Make sure host is up and SNMP is configured properly.
[1453759596] SERVICE ALERT: markjeffersonpk;Maintenance Kit HP 110V-C9152A;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.81.90! Make sure host is up and SNMP is configured properly.
[1453759642] SERVICE ALERT: markjeffersonpk;Paper Status Tray 5;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.81.90! Make sure host is up and SNMP is configured properly.
[1453759647] SERVICE ALERT: markjeffersonpk;Paper Status Tray 2;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.81.90! Make sure host is up and SNMP is configured properly.
[1453759647] SERVICE ALERT: markjeffersonpk;Paper Status Tray 3;WARNING;HARD;3;WARNING: No SNMP response from 164.76.81.90! Make sure host is up and SNMP is configured properly.
[1453759647] SERVICE FLAPPING ALERT: markjeffersonpk;Paper Status Tray 3;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1453759655] SERVICE ALERT: markjeffersonpk;Maintenance Kit HP 110V-C9152A;WARNING;HARD;3;WARNING: No SNMP response from 164.76.81.90! Make sure host is up and SNMP is configured properly.
[1453759702] SERVICE ALERT: markjeffersonpk;Paper Status Tray 5;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.81.90! Make sure host is up and SNMP is configured properly.
[1453759707] SERVICE ALERT: markjeffersonpk;Paper Status Tray 2;WARNING;HARD;3;WARNING: No SNMP response from 164.76.81.90! Make sure host is up and SNMP is configured properly.
[1453759732] HOST ALERT: markjeffersonpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1453759762] SERVICE ALERT: markjeffersonpk;Paper Status Tray 5;WARNING;HARD;3;WARNING: No SNMP response from 164.76.81.90! Make sure host is up and SNMP is configured properly.
[1453759766] HOST ALERT: markjeffersonpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.53 ms
[1453759811] SERVICE ALERT: halleground1;Paper Status Tray 3;OK;HARD;3;Tray 3 is at 40% - OK!
[1453760083] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.81.90! Make sure host is up and SNMP is configured properly.
[1453760143] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.81.90! Make sure host is up and SNMP is configured properly.
[1453760161] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at 5% - WARNING!
[1453760165] SERVICE ALERT: hallefirst1;Paper Status Tray 5;CRITICAL;SOFT;1;Tray 5 is at 5% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453760203] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;WARNING;HARD;3;WARNING: No SNMP response from 164.76.81.90! Make sure host is up and SNMP is configured properly.
[1453760221] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at 5% - WARNING!
[1453760224] SERVICE ALERT: hallefirst1;Paper Status Tray 5;CRITICAL;SOFT;2;Tray 5 is at 5% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453760281] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at 5% - WARNING!
[1453760285] SERVICE ALERT: hallefirst1;Paper Status Tray 5;CRITICAL;HARD;3;Tray 5 is at 5% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453760411] SERVICE ALERT: halleground1;Paper Status Tray 3;CRITICAL;SOFT;1;Tray 3 is at CRITICAL level - please refill with more 11.0 x 8.5 paper.
[1453760471] SERVICE ALERT: halleground1;Paper Status Tray 3;CRITICAL;SOFT;2;Tray 3 is at CRITICAL level - please refill with more 11.0 x 8.5 paper.
[1453760532] SERVICE ALERT: halleground1;Paper Status Tray 3;CRITICAL;HARD;3;Tray 3 is at CRITICAL level - please refill with more 11.0 x 8.5 paper.
[1453760549] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1453760696] HOST ALERT: markjeffersonpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1453760786] HOST ALERT: markjeffersonpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1453760833] HOST ALERT: markjeffersonpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1453760877] HOST ALERT: markjeffersonpk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1453760937] HOST ALERT: markjeffersonpk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1453760992] HOST ALERT: markjeffersonpk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1453761082] HOST ALERT: markjeffersonpk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1453761172] HOST ALERT: markjeffersonpk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1453761262] HOST ALERT: markjeffersonpk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1453761352] HOST ALERT: markjeffersonpk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1453762208] SERVICE ALERT: owen101lab3;Paper Status Tray 2;CRITICAL;SOFT;1;Tray 2 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453762269] SERVICE ALERT: owen101lab3;Paper Status Tray 2;CRITICAL;SOFT;2;Tray 2 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453762328] SERVICE ALERT: owen101lab3;Paper Status Tray 2;CRITICAL;HARD;3;Tray 2 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453762414] SERVICE ALERT: strongpk;Paper Status Tray 3;CRITICAL;SOFT;1;Tray 3 is at CRITICAL level - please refill with more 8.5 x 11.0 paper.
[1453762474] SERVICE ALERT: strongpk;Paper Status Tray 3;CRITICAL;SOFT;2;Tray 3 is at CRITICAL level - please refill with more 8.5 x 11.0 paper.
[1453762535] SERVICE ALERT: strongpk;Paper Status Tray 3;CRITICAL;HARD;3;Tray 3 is at CRITICAL level - please refill with more 8.5 x 11.0 paper.
[1453762683] SERVICE ALERT: hallefirst1;Paper Status Tray 5;OK;HARD;3;Tray 5 is at 40% - OK!
[1453762943] SERVICE ALERT: hallesecondsouth;Paper Status Tray 2;UNKNOWN;SOFT;1;Tray 2 status is UNKNOWN!
[1453763003] SERVICE ALERT: hallesecondsouth;Paper Status Tray 2;UNKNOWN;SOFT;2;Tray 2 status is UNKNOWN!
[1453763063] SERVICE ALERT: hallesecondsouth;Paper Status Tray 2;OK;SOFT;3;Tray 2 is at 40% - OK!
[1453763118] SERVICE ALERT: hallesecondsouth;Paper Status Tray 3;OK;HARD;3;Tray 3 is at 40% - OK!
[1453763528] SERVICE ALERT: owen101lab3;Paper Status Tray 2;OK;HARD;3;Tray 2 is at 40% - OK!
[1453763696] HOST ALERT: markjeffersonpk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.70 ms
[1453763797] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 9% - OK!
[1453763797] SERVICE FLAPPING ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1453763842] SERVICE ALERT: markjeffersonpk;Paper Status Tray 3;CRITICAL;HARD;3;Tray 3 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453763846] SERVICE ALERT: markjeffersonpk;Maintenance Kit HP 110V-C9152A;OK;HARD;3;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 76% - OK!
[1453763846] SERVICE FLAPPING ALERT: markjeffersonpk;Maintenance Kit HP 110V-C9152A;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1453763901] SERVICE ALERT: markjeffersonpk;Paper Status Tray 2;OK;HARD;3;Tray 2 is at 20% - OK!
[1453763901] SERVICE FLAPPING ALERT: markjeffersonpk;Paper Status Tray 2;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1453763956] SERVICE ALERT: markjeffersonpk;Paper Status Tray 5;OK;HARD;3;Tray 5 is at 40% - OK!
[1453763956] SERVICE FLAPPING ALERT: markjeffersonpk;Paper Status Tray 5;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1453764433] SERVICE ALERT: studentcenterpk;Paper Status Tray 3;OK;HARD;3;Tray 3 is at 10% - OK!
[1453764917] SERVICE ALERT: halleground1;Paper Status Tray 2;CRITICAL;SOFT;1;Tray 2 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453764976] SERVICE ALERT: halleground1;Paper Status Tray 2;CRITICAL;SOFT;2;Tray 2 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453765032] SERVICE ALERT: studentcenterpk;Paper Status Tray 3;CRITICAL;SOFT;1;Tray 3 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453765037] SERVICE ALERT: halleground1;Paper Status Tray 2;CRITICAL;HARD;3;Tray 2 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453765093] SERVICE ALERT: studentcenterpk;Paper Status Tray 3;CRITICAL;SOFT;2;Tray 3 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453765152] SERVICE ALERT: studentcenterpk;Paper Status Tray 3;CRITICAL;HARD;3;Tray 3 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453765478] SERVICE ALERT: hallefirst2;Paper Status Tray 3;CRITICAL;SOFT;1;Tray 3 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453765539] SERVICE ALERT: hallefirst2;Paper Status Tray 3;CRITICAL;SOFT;2;Tray 3 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453765598] SERVICE ALERT: hallefirst2;Paper Status Tray 3;CRITICAL;HARD;3;Tray 3 is at 0% - CRITICAL! Please refill with more 11.0 x 8.5 paper.
[1453766352] SERVICE ALERT: studentcenterpk;Paper Status Tray 3;UNKNOWN;HARD;3;Tray 3 status is UNKNOWN!
[1453766798] SERVICE ALERT: hallefirst2;Paper Status Tray 3;OK;HARD;3;Tray 3 is at 40% - OK!
[1453766889] SERVICE ALERT: studentcenterpk;Paper Status Tray 2;UNKNOWN;HARD;3;Tray 2 status is UNKNOWN!
[1453767489] SERVICE ALERT: studentcenterpk;Paper Status Tray 2;CRITICAL;HARD;3;Tray 2 is at CRITICAL level - please refill with more 11.0 x 8.5 paper.
[1453767552] SERVICE ALERT: studentcenterpk;Paper Status Tray 3;OK;HARD;3;Tray 3 is at 40% - OK!
[1453767552] SERVICE FLAPPING ALERT: studentcenterpk;Paper Status Tray 3;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1453768961] SERVICE ALERT: hallefirst2;Paper Status Tray 5;UNKNOWN;SOFT;1;UNKNOWN - Tray 5 not found. Use the TRAY TEST option to determine which trays may be monitored.
[1453768991] HOST ALERT: hallefirst2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1453769015] SERVICE ALERT: hallefirst2;Paper Status Tray 5;OK;SOFT;2;Tray 5 is at 100% - OK!
[1453769019] HOST ALERT: hallefirst2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.54 ms
[1453769095] SERVICE ALERT: hallefirst2;Paper Status Tray 2;CRITICAL;SOFT;1;Tray 2 is at CRITICAL level - please refill with more 11.0 x 8.5 paper.
[1453769154] SERVICE ALERT: hallefirst2;Paper Status Tray 2;CRITICAL;SOFT;2;Tray 2 is at CRITICAL level - please refill with more 11.0 x 8.5 paper.
[1453769198] SERVICE ALERT: hallefirst2;Paper Status Tray 3;CRITICAL;SOFT;1;Tray 3 is at CRITICAL level - please refill with more 11.0 x 8.5 paper.
[1453769215] SERVICE ALERT: hallefirst2;Paper Status Tray 2;OK;SOFT;3;Tray 2 is at 40% - OK!
[1453769259] SERVICE ALERT: hallefirst2;Paper Status Tray 3;OK;SOFT;2;Tray 3 is at 40% - OK!
[1453771641] SERVICE FLAPPING ALERT: markjeffersonpk;Paper Status Tray 3;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1453771645] SERVICE FLAPPING ALERT: markjeffersonpk;Maintenance Kit HP 110V-C9152A;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1453771701] SERVICE FLAPPING ALERT: markjeffersonpk;Paper Status Tray 2;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1453771756] SERVICE FLAPPING ALERT: markjeffersonpk;Paper Status Tray 5;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1453772196] SERVICE FLAPPING ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1453774569] SERVICE ALERT: halleground3;Paper Status Tray 5;OK;HARD;3;Tray 5 is at 40% - OK!
[1453774936] SERVICE ALERT: halleground1;Paper Status Tray 3;WARNING;HARD;3;WARNING: No SNMP response from 164.76.106.51! Make sure host is up and SNMP is configured properly.
[1453774966] HOST ALERT: halleground1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1453775054] HOST ALERT: halleground1;UP;SOFT;2;PING WARNING - Packet loss = 82%, RTA = 0.56 ms
[1453775169] SERVICE ALERT: halleground3;Paper Status Tray 5;CRITICAL;SOFT;1;Tray 5 is at CRITICAL level - please refill with more 11.0 x 8.5 paper.
[1453775229] SERVICE ALERT: halleground3;Paper Status Tray 5;CRITICAL;SOFT;2;Tray 5 is at CRITICAL level - please refill with more 11.0 x 8.5 paper.
[1453775236] SERVICE ALERT: halleground1;Paper Status Tray 2;OK;HARD;3;Tray 2 is at 40% - OK!
[1453775290] SERVICE ALERT: halleground3;Paper Status Tray 5;CRITICAL;HARD;3;Tray 5 is at CRITICAL level - please refill with more 11.0 x 8.5 paper.
[1453775530] SERVICE ALERT: halleground1;Paper Status Tray 3;OK;HARD;3;Tray 3 is at 40% - OK!
[1453781705] SERVICE ALERT: halleground2;Paper Status Tray 3;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.106.52! Make sure host is up and SNMP is configured properly.
[1453781735] HOST ALERT: halleground2;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1453781765] SERVICE ALERT: halleground2;Paper Status Tray 3;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.106.52! Make sure host is up and SNMP is configured properly.
[1453781774] HOST ALERT: halleground2;UP;SOFT;2;PING OK - Packet loss = 50%, RTA = 1.77 ms
[1453781819] SERVICE ALERT: halleground2;Paper Status Tray 3;OK;SOFT;3;Tray 3 is at 40% - OK!