[1460435823] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1460436423] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460436816] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;1;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1460436927] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;2;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1460437047] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;HARD;3;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1460440023] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 73% (5 min average)
[1460440140] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460440260] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 95% (5 min average)
[1460440380] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 94% (5 min average)
[1460440623] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460440743] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460440863] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1460440863] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1460440980] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460441580] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 93% (5 min average)
[1460441723] SERVICE ALERT: INTLAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460441843] SERVICE ALERT: INTLAPP;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1460442180] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 97% (5 min average)
[1460442393] SERVICE ALERT: VUM;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460442513] SERVICE ALERT: VUM;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1460444247] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;OK;HARD;3;OK: No entries found
[1460444580] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 93% (5 min average)
[1460445180] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1460445780] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1460446380] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 96% (5 min average)
[1460446980] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 80% (5 min average)
[1460449863] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 69% (5 min average)
[1460451663] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1460451783] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 84% (5 min average)
[1460452383] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1460452503] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460452623] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460454423] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 42% (5 min average)
[1460455023] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1460455143] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 91% (5 min average)
[1460455263] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 51% (5 min average)
[1460455863] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460455983] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1460456103] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1460458379] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1460459103] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 61% (5 min average)
[1460459703] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460459823] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460459943] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460460296] SERVICE ALERT: IMAGES;Service: MDT Monitor;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460460326] HOST ALERT: IMAGES;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1460460340] SERVICE ALERT: IMAGES;Service: PXE Boot;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460460370] HOST ALERT: IMAGES;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460460383] HOST ALERT: IMAGES;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1460460386] SERVICE ALERT: IMAGES;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460460387] SERVICE ALERT: IMAGES;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460460440] SERVICE ALERT: IMAGES;Service: PXE Boot;CRITICAL;SOFT;2;WDSServer: Stopped
[1460460486] SERVICE ALERT: IMAGES;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1460460487] SERVICE ALERT: IMAGES;C:\ Drive Space;OK;SOFT;2;c:\ - total: 1157.56 Gb - used: 547.99 Gb (47%) - free 609.58 Gb (53%)
[1460460526] SERVICE ALERT: IMAGES;Service: PXE Boot;CRITICAL;HARD;3;WDSServer: Stopped
[1460460526] SERVICE ALERT: IMAGES;PXE Boot Service;CRITICAL;SOFT;1;WDSServer: Stopped
[1460460526] SERVICE ALERT: IMAGES;Service: MDT Monitor;UNKNOWN;HARD;3;Failed to open service MDT_Monitor: 1060: The specified service does not exist as an installed service.
[1460460646] SERVICE ALERT: IMAGES;PXE Boot Service;CRITICAL;SOFT;2;WDSServer: Stopped
[1460460766] SERVICE ALERT: IMAGES;PXE Boot Service;OK;SOFT;3;WDSServer: Started
[1460460776] SERVICE ALERT: IMAGES;Service: PXE Boot;OK;HARD;3;WDSServer: Started
[1460461953] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1460462014] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1460462073] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1460463448] SERVICE ALERT: AD-RMAD;E:\ Drive Space;OK;HARD;3;e:\ - total: 499.87 Gb - used: 378.96 Gb (76%) - free 120.91 Gb (24%)
[1460463534] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1460463543] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (5 min average)
[1460463593] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1460463653] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;SOFT;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1460464474] SERVICE ALERT: marshallpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1460465073] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1460465134] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1460465193] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1460465794] SERVICE ALERT: marshallpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1460465794] SERVICE FLAPPING ALERT: marshallpk;Paper Tray Status;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1460465943] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460466063] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 87% (5 min average)
[1460466663] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460466783] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460466903] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460467994] SERVICE ALERT: SCCM-PRIM01;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.102.80 and port 12489: No route to host
[1460467994] SERVICE ALERT: SCCM-PRIM01;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.102.80 and port 12489: No route to host
[1460467997] SERVICE ALERT: SCCM-PRIM01;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.102.80 and port 12489: No route to host
[1460467997] HOST ALERT: SCCM-PRIM01;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.80)
[1460468008] SERVICE ALERT: SCCM-PRIM01;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.102.80 and port 12489: No route to host
[1460468031] HOST DOWNTIME ALERT: SCCM-PRIM01;STARTED; Host has entered a period of scheduled downtime
[1460468052] SERVICE ALERT: SCCM-PRIM01;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.102.80 and port 12489: No route to host
[1460468055] HOST ALERT: SCCM-PRIM01;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.80)
[1460468106] SERVICE ALERT: SCCM-PRIM01;Service: Firewall;CRITICAL;SOFT;2;connect to address 164.76.102.80 and port 12489: No route to host
[1460468109] HOST ALERT: SCCM-PRIM01;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.102.80)
[1460468109] SERVICE ALERT: SCCM-PRIM01;C:\ Drive Space;CRITICAL;SOFT;2;connect to address 164.76.102.80 and port 12489: No route to host
[1460468118] SERVICE ALERT: SCCM-PRIM01;Service: RPC;CRITICAL;SOFT;2;connect to address 164.76.102.80 and port 12489: No route to host
[1460468118] SERVICE ALERT: SCCM-PRIM01;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.102.80 and port 12489: No route to host
[1460468128] SERVICE ALERT: SCCM-PRIM01;Service: Windows Time;CRITICAL;SOFT;2;connect to address 164.76.102.80 and port 12489: No route to host
[1460468131] HOST ALERT: SCCM-PRIM01;DOWN;SOFT;4;CRITICAL - Host Unreachable (164.76.102.80)
[1460468166] SERVICE ALERT: SCCM-PRIM01;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.102.80 and port 12489: No route to host
[1460468169] HOST ALERT: SCCM-PRIM01;DOWN;SOFT;5;CRITICAL - Host Unreachable (164.76.102.80)
[1460468172] SERVICE ALERT: SCCM-PRIM01;Memory Usage;CRITICAL;SOFT;2;connect to address 164.76.102.80 and port 12489: No route to host
[1460468226] SERVICE ALERT: SCCM-PRIM01;Service: Firewall;CRITICAL;HARD;3;connect to address 164.76.102.80 and port 12489: No route to host
[1460468229] HOST ALERT: SCCM-PRIM01;DOWN;SOFT;6;CRITICAL - Host Unreachable (164.76.102.80)
[1460468229] SERVICE ALERT: SCCM-PRIM01;C:\ Drive Space;CRITICAL;HARD;3;connect to address 164.76.102.80 and port 12489: No route to host
[1460468238] SERVICE ALERT: SCCM-PRIM01;NSClient++ Version;CRITICAL;SOFT;2;connect to address 164.76.102.80 and port 12489: No route to host
[1460468238] SERVICE ALERT: SCCM-PRIM01;Service: RPC;CRITICAL;HARD;3;connect to address 164.76.102.80 and port 12489: No route to host
[1460468248] SERVICE ALERT: SCCM-PRIM01;Service: Windows Time;CRITICAL;HARD;3;connect to address 164.76.102.80 and port 12489: No route to host
[1460468251] HOST ALERT: SCCM-PRIM01;DOWN;SOFT;7;CRITICAL - Host Unreachable (164.76.102.80)
[1460468286] SERVICE ALERT: SCCM-PRIM01;Service: DNS Client;CRITICAL;SOFT;2;connect to address 164.76.102.80 and port 12489: No route to host
[1460468289] HOST ALERT: SCCM-PRIM01;DOWN;SOFT;8;CRITICAL - Host Unreachable (164.76.102.80)
[1460468292] SERVICE ALERT: SCCM-PRIM01;Memory Usage;CRITICAL;HARD;3;connect to address 164.76.102.80 and port 12489: No route to host
[1460468358] SERVICE ALERT: SCCM-PRIM01;NSClient++ Version;CRITICAL;HARD;3;connect to address 164.76.102.80 and port 12489: No route to host
[1460468361] HOST ALERT: SCCM-PRIM01;DOWN;SOFT;9;CRITICAL - Host Unreachable (164.76.102.80)
[1460468406] SERVICE ALERT: SCCM-PRIM01;Service: DNS Client;CRITICAL;HARD;3;connect to address 164.76.102.80 and port 12489: No route to host
[1460468409] HOST ALERT: SCCM-PRIM01;DOWN;HARD;10;CRITICAL - Host Unreachable (164.76.102.80)
[1460468496] SERVICE ALERT: SCCM-PRIM01;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.102.80 and port 12489: No route to host
[1460468531] SERVICE ALERT: SCCM-PRIM01;Uptime;CRITICAL;SOFT;1;connect to address 164.76.102.80 and port 12489: No route to host
[1460468534] SERVICE ALERT: SCCM-PRIM01;Service: User Profile;CRITICAL;SOFT;1;connect to address 164.76.102.80 and port 12489: No route to host
[1460468546] SERVICE ALERT: SCCM-PRIM01;Service: NetLogin;CRITICAL;SOFT;1;connect to address 164.76.102.80 and port 12489: No route to host
[1460468546] SERVICE ALERT: SCCM-PRIM01;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.102.80 and port 12489: No route to host
[1460468616] SERVICE ALERT: SCCM-PRIM01;Service: Group Policy;CRITICAL;SOFT;2;connect to address 164.76.102.80 and port 12489: No route to host
[1460468651] SERVICE ALERT: SCCM-PRIM01;Uptime;CRITICAL;SOFT;2;connect to address 164.76.102.80 and port 12489: No route to host
[1460468654] SERVICE ALERT: SCCM-PRIM01;Service: User Profile;CRITICAL;SOFT;2;connect to address 164.76.102.80 and port 12489: No route to host
[1460468666] SERVICE ALERT: SCCM-PRIM01;Service: WMI;CRITICAL;SOFT;2;connect to address 164.76.102.80 and port 12489: No route to host
[1460468675] SERVICE ALERT: SCCM-PRIM01;Service: NetLogin;CRITICAL;SOFT;2;connect to address 164.76.102.80 and port 12489: No route to host
[1460468736] SERVICE ALERT: SCCM-PRIM01;Service: Group Policy;CRITICAL;HARD;3;connect to address 164.76.102.80 and port 12489: No route to host
[1460468771] SERVICE ALERT: SCCM-PRIM01;Uptime;CRITICAL;HARD;3;connect to address 164.76.102.80 and port 12489: No route to host
[1460468774] SERVICE ALERT: SCCM-PRIM01;Service: User Profile;CRITICAL;HARD;3;connect to address 164.76.102.80 and port 12489: No route to host
[1460468786] SERVICE ALERT: SCCM-PRIM01;Service: NetLogin;CRITICAL;HARD;3;connect to address 164.76.102.80 and port 12489: No route to host
[1460468786] SERVICE ALERT: SCCM-PRIM01;Service: WMI;CRITICAL;HARD;3;connect to address 164.76.102.80 and port 12489: No route to host
[1460468939] SERVICE ALERT: EMU-SCCM01;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1460469059] SERVICE ALERT: EMU-SCCM01;CPU Load;CRITICAL;SOFT;2;CPU Load 97% (5 min average)
[1460469182] SERVICE ALERT: EMU-SCCM01;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460469199] SERVICE ALERT: EMU-SCCM01;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460469212] HOST ALERT: EMU-SCCM01;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460469301] SERVICE ALERT: EMU-SCCM01;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460469301] HOST ALERT: EMU-SCCM01;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460469333] SERVICE ALERT: owen101lab3;Printer Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1460469363] HOST ALERT: owen101lab3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460469380] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1460469391] HOST ALERT: EMU-SCCM01;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460469393] SERVICE ALERT: owen101lab3;Printer Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1460469410] SERVICE ALERT: EMU-SCCM01;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460469421] SERVICE ALERT: EMU-SCCM01;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460469423] HOST ALERT: owen101lab3;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460469440] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1460469440] HOST ALERT: EMU-SCCM01;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460469447] SERVICE ALERT: owen101lab3;Printer Status;OK;SOFT;3;Warming up
[1460469451] HOST ALERT: owen101lab3;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.52 ms
[1460469461] SERVICE ALERT: EMU-SCCM01;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460469486] SERVICE ALERT: EMU-SCCM01;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460469491] HOST ALERT: EMU-SCCM01;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460469494] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;OK;SOFT;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1460469503] SERVICE ALERT: EMU-SCCM01;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460469508] SERVICE ALERT: EMU-SCCM01;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460469509] SERVICE ALERT: EMU-SCCM01;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460469512] SERVICE ALERT: EMU-SCCM01;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460469530] SERVICE ALERT: EMU-SCCM01;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460469533] HOST ALERT: EMU-SCCM01;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1460469563] SERVICE ALERT: EMU-SCCM01;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460469581] SERVICE ALERT: EMU-SCCM01;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460469593] HOST ALERT: EMU-SCCM01;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1460469596] HOST DOWNTIME ALERT: EMU-SCCM01;STARTED; Host has entered a period of scheduled downtime
[1460469598] SERVICE ALERT: EMU-SCCM01;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460469603] SERVICE FLAPPING ALERT: marshallpk;Paper Tray Status;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1460469603] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1460469603] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1460469603] HOST DOWNTIME ALERT: SCCM-PRIM01;STARTED; Host has entered a period of scheduled downtime
[1460469696] SERVICE ALERT: SCCM-PRIM01;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1460469700] HOST ALERT: SCCM-PRIM01;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.47 ms
[1460469729] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1460469789] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1460469819] HOST ALERT: owen101lab3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460469849] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1460469879] HOST ALERT: owen101lab3;UP;SOFT;2;PING WARNING - Packet loss = 86%, RTA = 0.54 ms
[1460469879] HOST FLAPPING ALERT: owen101lab3;STARTED; Host appears to have started flapping (22.6% change > 20.0% threshold)
[1460469903] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 37% (5 min average)
[1460469903] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (24.9% change >= 20.0% threshold)
[1460469932] SERVICE ALERT: SCCM-PRIM01;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1460469968] SERVICE ALERT: SCCM-PRIM01;Uptime;OK;HARD;3;System Uptime - 0 day(s) 18 hour(s) 5 minute(s)
[1460469973] SERVICE ALERT: SCCM-PRIM01;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1460469983] SERVICE ALERT: SCCM-PRIM01;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1460469983] SERVICE ALERT: SCCM-PRIM01;Service: WMI;OK;HARD;3;Winmgmt: Started
[1460470023] SERVICE ALERT: SCCM-PRIM01;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1460470026] SERVICE ALERT: SCCM-PRIM01;C:\ Drive Space;OK;HARD;3;c:\ - total: 194.90 Gb - used: 129.98 Gb (67%) - free 64.92 Gb (33%)
[1460470037] SERVICE ALERT: SCCM-PRIM01;Service: RPC;OK;HARD;3;RpcSs: Started
[1460470045] SERVICE ALERT: SCCM-PRIM01;Service: Windows Time;OK;HARD;3;W32Time: Started
[1460470089] SERVICE ALERT: SCCM-PRIM01;Memory Usage;OK;HARD;3;Memory usage: total:32756.36 MB - used: 6041.13 MB (18%) - free: 26715.24 MB (82%)
[1460470155] SERVICE ALERT: SCCM-PRIM01;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1460470444] SERVICE ALERT: owen101lab3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1460470979] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1460471099] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 83% (5 min average)
[1460471341] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 1 Trays Empty
[1460471401] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 1 Trays Empty
[1460471460] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 1 Trays Empty
[1460472615] HOST FLAPPING ALERT: owen101lab3;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1460472661] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1460473816] SERVICE ALERT: owen101lab2;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1460473876] SERVICE ALERT: owen101lab2;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1460473900] SERVICE ALERT: owen101lab1;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460473905] SERVICE ALERT: owen101lab1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460473906] HOST ALERT: owen101lab2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460473914] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460473929] HOST ALERT: owen101lab2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.49 ms
[1460473931] SERVICE ALERT: owen101lab2;Maintenance Kit HP 110V-C9152A;OK;SOFT;3;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 58% - OK!
[1460473953] SERVICE ALERT: owen101lab1;Maintenance Kit HP 110V-C9152A;OK;SOFT;2;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 55% - OK!
[1460473959] SERVICE ALERT: owen101lab1;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! 2 of 3 Trays Still Full
[1460473968] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 36% - OK!
[1460475225] HOST DOWNTIME ALERT: SCCM-PRIM01;STOPPED; Host has exited from a period of scheduled downtime
[1460477194] SERVICE FLAPPING ALERT: marshallpk;Paper Tray Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1460477589] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1460477619] HOST ALERT: hallefirst1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460477649] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1460477677] HOST ALERT: hallefirst1;UP;SOFT;2;PING WARNING - Packet loss = 82%, RTA = 0.55 ms
[1460477703] SERVICE ALERT: hallefirst1;Paper Tray Status;OK;SOFT;3;Printer Paper Is OK! All Trays Still Full
[1460477997] SERVICE ALERT: hallefirst2;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460478057] SERVICE ALERT: hallefirst2;Printer Status;OK;SOFT;2;Clearing paper path
[1460478898] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1460478903] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1460479018] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 83% (5 min average)
[1460482069] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1460482099] HOST ALERT: markjeffersonpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460482129] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1460482151] HOST ALERT: markjeffersonpk;UP;SOFT;2;PING OK - Packet loss = 78%, RTA = 5.15 ms
[1460482183] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;OK;SOFT;3;Black Print Cartridge HP C8543X is at 24% - OK!
[1460483103] HOST ALERT: prayharrold205lab;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.56 ms
[1460483253] SERVICE ALERT: prayharrold205lab;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1460483363] SERVICE ALERT: prayharrold205lab;Printer Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460484063] SERVICE ALERT: prayharroldpk;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1460484081] HOST ALERT: prayharroldpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460484093] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1460484113] SERVICE ALERT: prayharroldpk;Printer Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1460484123] HOST ALERT: prayharroldpk;UP;SOFT;2;PING WARNING - Packet loss = 93%, RTA = 0.55 ms
[1460484123] SERVICE ALERT: prayharroldpk;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1460484147] SERVICE ALERT: prayharroldpk;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! 2 of 3 Trays Still Full
[1460484167] SERVICE ALERT: prayharroldpk;Printer Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460484177] SERVICE ALERT: prayharroldpk;Maintenance Kit HP 110V-C9152A;OK;SOFT;3;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 49% - OK!
[1460484227] SERVICE ALERT: prayharroldpk;Printer Status;OK;SOFT;3;50 72 65 73 73 20 17 20 74 6F 20 73 65 65 20 73  74 65 70  by step information
[1460484380] SERVICE ALERT: alexanderpk;Black Cartridge HP CC364X;WARNING;SOFT;1;Black Cartridge HP CC364X is at 2% - WARNING!
[1460484440] SERVICE ALERT: alexanderpk;Black Cartridge HP CC364X;WARNING;SOFT;2;Black Cartridge HP CC364X is at 2% - WARNING!
[1460484501] SERVICE ALERT: alexanderpk;Black Cartridge HP CC364X;WARNING;HARD;3;Black Cartridge HP CC364X is at 2% - WARNING!
[1460484656] SERVICE ALERT: hallefirst2;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460484716] SERVICE ALERT: hallefirst2;Printer Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460484777] SERVICE ALERT: hallefirst2;Printer Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460484843] SERVICE ALERT: sillpk;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460484903] SERVICE ALERT: sillpk;Printer Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460484963] SERVICE ALERT: sillpk;Printer Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460485564] SERVICE ALERT: sillpk;Printer Status;OK;HARD;3;Open tray indicated.
[1460486338] HOST ALERT: idmrbstest.emich.edu;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.102)
[1460486401] HOST ALERT: idmrbstest.emich.edu;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.102)
[1460486464] HOST ALERT: idmrbstest.emich.edu;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.102.102)
[1460486527] HOST ALERT: idmrbstest.emich.edu;DOWN;SOFT;4;CRITICAL - Host Unreachable (164.76.102.102)
[1460486563] SERVICE ALERT: idmrbstest.emich.edu;Banner Driver;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 60 seconds.
[1460486566] HOST ALERT: idmrbstest.emich.edu;DOWN;SOFT;5;CRITICAL - Host Unreachable (164.76.102.102)
[1460486628] HOST ALERT: idmrbstest.emich.edu;DOWN;SOFT;6;CRITICAL - Host Unreachable (164.76.102.102)
[1460486663] SERVICE ALERT: idmrbstest.emich.edu;PING;CRITICAL;SOFT;1;CRITICAL - Host Unreachable (164.76.102.102)
[1460486666] HOST ALERT: idmrbstest.emich.edu;DOWN;SOFT;7;CRITICAL - Host Unreachable (164.76.102.102)
[1460486683] SERVICE ALERT: idmrbstest.emich.edu;Banner Driver;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 60 seconds.
[1460486686] HOST ALERT: idmrbstest.emich.edu;DOWN;SOFT;8;CRITICAL - Host Unreachable (164.76.102.102)
[1460486696] SERVICE ALERT: idmrbstest.emich.edu;eDirectory Service;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1460486708] SERVICE ALERT: idmrbs.emich.edu;Banner Driver;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 60 seconds.
[1460486743] SERVICE ALERT: idmrbstest.emich.edu;Banner Driver;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1460486746] HOST ALERT: idmrbstest.emich.edu;DOWN;SOFT;9;CRITICAL - Host Unreachable (164.76.102.102)
[1460486769] SERVICE ALERT: sillpk;Printer Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.126.102. Is the printer off?  Is the network down?
[1460486783] SERVICE ALERT: sillpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.126.102. Is the printer off?  Is the network down?
[1460486783] SERVICE ALERT: idmrbstest.emich.edu;PING;CRITICAL;SOFT;2;CRITICAL - Host Unreachable (164.76.102.102)
[1460486786] HOST ALERT: idmrbstest.emich.edu;DOWN;HARD;10;CRITICAL - Host Unreachable (164.76.102.102)
[1460486799] HOST ALERT: sillpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460486816] SERVICE ALERT: idmrbstest.emich.edu;eDirectory Service;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1460486823] SERVICE ALERT: sillpk;Printer Status;OK;SOFT;2;Open tray indicated.
[1460486827] HOST ALERT: sillpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.41 ms
[1460486828] SERVICE ALERT: idmrbs.emich.edu;Banner Driver;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 60 seconds.
[1460486836] SERVICE ALERT: sillpk;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! All Trays Still Full
[1460486903] SERVICE ALERT: idmrbstest.emich.edu;PING;CRITICAL;HARD;3;CRITICAL - Host Unreachable (164.76.102.102)
[1460486937] SERVICE ALERT: idmrbstest.emich.edu;eDirectory Service;OK;SOFT;3;OK: ndsd status is good.
[1460486948] SERVICE ALERT: idmrbs.emich.edu;Banner Driver;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 60 seconds.
[1460486973] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1460487003] HOST ALERT: hallefirst2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460487028] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! All Trays Still Full
[1460487032] HOST ALERT: hallefirst2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1460487060] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 1 Trays Empty
[1460487097] HOST DOWNTIME ALERT: idm1.emich.edu;STARTED; Host has entered a period of scheduled downtime
[1460487097] HOST DOWNTIME ALERT: idm2.emich.edu;STARTED; Host has entered a period of scheduled downtime
[1460487097] HOST DOWNTIME ALERT: idm3.emich.edu;STARTED; Host has entered a period of scheduled downtime
[1460487097] HOST DOWNTIME ALERT: idmrbs.emich.edu;STARTED; Host has entered a period of scheduled downtime
[1460487097] HOST DOWNTIME ALERT: idmrbstest.emich.edu;STARTED; Host has entered a period of scheduled downtime
[1460487120] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 1 Trays Empty
[1460487129] SERVICE ALERT: idmrbstest.emich.edu;LDAP Stats;CRITICAL;SOFT;1;CRITICAL: Failed to connect to LDAP!
[1460487176] SERVICE ALERT: hallefirst2;Printer Status;OK;HARD;3;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1460487181] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;SOFT;3;Printer Paper Is OK! All Trays Still Full
[1460487249] SERVICE ALERT: idmrbstest.emich.edu;LDAP Stats;CRITICAL;SOFT;2;CRITICAL: Failed to connect to LDAP!
[1460487369] SERVICE ALERT: idmrbstest.emich.edu;LDAP Stats;OK;SOFT;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1460487816] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1460487846] HOST ALERT: hallefirstcolor;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460487876] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1460487883] SERVICE ALERT: hallefirstcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Image Fuser Kit HP 110V-CE246A, 220V-CE247A'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1460487888] HOST ALERT: hallefirstcolor;UP;SOFT;2;PING OK - Packet loss = 61%, RTA = 0.44 ms
[1460487930] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;OK;SOFT;3;Black Cartridge HP CE260X is at 69% - OK!
[1460487944] SERVICE ALERT: hallefirstcolor;Image Fuser Kit HP 110V-CE246A;OK;SOFT;2;Image Fuser Kit HP 110V-CE246A, 220V-CE247A is at 100% - OK!
[1460487964] SERVICE ALERT: studentcenterpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1460487968] HOST ALERT: studentcenterpk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.59 ms
[1460487969] SERVICE ALERT: idmrbstest.emich.edu;LDAP Stats;CRITICAL;SOFT;1;CRITICAL: Failed to connect to LDAP!
[1460487985] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 39% - OK!
[1460488018] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 94% (5 min average)
[1460488023] SERVICE ALERT: sillpk;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460488083] SERVICE ALERT: sillpk;Printer Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460488089] SERVICE ALERT: idmrbstest.emich.edu;LDAP Stats;CRITICAL;SOFT;2;CRITICAL: Failed to connect to LDAP!
[1460488138] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 94% (5 min average)
[1460488143] SERVICE ALERT: sillpk;Printer Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460488209] SERVICE ALERT: idmrbstest.emich.edu;LDAP Stats;OK;SOFT;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1460488258] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 81% (5 min average)
[1460488263] SERVICE ALERT: studentcenterpk;Printer Status;OK;HARD;3;Processing job from tray 4
[1460488376] SERVICE ALERT: hallefirst2;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460488393] SERVICE ALERT: studentcenterpk;Maintenance Kit HP 110V-C9152A;OK;HARD;3;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 88% - OK!
[1460488436] SERVICE ALERT: hallefirst2;Printer Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460488497] SERVICE ALERT: hallefirst2;Printer Status;OK;SOFT;3;Clearing paper path
[1460488743] SERVICE ALERT: sillpk;Printer Status;OK;HARD;3;Open tray indicated.
[1460488743] SERVICE FLAPPING ALERT: sillpk;Printer Status;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1460488809] SERVICE ALERT: idmrbstest.emich.edu;LDAP Stats;CRITICAL;SOFT;1;CRITICAL: Failed to connect to LDAP!
[1460488929] SERVICE ALERT: idmrbstest.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1460489580] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 1 Trays Empty
[1460489640] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 1 Trays Empty
[1460489701] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 1 Trays Empty
[1460489943] SERVICE ALERT: sillpk;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460490004] SERVICE ALERT: sillpk;Printer Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460490063] SERVICE ALERT: sillpk;Printer Status;OK;SOFT;3;Open tray indicated.
[1460490300] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1460490344] SERVICE ALERT: ADWSUSCLIENTS;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.4.77 and port 12489: Connection refused
[1460490346] SERVICE ALERT: ADWSUSCLIENTS;Uptime;CRITICAL;SOFT;1;connect to address 164.76.4.77 and port 12489: Connection refused
[1460490455] SERVICE ALERT: ADWSUSCLIENTS;C:\ Drive Space;OK;SOFT;2;c:\ - total: 79.66 Gb - used: 38.54 Gb (48%) - free 41.11 Gb (52%)
[1460490457] SERVICE ALERT: ADWSUSCLIENTS;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 2 minute(s)
[1460490530] SERVICE ALERT: ADWSUSSERVERS;Service: User Profile;CRITICAL;SOFT;1;connect to address 164.76.4.79 and port 12489: Connection refused
[1460490641] SERVICE ALERT: ADWSUSSERVERS;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1460491099] SERVICE ALERT: alexanderpk;Black Cartridge HP CC364X;CRITICAL;HARD;3;Black Cartridge HP CC364X is at 1% - CRITICAL!
[1460491263] SERVICE ALERT: sillpk;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460491323] SERVICE ALERT: sillpk;Printer Status;OK;SOFT;2;Then close tray.
[1460491482] HOST ALERT: idmrbs.emich.edu;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460491507] SERVICE ALERT: idmrbs.emich.edu;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460491537] HOST ALERT: idmrbs.emich.edu;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460491627] SERVICE ALERT: idmrbs.emich.edu;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460491627] HOST ALERT: idmrbs.emich.edu;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460491717] HOST ALERT: idmrbs.emich.edu;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460491747] SERVICE ALERT: idmrbs.emich.edu;PING;CRITICAL;HARD;3;PING CRITICAL - Packet loss = 100%
[1460491773] SERVICE ALERT: idmrbs.emich.edu;LDAP Stats;UNKNOWN;SOFT;1;Difference is negative. LDAP recently rebooted.
[1460491777] HOST ALERT: idmrbs.emich.edu;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460491867] HOST ALERT: idmrbs.emich.edu;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1460491893] SERVICE ALERT: idmrbs.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1460491923] HOST ALERT: idmrbs.emich.edu;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1460492013] HOST ALERT: idmrbs.emich.edu;DOWN;SOFT;8;(Host check timed out after 30.01 seconds)
[1460492103] HOST ALERT: idmrbs.emich.edu;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1460492193] HOST ALERT: idmrbs.emich.edu;DOWN;HARD;10;(Host check timed out after 30.01 seconds)
[1460493693] SERVICE ALERT: idmrbs.emich.edu;LDAP Stats;UNKNOWN;SOFT;1;Difference is negative. LDAP recently rebooted.
[1460493813] SERVICE ALERT: idmrbs.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1460493852] SERVICE ALERT: DB01-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460493900] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 1 Trays Empty
[1460493960] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 1 Trays Empty
[1460493972] SERVICE ALERT: DB01-2012;CPU Load;OK;SOFT;2;CPU Load 16% (5 min average)
[1460494021] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 1 Trays Empty
[1460494103] SERVICE ALERT: hallefirst1;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460494163] SERVICE ALERT: hallefirst1;Printer Status;OK;SOFT;2;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1460494253] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1460494314] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;CRITICAL;SOFT;2;Black Print Cartridge HP C8543X is at 1% - CRITICAL!
[1460494373] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;CRITICAL;HARD;3;Black Print Cartridge HP C8543X is at 1% - CRITICAL!
[1460494620] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1460494620] SERVICE FLAPPING ALERT: hallefirstcolor;Paper Tray Status;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1460495963] SERVICE ALERT: prayharrold205lab;Printer Status;OK;HARD;3;Clearing paper path
[1460496173] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 55% - OK!
[1460496843] SERVICE ALERT: hallesecondsouth;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460496903] SERVICE ALERT: hallesecondsouth;Printer Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460496913] HOST ALERT: prayharrold205lab;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460496963] SERVICE ALERT: hallesecondsouth;Printer Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460497003] HOST ALERT: prayharrold205lab;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460497057] SERVICE ALERT: prayharrold205lab;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.205.183. Is the printer off?  Is the network down?
[1460497087] HOST ALERT: prayharrold205lab;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460497117] SERVICE ALERT: prayharrold205lab;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.205.183. Is the printer off?  Is the network down?
[1460497147] HOST ALERT: prayharrold205lab;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460497169] SERVICE ALERT: prayharrold205lab;Printer Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.205.183. Is the printer off?  Is the network down?
[1460497177] SERVICE ALERT: prayharrold205lab;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.205.183. Is the printer off?  Is the network down?
[1460497199] HOST ALERT: prayharrold205lab;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460497229] SERVICE ALERT: prayharrold205lab;Printer Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.205.183. Is the printer off?  Is the network down?
[1460497259] HOST ALERT: prayharrold205lab;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1460497289] SERVICE ALERT: prayharrold205lab;Printer Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.205.183. Is the printer off?  Is the network down?
[1460497319] HOST ALERT: prayharrold205lab;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1460497409] HOST ALERT: prayharrold205lab;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1460497499] HOST ALERT: prayharrold205lab;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1460497503] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460497563] SERVICE ALERT: hallesecondsouth;Printer Status;OK;HARD;3;Ready
[1460497589] HOST ALERT: prayharrold205lab;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1460497623] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1460497743] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1460498458] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1460498578] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 84% (5 min average)
[1460498900] SERVICE ALERT: alexanderpk;Black Cartridge HP CC364X;OK;HARD;3;Black Cartridge HP CC364X is at 100% - OK!
[1460499645] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1460499675] HOST ALERT: hallefirstcolor;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460499705] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1460499720] HOST ALERT: hallefirstcolor;UP;SOFT;2;PING OK - Packet loss = 68%, RTA = 0.38 ms
[1460499759] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;OK;SOFT;3;Yellow Cartridge HP CE262A is at 31% - OK!
[1460500323] SERVICE FLAPPING ALERT: sillpk;Printer Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1460501149] SERVICE ALERT: porterpk;Printer Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.133.4. Is the printer off?  Is the network down?
[1460501209] SERVICE ALERT: porterpk;Printer Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.133.4. Is the printer off?  Is the network down?
[1460501219] SERVICE ALERT: porterpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.133.4. Is the printer off?  Is the network down?
[1460501239] HOST ALERT: porterpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460501263] SERVICE ALERT: porterpk;Printer Status;OK;SOFT;3;Ready
[1460501267] HOST ALERT: porterpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1460501273] SERVICE ALERT: porterpk;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! All Trays Still Full
[1460501696] SERVICE ALERT: hallefirst2;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460501744] SERVICE ALERT: porterpk;Black Cartridge HP CC364X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.133.4. Is the printer off?  Is the network down?
[1460501756] SERVICE ALERT: hallefirst2;Printer Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460501804] SERVICE ALERT: porterpk;Black Cartridge HP CC364X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.133.4. Is the printer off?  Is the network down?
[1460501817] SERVICE ALERT: hallefirst2;Printer Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460501864] SERVICE ALERT: porterpk;Black Cartridge HP CC364X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.133.4. Is the printer off?  Is the network down?
[1460501869] SERVICE ALERT: porterpk;Printer Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.133.4. Is the printer off?  Is the network down?
[1460501879] SERVICE ALERT: porterpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.133.4. Is the printer off?  Is the network down?
[1460501929] SERVICE ALERT: porterpk;Printer Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.133.4. Is the printer off?  Is the network down?
[1460501939] SERVICE ALERT: porterpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.133.4. Is the printer off?  Is the network down?
[1460501989] SERVICE ALERT: porterpk;Printer Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.133.4. Is the printer off?  Is the network down?
[1460501999] SERVICE ALERT: porterpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.133.4. Is the printer off?  Is the network down?
[1460502416] SERVICE ALERT: hallefirst2;Printer Status;OK;HARD;3;Ready
[1460506021] SERVICE FLAPPING ALERT: hallefirstcolor;Paper Tray Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1460510050] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1460510099] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1460510110] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1460510159] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1460510170] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1460510172] SERVICE ALERT: DB01-2012;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1460510219] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1460510292] SERVICE ALERT: DB01-2012;CPU Load;CRITICAL;SOFT;2;CPU Load 96% (5 min average)
[1460510412] SERVICE ALERT: DB01-2012;CPU Load;OK;SOFT;3;CPU Load 73% (5 min average)
[1460510433] SERVICE ALERT: halleground2;Printer Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1460510456] SERVICE ALERT: TARGETX-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1460510463] SERVICE ALERT: AXIOM-TEST-2012;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1460510463] HOST ALERT: halleground2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460510493] SERVICE ALERT: halleground2;Printer Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1460510497] HOST ALERT: halleground2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.53 ms
[1460510499] SERVICE ALERT: halleground2;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1460510553] SERVICE ALERT: halleground2;Printer Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1460510559] SERVICE ALERT: halleground2;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1460510576] SERVICE ALERT: TARGETX-2012;CPU Load;OK;SOFT;2;CPU Load 28% (5 min average)
[1460510583] SERVICE ALERT: AXIOM-TEST-2012;CPU Load;OK;SOFT;2;CPU Load 86% (5 min average)
[1460510619] SERVICE ALERT: halleground2;Maintenance Kit HP 110V-C9152A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1460510649] HOST ALERT: halleground2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460510685] SERVICE ALERT: eateriespk;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460510713] HOST ALERT: halleground2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1460510713] HOST FLAPPING ALERT: halleground2;STARTED; Host appears to have started flapping (22.8% change > 20.0% threshold)
[1460510745] SERVICE ALERT: eateriespk;Printer Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460510805] SERVICE ALERT: eateriespk;Printer Status;OK;SOFT;3;TRAY 2 TYPE=  ANY TYPE *PLAIN  LIGHT 60-75 G/M2
[1460511415] SERVICE ALERT: hallefirst2;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460511476] SERVICE ALERT: hallefirst2;Printer Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460511535] SERVICE ALERT: hallefirst2;Printer Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460511970] HOST FLAPPING ALERT: halleground2;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1460512653] SERVICE ALERT: halleground1;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460512713] SERVICE ALERT: halleground1;Printer Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460512774] SERVICE ALERT: halleground1;Printer Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460513219] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 1 Trays Empty
[1460513280] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 1 Trays Empty
[1460513339] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 1 Trays Empty
[1460513373] SERVICE ALERT: halleground1;Printer Status;OK;HARD;3;ORDER CARTRIDGE LESS THAN 1500 PAGES 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Processing job from tray 4
[1460515140] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1460517176] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1460517296] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 83% (5 min average)
[1460517543] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 87% (5 min average)
[1460518143] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460518263] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460518383] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460518443] SERVICE ALERT: hallefirst2;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Maintenance Kit HP 110V-C9152A, 220V-C9153A'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1460518503] SERVICE ALERT: hallefirst2;Maintenance Kit HP 110V-C9152A;OK;SOFT;2;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 89% - OK!
[1460518734] SERVICE ALERT: hallefirst2;Printer Status;OK;HARD;3;Ready