[1453878282] SERVICE ALERT: eateriespk;Paper Tray Status;UNKNOWN;SOFT;1;Printer needs paper!
[1453878342] SERVICE ALERT: eateriespk;Paper Tray Status;OK;SOFT;2;Printer Paper is OK! -  Total Trays: 3, Trays Needing Paper: 2
[1453883604] HOST ALERT: owen101lab1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1453883694] HOST ALERT: owen101lab1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1453883784] HOST ALERT: owen101lab1;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1453883841] HOST ALERT: owen101lab1;DOWN;SOFT;4;(Host check timed out after 30.00 seconds)
[1453883928] SERVICE ALERT: halleground2;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Maintenance Kit HP 110V-C9152A, 220V-C9153A'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1453883931] HOST ALERT: owen101lab1;DOWN;SOFT;5;(Host check timed out after 30.00 seconds)
[1453883988] SERVICE ALERT: halleground2;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Maintenance Kit HP 110V-C9152A, 220V-C9153A'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1453884021] HOST ALERT: owen101lab1;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1453884048] SERVICE ALERT: halleground2;Maintenance Kit HP 110V-C9152A;OK;SOFT;3;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 61% - OK!
[1453884085] HOST ALERT: owen101lab1;UP;SOFT;7;PING OK - Packet loss = 0%, RTA = 0.69 ms
[1453884204] HOST ALERT: owen101lab1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1453884294] HOST ALERT: owen101lab1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1453884384] HOST ALERT: owen101lab1;UP;SOFT;3;PING WARNING - Packet loss = 93%, RTA = 0.62 ms
[1453884384] HOST FLAPPING ALERT: owen101lab1;STARTED; Host appears to have started flapping (22.5% change > 20.0% threshold)
[1453884441] HOST ALERT: owen101lab1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1453884469] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1453884531] HOST ALERT: owen101lab1;DOWN;SOFT;2;(Host check timed out after 30.00 seconds)
[1453884621] HOST ALERT: owen101lab1;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1453884711] HOST ALERT: owen101lab1;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1453884801] HOST ALERT: owen101lab1;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1453884891] HOST ALERT: owen101lab1;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1453884981] HOST ALERT: owen101lab1;DOWN;SOFT;7;(Host check timed out after 30.00 seconds)
[1453885041] HOST ALERT: owen101lab1;DOWN;SOFT;8;(Host check timed out after 30.00 seconds)
[1453885131] HOST ALERT: owen101lab1;DOWN;SOFT;9;(Host check timed out after 30.00 seconds)
[1453885221] HOST ALERT: owen101lab1;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1453885971] HOST FLAPPING ALERT: owen101lab1;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1453899653] SERVICE ALERT: strongpk;Paper Tray Status;WARNING;HARD;3;Printer needs paper!
[1453901452] SERVICE ALERT: strongpk;Paper Tray Status;OK;HARD;3;Printer Paper is OK! -  Total Trays: 2, Trays Needing Paper: 0
[1453901822] SERVICE ALERT: marshallpk;Black Cartridge HP CC364X;OK;HARD;3;Black Cartridge HP CC364X is at 100% - OK!
[1453901967] HOST ALERT: owen101lab3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1453902031] HOST ALERT: owen101lab3;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.63 ms
[1453902144] HOST ALERT: owen101lab1;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.59 ms
[1453902168] SERVICE ALERT: owen101lab1;Maintenance Kit HP 110V-C9152A;OK;HARD;3;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 74% - OK!
[1453902399] SERVICE ALERT: owen101lab3;Paper Tray Status;OK;HARD;3;Printer Paper is OK! -  Total Trays: 2, Trays Needing Paper: 0
[1453902404] SERVICE ALERT: owen101lab1;Paper Tray Status;OK;HARD;3;Printer Paper is OK! -  Total Trays: 2, Trays Needing Paper: 0
[1453902406] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 42% - OK!
[1453902531] SERVICE ALERT: owen101lab3;Maintenance Kit HP 110V-C9152A;OK;HARD;3;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 72% - OK!
[1453903418] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;SOFT;1;Printer needs paper!
[1453903477] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;SOFT;2;Printer needs paper!
[1453903538] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;HARD;3;Printer needs paper!
[1453910087] SERVICE ALERT: crossroadspk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1453916566] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;CRITICAL;HARD;3;Black Print Cartridge HP C8543X is at 1% - CRITICAL!
[1453922803] SERVICE ALERT: owen101lab1;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1453922864] SERVICE ALERT: owen101lab1;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1453922924] SERVICE ALERT: owen101lab1;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1453923794] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;WARNING;SOFT;1;Magenta Cartridge HP CE263A is at 4% - WARNING!
[1453923855] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;WARNING;SOFT;2;Magenta Cartridge HP CE263A is at 4% - WARNING!
[1453923914] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;WARNING;HARD;3;Magenta Cartridge HP CE263A is at 4% - WARNING!
[1453925487] SERVICE ALERT: hallefirst1;Maintenance Kit HP 110V-C9152A;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.104.232! Make sure host is up and SNMP is configured properly.
[1453925545] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.104.232! Make sure host is up and SNMP is configured properly.
[1453925547] SERVICE ALERT: hallefirst1;Maintenance Kit HP 110V-C9152A;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.104.232! Make sure host is up and SNMP is configured properly.
[1453925605] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.104.232! Make sure host is up and SNMP is configured properly.
[1453925607] SERVICE ALERT: hallefirst1;Maintenance Kit HP 110V-C9152A;WARNING;HARD;3;WARNING: No SNMP response from 164.76.104.232! Make sure host is up and SNMP is configured properly.
[1453925665] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;HARD;3;WARNING: No SNMP response from 164.76.104.232! Make sure host is up and SNMP is configured properly.
[1453925695] HOST ALERT: hallefirst1;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1453925749] SERVICE ALERT: hallefirst1;Paper Tray Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.104.232! Make sure host is up and SNMP is configured properly.
[1453925755] HOST ALERT: hallefirst1;UP;SOFT;2;PING OK - Packet loss = 28%, RTA = 1.21 ms
[1453925803] SERVICE ALERT: hallefirst1;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! All Trays Still Full
[1453925809] SERVICE ALERT: hallefirst2;Paper Tray Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1453925869] SERVICE ALERT: hallefirst2;Paper Tray Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1453925878] SERVICE ALERT: hallefirst2;Maintenance Kit HP 110V-C9152A;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1453925882] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1453925923] SERVICE ALERT: owen101lab1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1453925930] SERVICE ALERT: hallefirst2;Paper Tray Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1453925938] SERVICE ALERT: hallefirst2;Maintenance Kit HP 110V-C9152A;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1453925942] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1453925998] SERVICE ALERT: hallefirst2;Maintenance Kit HP 110V-C9152A;WARNING;HARD;3;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1453926002] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;WARNING;HARD;3;WARNING: No SNMP response from 164.76.104.163! Make sure host is up and SNMP is configured properly.
[1453926202] SERVICE ALERT: hallefirst1;Maintenance Kit HP 110V-C9152A;OK;HARD;3;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 92% - OK!
[1453926260] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 43% - OK!
[1453926523] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1453926592] SERVICE ALERT: hallefirst2;Maintenance Kit HP 110V-C9152A;OK;HARD;3;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 12% - OK!
[1453926596] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 71% - OK!
[1453931928] SERVICE ALERT: owen101lab3;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Maintenance Kit HP 110V-C9152A, 220V-C9153A'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1453931988] SERVICE ALERT: owen101lab3;Maintenance Kit HP 110V-C9152A;OK;SOFT;2;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 71% - OK!
[1453933723] SERVICE ALERT: owen101lab1;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1453933782] SERVICE ALERT: owen101lab1;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1453933843] SERVICE ALERT: owen101lab1;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1453934565] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1453936243] SERVICE ALERT: owen101lab1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1453939841] SERVICE ALERT: owen101lab1;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1453939902] SERVICE ALERT: owen101lab1;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1453939962] SERVICE ALERT: owen101lab1;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1453944914] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;CRITICAL;HARD;3;Magenta Cartridge HP CE263A is at CRITICAL level!
[1453946560] SERVICE ALERT: owen101lab1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1453953957] SERVICE ALERT: MBAM;CPU Load;CRITICAL;SOFT;1;CPU Load 90% (5 min average)
[1453954078] SERVICE ALERT: MBAM;CPU Load;CRITICAL;SOFT;2;CPU Load 95% (5 min average)
[1453954197] SERVICE ALERT: MBAM;CPU Load;OK;SOFT;3;CPU Load 14% (5 min average)
[1453954503] SERVICE ALERT: LANSWEEPER;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1453954504] SERVICE ALERT: LANSWEEPER;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1453954533] HOST ALERT: LANSWEEPER;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1453954549] SERVICE ALERT: LANSWEEPER;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1453954555] SERVICE ALERT: LANSWEEPER;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1453954579] HOST ALERT: LANSWEEPER;DOWN;SOFT;2;(Host check timed out after 30.00 seconds)
[1453954612] SERVICE ALERT: LANSWEEPER;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1453954623] SERVICE ALERT: LANSWEEPER;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1453954624] SERVICE ALERT: LANSWEEPER;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1453954630] SERVICE ALERT: LANSWEEPER;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1453954642] HOST ALERT: LANSWEEPER;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1453954669] SERVICE ALERT: LANSWEEPER;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1453954675] SERVICE ALERT: LANSWEEPER;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1453954689] SERVICE ALERT: LANSWEEPER;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1453954699] HOST ALERT: LANSWEEPER;DOWN;SOFT;4;(Host check timed out after 30.00 seconds)
[1453954732] SERVICE ALERT: LANSWEEPER;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1453954743] SERVICE ALERT: LANSWEEPER;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1453954744] SERVICE ALERT: LANSWEEPER;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1453954748] SERVICE ALERT: LANSWEEPER;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1453954750] SERVICE ALERT: LANSWEEPER;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1453954762] HOST ALERT: LANSWEEPER;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1453954788] SERVICE ALERT: LANSWEEPER;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1453954795] SERVICE ALERT: LANSWEEPER;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1453954808] SERVICE ALERT: LANSWEEPER;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1453954818] HOST ALERT: LANSWEEPER;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1453954831] SERVICE ALERT: LANSWEEPER;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1453954852] SERVICE ALERT: LANSWEEPER;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1453954861] HOST ALERT: LANSWEEPER;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1453954868] SERVICE ALERT: LANSWEEPER;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1453954870] SERVICE ALERT: LANSWEEPER;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1453954884] SERVICE ALERT: LANSWEEPER;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1453954898] HOST ALERT: LANSWEEPER;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1453954928] SERVICE ALERT: LANSWEEPER;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1453954939] SERVICE ALERT: LANSWEEPER;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1453954951] SERVICE ALERT: LANSWEEPER;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1453954954] SERVICE ALERT: LANSWEEPER;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1453954958] HOST ALERT: LANSWEEPER;DOWN;SOFT;9;(Host check timed out after 30.00 seconds)
[1453954988] SERVICE ALERT: LANSWEEPER;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1453955004] SERVICE ALERT: LANSWEEPER;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1453955018] HOST ALERT: LANSWEEPER;DOWN;HARD;10;(Host check timed out after 30.00 seconds)
[1453955043] SERVICE ALERT: LANSWEEPER;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1453955059] SERVICE ALERT: LANSWEEPER;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1453955071] SERVICE ALERT: LANSWEEPER;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1453955074] SERVICE ALERT: LANSWEEPER;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1453955124] SERVICE ALERT: LANSWEEPER;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1453955163] SERVICE ALERT: LANSWEEPER;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1453955179] SERVICE ALERT: LANSWEEPER;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1453955194] SERVICE ALERT: LANSWEEPER;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1453955283] SERVICE ALERT: LANSWEEPER;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds