[1463803225] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1463803284] SERVICE ALERT: GOVGIS;CPU Load;OK;SOFT;2;CPU Load 0% (15 min average)
[1463803345] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463803766] SERVICE ALERT: TARGETX-2012;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463803866] SERVICE ALERT: TARGETX-2012;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1463803945] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 88% (15 min average)
[1463804353] SERVICE ALERT: idm2.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1463804401] SERVICE ALERT: idm3.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1463804545] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463804626] SERVICE ALERT: idm2.emich.edu;LDAP Stats;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1463804665] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 90% (15 min average)
[1463804746] SERVICE ALERT: idm2.emich.edu;LDAP Stats;CRITICAL;SOFT;2;CRITICAL: Failed to connect to LDAP!
[1463804778] SERVICE FLAPPING ALERT: C4SAS-SE;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1463804866] SERVICE ALERT: idm2.emich.edu;LDAP Stats;OK;SOFT;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463804953] SERVICE ALERT: idm2.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 32
[1463805001] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 29
[1463805073] SERVICE ALERT: idm2.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 31
[1463805121] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 29
[1463805193] SERVICE ALERT: idm2.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 31
[1463805241] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 23
[1463805265] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1463805385] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 91% (15 min average)
[1463805770] SERVICE FLAPPING ALERT: FITNESSGRAM;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1463805985] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1463806029] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;1;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1463806105] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 48% (15 min average)
[1463806149] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;2;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1463806268] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;HARD;3;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1463806705] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463806825] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1463806942] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 140 found since last run. wholeSubtreeSearchOps is too high! 510 found since last run. oneLevelSearchOps OK. searchOps is too high! 512 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 930 found since last run.
[1463806945] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463807062] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463807662] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 648 found since last run. oneLevelSearchOps OK. searchOps is too high! 662 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 901 found since last run.
[1463807783] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463807955] SERVICE ALERT: idm1.emich.edu;AD Driver;CRITICAL;SOFT;1;CRITICAL: cn=Active Directory,cn=driverset1,o=system is stopped!
[1463808074] SERVICE ALERT: idm1.emich.edu;AD Driver;CRITICAL;SOFT;2;CRITICAL: cn=Active Directory,cn=driverset1,o=system is stopped!
[1463808195] SERVICE ALERT: idm1.emich.edu;AD Driver;CRITICAL;HARD;3;CRITICAL: cn=Active Directory,cn=driverset1,o=system is stopped!
[1463808382] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 733 found since last run. oneLevelSearchOps OK. searchOps is too high! 734 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 841 found since last run.
[1463808501] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463808733] SERVICE FLAPPING ALERT: INTLDB;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1463809100] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 102 found since last run. wholeSubtreeSearchOps is too high! 751 found since last run. oneLevelSearchOps OK. searchOps is too high! 811 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 1002 found since last run.
[1463809174] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 219.90 Gb - used: 198.41 Gb (90%) - free 21.49 Gb (10%)
[1463809220] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463809294] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 219.90 Gb - used: 204.51 Gb (93%) - free 15.39 Gb (7%)
[1463809414] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 218.93 Gb (100%) - free 0.97 Gb (0%)
[1463809820] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 715 found since last run. oneLevelSearchOps OK. searchOps is too high! 721 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 827 found since last run.
[1463809940] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463810014] SERVICE ALERT: BOEAPP;C:\ Drive Space;OK;HARD;3;c:\ - total: 219.90 Gb - used: 162.88 Gb (74%) - free 57.02 Gb (26%)
[1463810540] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 125 found since last run. wholeSubtreeSearchOps is too high! 856 found since last run. oneLevelSearchOps OK. searchOps is too high! 898 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 983 found since last run.
[1463810660] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463811260] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 697 found since last run. oneLevelSearchOps OK. searchOps is too high! 702 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 971 found since last run.
[1463811380] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463811656] SERVICE ALERT: idm1.emich.edu;NSDS Errors;OK;HARD;3;OK: 0 Errors Reported
[1463811980] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 755 found since last run. oneLevelSearchOps OK. searchOps is too high! 763 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 950 found since last run.
[1463812100] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463812256] SERVICE ALERT: idm1.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 15
[1463812376] SERVICE ALERT: idm1.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 16
[1463812394] SERVICE ALERT: idm1.emich.edu;AD Driver;OK;HARD;3;OK: cn=Active Directory,cn=driverset1,o=system is running.
[1463812496] SERVICE ALERT: idm1.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 16
[1463812666] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;CRITICAL;SOFT;1;CRITICAL: cn=ADAUTH Driver,cn=driverset1,o=system is stopped!
[1463812700] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 725 found since last run. oneLevelSearchOps OK. searchOps is too high! 731 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 939 found since last run.
[1463812785] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;CRITICAL;SOFT;2;CRITICAL: cn=ADAUTH Driver,cn=driverset1,o=system is stopped!
[1463812820] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463812906] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;CRITICAL;HARD;3;CRITICAL: cn=ADAUTH Driver,cn=driverset1,o=system is stopped!
[1463813040] SERVICE ALERT: idm3.emich.edu;NSDS Errors;OK;HARD;3;OK:
[1463813096] SERVICE ALERT: idm1.emich.edu;NSDS Errors;OK;HARD;3;OK: 0 Errors Reported
[1463813420] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 102 found since last run. wholeSubtreeSearchOps is too high! 702 found since last run. oneLevelSearchOps OK. searchOps is too high! 712 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 994 found since last run.
[1463813468] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;OK;HARD;3;OK: No entries found
[1463813540] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463813545] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 64% (15 min average)
[1463813640] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 7
[1463813653] SERVICE ALERT: idm2.emich.edu;NSDS Errors;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 60 seconds.
[1463813696] SERVICE ALERT: idm1.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 6
[1463813760] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 6
[1463813816] SERVICE ALERT: idm1.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 6
[1463813866] SERVICE ALERT: idm2.emich.edu;LDAP Stats;UNKNOWN;SOFT;1;Difference is negative. LDAP recently rebooted.
[1463813880] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 6
[1463813880] SERVICE FLAPPING ALERT: idm3.emich.edu;NSDS Errors;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1463813936] SERVICE ALERT: idm1.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 6
[1463813986] SERVICE ALERT: idm2.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463814140] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 4275 found since last run. oneLevelSearchOps OK. searchOps is too high! 4284 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 3210 found since last run.
[1463814145] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463814177] SERVICE ALERT: SNAP;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1463814177] SERVICE ALERT: SNAP;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1463814180] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.40)
[1463814193] SERVICE ALERT: idm2.emich.edu;NSDS Errors;OK;HARD;3;OK:
[1463814193] SERVICE FLAPPING ALERT: idm2.emich.edu;NSDS Errors;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1463814202] HOST ALERT: SNAP;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1463814253] SERVICE ALERT: SNAP;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463814265] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1463814270] SERVICE ALERT: idm4.emich.edu;LDAP Stats;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1463814288] SERVICE ALERT: SNAP;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463814294] SERVICE ALERT: SNAP;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1463814294] SERVICE ALERT: SNAP;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1463814353] SERVICE ALERT: SNAP;CPU Load;OK;SOFT;2;CPU Load 0% (15 min average)
[1463814388] SERVICE ALERT: SNAP;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1463814390] SERVICE ALERT: idm4.emich.edu;LDAP Stats;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1463814420] HOST ALERT: idm4.emich.edu;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1463814441] SERVICE ALERT: idm4.emich.edu;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1463814471] HOST ALERT: idm4.emich.edu;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1463814490] SERVICE ALERT: idm4.emich.edu;eDirectory Service;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1463814493] SERVICE ALERT: idm3.emich.edu;LDAP Stats;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1463814518] HOST ALERT: idm3.emich.edu;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1463814521] HOST ALERT: idm4.emich.edu;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1463814540] SERVICE ALERT: idm3.emich.edu;NSDS Errors;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 60 seconds.
[1463814543] HOST ALERT: idm1.emich.edu;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1463814561] SERVICE ALERT: idm4.emich.edu;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1463814570] HOST ALERT: idm3.emich.edu;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1463814577] SERVICE ALERT: idm3.emich.edu;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1463814585] SERVICE ALERT: idm1.emich.edu;Google Driver;CRITICAL;SOFT;1;Unknown: Error reported?
[1463814589] HOST ALERT: idm1.emich.edu;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1463814591] HOST ALERT: idm4.emich.edu;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1463814596] SERVICE ALERT: idm1.emich.edu;NSDS Errors;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 60 seconds.
[1463814607] HOST ALERT: idm3.emich.edu;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1463814609] SERVICE ALERT: idm4.emich.edu;eDirectory Service;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1463814613] SERVICE ALERT: idm3.emich.edu;LDAP Stats;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1463814639] HOST ALERT: idm4.emich.edu;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1463814643] HOST ALERT: idm3.emich.edu;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1463814663] SERVICE ALERT: idm3.emich.edu;eDirectory Service;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1463814681] SERVICE ALERT: idm1.emich.edu;Google Driver;CRITICAL;SOFT;2;Unknown: Error reported?
[1463814681] SERVICE ALERT: idm4.emich.edu;PING;CRITICAL;HARD;3;PING CRITICAL - Packet loss = 100%
[1463814693] HOST ALERT: idm3.emich.edu;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1463814697] SERVICE ALERT: idm3.emich.edu;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1463814711] HOST ALERT: idm4.emich.edu;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1463814727] HOST ALERT: idm3.emich.edu;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1463814729] SERVICE ALERT: idm4.emich.edu;eDirectory Service;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1463814733] SERVICE ALERT: idm3.emich.edu;LDAP Stats;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1463814744] SERVICE ALERT: idm1.emich.edu;OfficeAD Driver;CRITICAL;SOFT;1;CRITICAL: cn=OfficeAD,cn=driverset1,o=system is stopped!
[1463814759] HOST ALERT: idm4.emich.edu;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1463814763] HOST ALERT: idm3.emich.edu;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1463814767] SERVICE ALERT: idm4.emich.edu;NSDS Errors;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 60 seconds.
[1463814783] SERVICE ALERT: idm3.emich.edu;eDirectory Service;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1463814793] SERVICE ALERT: idm2.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 52
[1463814794] SERVICE ALERT: idm1.emich.edu;Loopback Driver;CRITICAL;SOFT;1;CRITICAL: cn=Generic Loopback,cn=driverset1,o=system is stopped!
[1463814795] SERVICE ALERT: idm1.emich.edu;AD Driver;CRITICAL;SOFT;1;CRITICAL: cn=Active Directory,cn=driverset1,o=system is stopped!
[1463814797] HOST ALERT: idm4.emich.edu;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1463814800] SERVICE ALERT: idm1.emich.edu;Google Driver;CRITICAL;HARD;3;CRITICAL: cn=gdap,cn=driverset1,o=system is stopped!
[1463814813] HOST ALERT: idm3.emich.edu;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1463814817] SERVICE ALERT: idm3.emich.edu;PING;CRITICAL;HARD;3;PING CRITICAL - Packet loss = 100%
[1463814840] SERVICE ALERT: idm1.emich.edu;LDAP Stats;UNKNOWN;SOFT;1;Difference is negative. LDAP recently rebooted.
[1463814847] HOST ALERT: idm3.emich.edu;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1463814864] SERVICE ALERT: idm1.emich.edu;OfficeAD Driver;CRITICAL;SOFT;2;CRITICAL: cn=OfficeAD,cn=driverset1,o=system is stopped!
[1463814887] HOST ALERT: idm4.emich.edu;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1463814903] SERVICE ALERT: idm3.emich.edu;eDirectory Service;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1463814909] HOST ALERT: idm3.emich.edu;UP;SOFT;10;PING OK - Packet loss = 28%, RTA = 1.11 ms
[1463814913] SERVICE ALERT: idm2.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 52
[1463814913] SERVICE ALERT: idm1.emich.edu;Loopback Driver;OK;SOFT;2;OK: cn=Generic Loopback,cn=driverset1,o=system is running.
[1463814914] SERVICE ALERT: idm1.emich.edu;AD Driver;OK;SOFT;2;OK: cn=Active Directory,cn=driverset1,o=system is running.
[1463814960] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463814977] HOST ALERT: idm4.emich.edu;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1463814985] SERVICE ALERT: idm1.emich.edu;OfficeAD Driver;OK;SOFT;3;OK: cn=OfficeAD,cn=driverset1,o=system is running.
[1463815033] SERVICE ALERT: idm2.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 32
[1463815136] SERVICE ALERT: idm1.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 15
[1463815275] SERVICE ALERT: idm4.emich.edu;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1463815279] HOST ALERT: idm4.emich.edu;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1463815306] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;OK;HARD;3;OK: cn=ADAUTH Driver,cn=driverset1,o=system is running.
[1463815307] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 19
[1463815320] SERVICE ALERT: idm4.emich.edu;eDirectory Service;OK;HARD;3;OK: ndsd status is good.
[1463815363] HOST ALERT: idm3.emich.edu;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1463815401] SERVICE ALERT: idm1.emich.edu;Google Driver;OK;HARD;3;OK: cn=gdap,cn=driverset1,o=system is running.
[1463815411] SERVICE ALERT: idm3.emich.edu;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1463815415] HOST ALERT: idm3.emich.edu;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1463815493] SERVICE ALERT: idm3.emich.edu;eDirectory Service;OK;HARD;3;OK: ndsd status is good.
[1463815580] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 295 found since last run. oneLevelSearchOps OK. searchOps is too high! 300 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 816 found since last run.
[1463815923] SERVICE ALERT: idm3.emich.edu;LDAP Stats;UNKNOWN;HARD;3;Difference is negative. LDAP recently rebooted.
[1463816165] SERVICE ALERT: SNAP;CPU Load;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1463816168] SERVICE ALERT: SNAP;Uptime;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1463816168] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.40)
[1463816191] SERVICE ALERT: SNAP;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1463816194] HOST ALERT: SNAP;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.40)
[1463816233] SERVICE ALERT: SNAP;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1463816233] SERVICE ALERT: SNAP;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1463816236] HOST ALERT: SNAP;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.102.40)
[1463816261] HOST ALERT: SNAP;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.40 ms
[1463816273] SERVICE ALERT: SNAP;CPU Load;OK;SOFT;2;CPU Load 0% (15 min average)
[1463816276] SERVICE ALERT: SNAP;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 6 hour(s) 29 minute(s)
[1463816280] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 10
[1463816308] SERVICE ALERT: SNAP;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1463816350] SERVICE ALERT: SNAP;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1463816353] SERVICE ALERT: SNAP;C:\ Drive Space;OK;SOFT;2;c:\ - total: 1116.65 Gb - used: 36.15 Gb (3%) - free 1080.50 Gb (97%)
[1463816523] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463817214] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;SOFT;1;c:\ - total: 219.90 Gb - used: 211.47 Gb (96%) - free 8.43 Gb (4%)
[1463817265] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1463817334] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;SOFT;2;c:\ - total: 219.90 Gb - used: 209.05 Gb (95%) - free 10.85 Gb (5%)
[1463817385] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 83% (15 min average)
[1463817454] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 218.91 Gb (100%) - free 0.99 Gb (0%)
[1463818054] SERVICE ALERT: BOEAPP;C:\ Drive Space;OK;HARD;3;c:\ - total: 219.90 Gb - used: 163.23 Gb (74%) - free 56.67 Gb (26%)
[1463818054] SERVICE FLAPPING ALERT: BOEAPP;C:\ Drive Space;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1463818585] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1463818633] SERVICE ALERT: idm2.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1463818681] SERVICE ALERT: idm3.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1463818705] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 86% (15 min average)
[1463818736] SERVICE ALERT: idm1.emich.edu;NSDS Errors;OK;HARD;3;OK: 0 Errors Reported
[1463818906] SERVICE ALERT: idm4.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1463819610] SERVICE ALERT: SNAP;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1463819610] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.40)
[1463819655] HOST ALERT: SNAP;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1463819718] SERVICE ALERT: SNAP;Memory Usage;OK;SOFT;2;Memory usage: total:16362.75 MB - used: 1335.61 MB (8%) - free: 15027.14 MB (92%)
[1463819833] SERVICE ALERT: idm2.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 1
[1463819880] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 2
[1463819905] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1463819936] SERVICE ALERT: idm1.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 2
[1463819953] SERVICE ALERT: idm2.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 2
[1463820000] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 2
[1463820056] SERVICE ALERT: idm1.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 2
[1463820073] SERVICE ALERT: idm2.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 2
[1463820107] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 2
[1463820120] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 2
[1463820123] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 241 found since last run. oneLevelSearchOps OK. searchOps is too high! 246 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463820177] SERVICE ALERT: idm1.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 2
[1463820227] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 2
[1463820243] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463820348] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 3
[1463820348] SERVICE FLAPPING ALERT: idm4.emich.edu;NSDS Errors;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1463821105] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1463821225] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1463821315] SERVICE ALERT: EVENTSTEST-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1463821345] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 71% (15 min average)
[1463821435] SERVICE ALERT: EVENTSTEST-2012;CPU Load;OK;SOFT;2;CPU Load 90% (15 min average)
[1463821980] SERVICE ALERT: idm3.emich.edu;NSDS Errors;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 60 seconds.
[1463822043] SERVICE ALERT: idm3.emich.edu;LDAP Stats;CRITICAL;SOFT;1;CRITICAL: Failed to connect to LDAP!
[1463822163] SERVICE ALERT: idm3.emich.edu;LDAP Stats;CRITICAL;SOFT;2;CRITICAL: Failed to connect to LDAP!
[1463822283] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463822520] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 4
[1463823777] SERVICE ALERT: idm1.emich.edu;NSDS Errors;OK;HARD;3;OK: 0 Errors Reported
[1463823947] SERVICE ALERT: idm4.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1463824273] SERVICE ALERT: idm2.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1463824320] SERVICE ALERT: idm3.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1463824345] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463824465] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1463824585] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1463826454] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;SOFT;1;c:\ - total: 219.90 Gb - used: 216.56 Gb (98%) - free 3.34 Gb (2%)
[1463826574] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;SOFT;2;c:\ - total: 219.90 Gb - used: 214.70 Gb (98%) - free 5.20 Gb (2%)
[1463826693] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 218.56 Gb (99%) - free 1.33 Gb (1%)
[1463827293] SERVICE ALERT: BOEAPP;C:\ Drive Space;OK;HARD;3;c:\ - total: 219.90 Gb - used: 163.44 Gb (74%) - free 56.46 Gb (26%)
[1463828397] SERVICE ALERT: DB01-2012;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1463828517] SERVICE ALERT: DB01-2012;CPU Load;OK;SOFT;2;CPU Load 0% (15 min average)
[1463828785] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1463829385] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1463829505] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 62% (15 min average)
[1463832073] SERVICE FLAPPING ALERT: idm2.emich.edu;NSDS Errors;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1463832175] SERVICE FLAPPING ALERT: idm1.emich.edu;NSDS Errors;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1463832345] SERVICE FLAPPING ALERT: idm4.emich.edu;NSDS Errors;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1463834520] SERVICE FLAPPING ALERT: idm3.emich.edu;NSDS Errors;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1463837905] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463838025] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1463838145] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1463838693] SERVICE FLAPPING ALERT: BOEAPP;C:\ Drive Space;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1463838745] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (15 min average)
[1463840071] SERVICE ALERT: NETIDAD;Service: Group Policy;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1463840191] SERVICE ALERT: NETIDAD;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1463841145] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463841265] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 83% (15 min average)
[1463841543] SERVICE ALERT: EMUSCRIPT;Service: User Profile;WARNING;SOFT;1;ProfSvc: Error
[1463841543] SERVICE ALERT: EMUSCRIPT;C:\ Drive Space;UNKNOWN;SOFT;1;Free disk space : Invalid drive
[1463841663] SERVICE ALERT: EMUSCRIPT;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1463841663] SERVICE ALERT: EMUSCRIPT;C:\ Drive Space;OK;SOFT;2;c:\ - total: 79.66 Gb - used: 20.77 Gb (26%) - free 58.89 Gb (74%)
[1463842465] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463842585] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1463842705] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1463843305] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 85% (15 min average)
[1463843305] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1463843714] SERVICE ALERT: idm1.emich.edu;AD Driver;CRITICAL;SOFT;1;CRITICAL: cn=Active Directory,cn=driverset1,o=system is stopped!
[1463843833] SERVICE ALERT: idm1.emich.edu;AD Driver;CRITICAL;SOFT;2;CRITICAL: cn=Active Directory,cn=driverset1,o=system is stopped!
[1463843954] SERVICE ALERT: idm1.emich.edu;AD Driver;CRITICAL;HARD;3;CRITICAL: cn=Active Directory,cn=driverset1,o=system is stopped!
[1463845105] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1463845225] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1463845345] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 90% (15 min average)
[1463845945] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1463846065] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 86% (15 min average)
[1463846665] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463846784] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 74% (15 min average)
[1463847093] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 219.90 Gb - used: 203.67 Gb (93%) - free 16.23 Gb (7%)
[1463847213] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;SOFT;2;c:\ - total: 219.90 Gb - used: 218.10 Gb (99%) - free 1.80 Gb (1%)
[1463847333] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 219.76 Gb (100%) - free 0.14 Gb (0%)
[1463847384] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1463847503] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1463847623] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1463847933] SERVICE ALERT: BOEAPP;C:\ Drive Space;OK;HARD;3;c:\ - total: 219.90 Gb - used: 165.55 Gb (75%) - free 54.35 Gb (25%)
[1463850023] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 90% (15 min average)
[1463851184] SERVICE ALERT: idm2.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 78365 found since last run. oneLevelSearchOps OK. searchOps is too high! 78368 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463851223] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463851304] SERVICE ALERT: idm2.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 49779 found since last run. oneLevelSearchOps OK. searchOps is too high! 49781 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463851343] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463851424] SERVICE ALERT: idm2.emich.edu;LDAP Stats;WARNING;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 50472 found since last run. oneLevelSearchOps OK. searchOps is too high! 50473 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463851463] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1463852063] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 47% (15 min average)
[1463852624] SERVICE ALERT: idm2.emich.edu;LDAP Stats;OK;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463853263] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1463853333] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;SOFT;1;c:\ - total: 219.90 Gb - used: 215.56 Gb (98%) - free 4.34 Gb (2%)
[1463853383] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463853453] SERVICE ALERT: BOEAPP;C:\ Drive Space;OK;SOFT;2;c:\ - total: 219.90 Gb - used: 186.67 Gb (85%) - free 33.23 Gb (15%)
[1463853503] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 94% (15 min average)
[1463854103] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463854223] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463854343] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463855024] SERVICE ALERT: idm2.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 115826 found since last run. oneLevelSearchOps OK. searchOps is too high! 115829 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463855144] SERVICE ALERT: idm2.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 51922 found since last run. oneLevelSearchOps OK. searchOps is too high! 51924 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463855264] SERVICE ALERT: idm2.emich.edu;LDAP Stats;WARNING;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 48367 found since last run. oneLevelSearchOps OK. searchOps is too high! 48368 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463855543] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 53% (15 min average)
[1463856143] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1463856263] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463856383] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463858183] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1463858253] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 219.90 Gb - used: 198.06 Gb (90%) - free 21.84 Gb (10%)
[1463858372] SERVICE ALERT: BOEAPP;C:\ Drive Space;OK;SOFT;2;c:\ - total: 219.90 Gb - used: 197.88 Gb (90%) - free 22.02 Gb (10%)
[1463858864] SERVICE ALERT: idm2.emich.edu;LDAP Stats;OK;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463858864] SERVICE FLAPPING ALERT: idm2.emich.edu;LDAP Stats;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1463858972] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 219.90 Gb - used: 202.01 Gb (92%) - free 17.89 Gb (8%)
[1463859092] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 219.90 Gb - used: 202.47 Gb (92%) - free 17.43 Gb (8%)
[1463859212] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 203.25 Gb (92%) - free 16.65 Gb (8%)
[1463859383] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463859503] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1463859623] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463860412] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 210.73 Gb (96%) - free 9.17 Gb (4%)
[1463861423] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 54% (15 min average)
[1463862212] SERVICE ALERT: BOEAPP;C:\ Drive Space;OK;HARD;3;c:\ - total: 219.90 Gb - used: 186.77 Gb (85%) - free 33.13 Gb (15%)
[1463863223] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463863343] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 91% (15 min average)
[1463863412] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 219.90 Gb - used: 205.03 Gb (93%) - free 14.87 Gb (7%)
[1463863532] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 219.90 Gb - used: 205.03 Gb (93%) - free 14.87 Gb (7%)
[1463863652] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 205.03 Gb (93%) - free 14.87 Gb (7%)
[1463863652] SERVICE FLAPPING ALERT: BOEAPP;C:\ Drive Space;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1463863943] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1463864063] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1463864183] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463864184] SERVICE ALERT: idm1.emich.edu;OfficeAD Driver;CRITICAL;SOFT;1;CRITICAL: cn=OfficeAD,cn=driverset1,o=system is stopped!
[1463864303] SERVICE ALERT: idm1.emich.edu;OfficeAD Driver;CRITICAL;SOFT;2;CRITICAL: cn=OfficeAD,cn=driverset1,o=system is stopped!
[1463864424] SERVICE ALERT: idm1.emich.edu;OfficeAD Driver;CRITICAL;HARD;3;CRITICAL: cn=OfficeAD,cn=driverset1,o=system is stopped!
[1463865383] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 44% (15 min average)
[1463865483] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 403 found since last run. oneLevelSearchOps OK. searchOps is too high! 558 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463865603] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463865983] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1463866052] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 212.53 Gb (97%) - free 7.37 Gb (3%)
[1463866103] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 83% (15 min average)
[1463866203] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 444 found since last run. oneLevelSearchOps OK. searchOps is too high! 500 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 677 found since last run.
[1463866323] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463866544] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 2
[1463866664] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 2
[1463866703] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1463866785] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 2
[1463866823] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463866873] SERVICE ALERT: idm2.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 2
[1463866918] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 1
[1463866923] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 107 found since last run. wholeSubtreeSearchOps is too high! 745 found since last run. oneLevelSearchOps OK. searchOps is too high! 783 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 859 found since last run.
[1463866943] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463866973] SERVICE ALERT: idm1.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 1
[1463866993] SERVICE ALERT: idm2.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 1
[1463867039] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 1
[1463867043] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463867093] SERVICE ALERT: idm1.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 1
[1463867113] SERVICE ALERT: idm2.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 1
[1463867158] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 1
[1463867213] SERVICE ALERT: idm1.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 1
[1463867252] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 205.02 Gb (93%) - free 14.88 Gb (7%)
[1463867264] SERVICE FLAPPING ALERT: idm2.emich.edu;LDAP Stats;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1463867543] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 76% (15 min average)
[1463867643] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 534 found since last run. oneLevelSearchOps OK. searchOps is too high! 565 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 784 found since last run.
[1463867763] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463867852] SERVICE ALERT: BOEAPP;C:\ Drive Space;OK;HARD;3;c:\ - total: 219.90 Gb - used: 187.94 Gb (85%) - free 31.96 Gb (15%)
[1463868143] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1463868263] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 78% (15 min average)
[1463868363] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 731 found since last run. oneLevelSearchOps OK. searchOps is too high! 767 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 790 found since last run.
[1463868483] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463868863] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1463868983] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1463869083] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 120 found since last run. wholeSubtreeSearchOps is too high! 800 found since last run. oneLevelSearchOps OK. searchOps is too high! 824 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 856 found since last run.
[1463869103] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 93% (15 min average)
[1463869203] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463869803] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 684 found since last run. oneLevelSearchOps OK. searchOps is too high! 725 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 765 found since last run.
[1463869923] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463870252] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 219.90 Gb - used: 199.88 Gb (91%) - free 20.02 Gb (9%)
[1463870303] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463870372] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 219.90 Gb - used: 200.56 Gb (91%) - free 19.34 Gb (9%)
[1463870423] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1463870492] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 201.27 Gb (92%) - free 18.62 Gb (8%)
[1463870523] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 1244 found since last run. oneLevelSearchOps OK. searchOps is too high! 1273 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 798 found since last run.
[1463870543] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 89% (15 min average)
[1463870643] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463871143] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463871243] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 108 found since last run. wholeSubtreeSearchOps is too high! 577 found since last run. oneLevelSearchOps OK. searchOps is too high! 609 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 943 found since last run.
[1463871263] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1463871364] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463871383] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 82% (15 min average)
[1463871963] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 578 found since last run. oneLevelSearchOps OK. searchOps is too high! 599 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 801 found since last run.
[1463872083] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463872292] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 209.61 Gb (95%) - free 10.28 Gb (5%)
[1463872683] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 630 found since last run. oneLevelSearchOps OK. searchOps is too high! 658 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 774 found since last run.
[1463872803] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463873183] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1463873303] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 87% (15 min average)
[1463873403] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 691 found since last run. oneLevelSearchOps OK. searchOps is too high! 723 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 878 found since last run.
[1463873492] SERVICE ALERT: BOEAPP;C:\ Drive Space;OK;HARD;3;c:\ - total: 219.90 Gb - used: 188.04 Gb (86%) - free 31.86 Gb (14%)
[1463873523] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463873903] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463874023] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1463874123] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 809 found since last run. oneLevelSearchOps OK. searchOps is too high! 839 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 731 found since last run.
[1463874243] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463874623] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463874743] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1463874843] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 739 found since last run. oneLevelSearchOps OK. searchOps is too high! 760 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 896 found since last run.
[1463874863] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 94% (15 min average)
[1463874963] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463875563] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 570 found since last run. oneLevelSearchOps OK. searchOps is too high! 595 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 726 found since last run.
[1463875683] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463876063] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463876183] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 90% (15 min average)
[1463876283] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 621 found since last run. oneLevelSearchOps OK. searchOps is too high! 646 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 829 found since last run.
[1463876403] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463876783] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463876903] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1463877003] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 123 found since last run. wholeSubtreeSearchOps is too high! 733 found since last run. oneLevelSearchOps OK. searchOps is too high! 754 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 853 found since last run.
[1463877123] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463877503] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1463877623] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1463877692] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 219.90 Gb - used: 200.97 Gb (91%) - free 18.93 Gb (9%)
[1463877723] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 824 found since last run. oneLevelSearchOps OK. searchOps is too high! 866 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 866 found since last run.
[1463877743] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 93% (15 min average)
[1463877812] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 219.90 Gb - used: 201.74 Gb (92%) - free 18.16 Gb (8%)
[1463877843] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463877932] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 202.48 Gb (92%) - free 17.42 Gb (8%)
[1463878343] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463878443] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 1026 found since last run. oneLevelSearchOps OK. searchOps is too high! 1057 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 884 found since last run.
[1463878463] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463878563] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463878583] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 91% (15 min average)
[1463879131] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 210.11 Gb (96%) - free 9.79 Gb (4%)
[1463879163] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 1076 found since last run. oneLevelSearchOps OK. searchOps is too high! 1113 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 998 found since last run.
[1463879183] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463879283] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 211 found since last run. oneLevelSearchOps OK. searchOps is too high! 219 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463879303] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463879403] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463879423] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463880003] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 1152 found since last run. oneLevelSearchOps OK. searchOps is too high! 1178 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 937 found since last run.
[1463880123] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 202 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463880243] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 240 found since last run. oneLevelSearchOps OK. searchOps is too high! 251 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463880713] SERVICE ALERT: DB01-2012;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1463880833] SERVICE ALERT: DB01-2012;CPU Load;OK;SOFT;2;CPU Load 83% (15 min average)
[1463880931] SERVICE ALERT: BOEAPP;C:\ Drive Space;OK;HARD;3;c:\ - total: 219.90 Gb - used: 188.40 Gb (86%) - free 31.50 Gb (14%)
[1463881223] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463881823] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1463882423] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463883023] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463883331] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 219.90 Gb - used: 199.57 Gb (91%) - free 20.33 Gb (9%)
[1463883451] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 219.90 Gb - used: 202.03 Gb (92%) - free 17.87 Gb (8%)
[1463883571] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 219.90 Gb - used: 201.10 Gb (91%) - free 18.80 Gb (9%)
[1463884223] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463884823] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1463885371] SERVICE ALERT: BOEAPP;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 219.90 Gb - used: 219.27 Gb (100%) - free 0.63 Gb (0%)
[1463885971] SERVICE ALERT: BOEAPP;C:\ Drive Space;OK;HARD;3;c:\ - total: 219.90 Gb - used: 188.63 Gb (86%) - free 31.27 Gb (14%)
[1463886023] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463886623] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463888423] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463889023] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463889571] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 219.90 Gb - used: 200.49 Gb (91%) - free 19.41 Gb (9%)