[1456983000] HOST DOWNTIME ALERT: EMU-VALUE01;STOPPED; Host has exited from a period of scheduled downtime
[1456983000] HOST DOWNTIME ALERT: LANSWEEPER2;STOPPED; Host has exited from a period of scheduled downtime
[1456983000] HOST DOWNTIME ALERT: EMUSQLEXPRESS;STOPPED; Host has exited from a period of scheduled downtime
[1456983000] HOST DOWNTIME ALERT: HDBACKUPS;STOPPED; Host has exited from a period of scheduled downtime
[1456983000] HOST DOWNTIME ALERT: MBAMDB;STOPPED; Host has exited from a period of scheduled downtime
[1456983000] HOST DOWNTIME ALERT: EMUSCRIPT;STOPPED; Host has exited from a period of scheduled downtime
[1456983000] HOST DOWNTIME ALERT: MBAM;STOPPED; Host has exited from a period of scheduled downtime
[1456983000] HOST DOWNTIME ALERT: REFWIN;STOPPED; Host has exited from a period of scheduled downtime
[1456983000] HOST DOWNTIME ALERT: OFFICEKMS;STOPPED; Host has exited from a period of scheduled downtime
[1456983000] HOST DOWNTIME ALERT: IMAGES;STOPPED; Host has exited from a period of scheduled downtime
[1456983000] HOST DOWNTIME ALERT: SANDBOX-DB;STOPPED; Host has exited from a period of scheduled downtime
[1456983000] HOST DOWNTIME ALERT: LANSWEEPER;STOPPED; Host has exited from a period of scheduled downtime
[1456983000] HOST DOWNTIME ALERT: AD-RMAD;STOPPED; Host has exited from a period of scheduled downtime
[1456990879] SERVICE ALERT: SANDBOX-DB;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1456990999] SERVICE ALERT: SANDBOX-DB;CPU Load;OK;SOFT;2;CPU Load 76% (5 min average)
[1457004673] SERVICE ALERT: DB01-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457004793] SERVICE ALERT: DB01-2012;CPU Load;OK;SOFT;2;CPU Load 1% (5 min average)
[1457006729] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1457011327] SERVICE ALERT: porterpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1457011387] SERVICE ALERT: porterpk;Device Status;OK;SOFT;2;Open tray indicated.
[1457013966] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1457014384] SERVICE ALERT: prayharroldpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1457014387] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;1;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1457014443] SERVICE ALERT: prayharroldpk;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! All Trays Still Full
[1457014507] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;2;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1457014627] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;HARD;3;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1457014909] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1457014968] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 3% - OK!
[1457018077] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1457018137] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1457018197] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1457018798] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;PERFORM PRINTER MAINTENANCE 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Processing job from tray 4
[1457021122] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1457021182] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1457021243] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1457021827] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;OK;HARD;3;OK: No entries found
[1457026069] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;CRITICAL;SOFT;1;Yellow Cartridge HP CE262A is at CRITICAL level!
[1457026129] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;CRITICAL;SOFT;2;Yellow Cartridge HP CE262A is at CRITICAL level!
[1457026190] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;CRITICAL;HARD;3;Yellow Cartridge HP CE262A is at CRITICAL level!
[1457027127] SERVICE ALERT: hallesecondsouth;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1457027157] HOST ALERT: hallesecondsouth;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1457027187] SERVICE ALERT: hallesecondsouth;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1457027217] HOST ALERT: hallesecondsouth;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1457027247] SERVICE ALERT: hallesecondsouth;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1457027248] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1457027248] SERVICE ALERT: hallesecondsouth;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1457027277] HOST ALERT: hallesecondsouth;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1457027308] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1457027367] HOST ALERT: hallesecondsouth;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1457027368] SERVICE ALERT: sillpk;Device Status;OK;SOFT;3;LOAD TRAY 3 PLAIN LETTER To use another tray press OK LOAD TRAY 3 PLAIN LETTER 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1457027457] HOST ALERT: hallesecondsouth;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1457027534] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1457027547] HOST ALERT: hallesecondsouth;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1457027594] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1457027624] HOST ALERT: hallesecondsouth;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1457027654] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1457027684] HOST ALERT: hallesecondsouth;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1457027774] HOST ALERT: hallesecondsouth;DOWN;SOFT;9;(Host check timed out after 30.00 seconds)
[1457027864] HOST ALERT: hallesecondsouth;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1457028522] SERVICE ALERT: strongpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1457028582] SERVICE ALERT: strongpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1457028641] SERVICE ALERT: strongpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1457030068] HOST ALERT: hallesecondsouth;UP;HARD;10;PING OK - Packet loss = 66%, RTA = 0.63 ms
[1457030242] SERVICE ALERT: hallesecondsouth;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1457030242] SERVICE ALERT: hallesecondsouth;Device Status;OK;HARD;3;Processing job from tray 4
[1457030561] SERVICE ALERT: AD-RMAD;E:\ Drive Space;CRITICAL;HARD;3;e:\ - total: 299.87 Gb - used: 290.33 Gb (97%) - free 9.54 Gb (3%)
[1457030648] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 42% - OK!
[1457031228] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;HARD;3;e:\ - total: 299.87 Gb - used: 248.44 Gb (83%) - free 51.44 Gb (17%)
[1457031421] SERVICE ALERT: AD-RMAD;E:\ Drive Space;CRITICAL;HARD;3;e:\ - total: 299.87 Gb - used: 270.92 Gb (90%) - free 28.95 Gb (10%)
[1457031642] SERVICE ALERT: strongpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1457031731] SERVICE ALERT: AD-RMAD;E:\ Drive Space;OK;HARD;3;e:\ - total: 399.87 Gb - used: 248.44 Gb (62%) - free 151.43 Gb (38%)
[1457031731] SERVICE FLAPPING ALERT: AD-RMAD;E:\ Drive Space;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1457035765] SERVICE ALERT: IMAGES;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457035885] SERVICE ALERT: IMAGES;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1457037589] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;OK;HARD;3;Yellow Cartridge HP CE262A is at 100% - OK!
[1457041482] HOST ALERT: hallethirdsouth;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1457041593] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 1 Trays Empty
[1457041654] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 1 Trays Empty
[1457041694] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1457041713] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 1 Trays Empty
[1457041753] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1457041814] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1457042969] HOST ALERT: porterpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1457043013] SERVICE ALERT: marshallpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.137.101. Is the printer off?  Is the network down?
[1457043032] SERVICE ALERT: marshallpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.137.101. Is the printer off?  Is the network down?
[1457043033] HOST ALERT: porterpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.40 ms
[1457043067] SERVICE ALERT: marshallpk;Device Status;OK;SOFT;2;Checking printer
[1457043092] SERVICE ALERT: marshallpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.137.101. Is the printer off?  Is the network down?
[1457043130] SERVICE FLAPPING ALERT: AD-RMAD;E:\ Drive Space;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1457043152] SERVICE ALERT: marshallpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.137.101. Is the printer off?  Is the network down?
[1457043192] SERVICE ALERT: porterpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.133.4. Is the printer off?  Is the network down?
[1457043252] SERVICE ALERT: porterpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.133.4. Is the printer off?  Is the network down?
[1457043265] SERVICE ALERT: porterpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.133.4. Is the printer off?  Is the network down?
[1457043298] SERVICE ALERT: marshallpk;Black Cartridge HP CC364X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.137.101. Is the printer off?  Is the network down?
[1457043312] SERVICE ALERT: porterpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.133.4. Is the printer off?  Is the network down?
[1457043325] SERVICE ALERT: porterpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.133.4. Is the printer off?  Is the network down?
[1457043358] SERVICE ALERT: marshallpk;Black Cartridge HP CC364X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.137.101. Is the printer off?  Is the network down?
[1457043386] SERVICE ALERT: porterpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.133.4. Is the printer off?  Is the network down?
[1457043418] SERVICE ALERT: marshallpk;Black Cartridge HP CC364X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.137.101. Is the printer off?  Is the network down?
[1457043480] SERVICE ALERT: porterpk;Black Cartridge HP CC364X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.133.4. Is the printer off?  Is the network down?
[1457043539] SERVICE ALERT: porterpk;Black Cartridge HP CC364X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.133.4. Is the printer off?  Is the network down?
[1457043598] SERVICE ALERT: porterpk;Black Cartridge HP CC364X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.133.4. Is the printer off?  Is the network down?
[1457043614] SERVICE ALERT: crossroadspk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1457043673] SERVICE ALERT: marshallpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.137.101. Is the printer off?  Is the network down?
[1457043733] SERVICE ALERT: marshallpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.137.101. Is the printer off?  Is the network down?
[1457043793] SERVICE ALERT: marshallpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.137.101. Is the printer off?  Is the network down?
[1457043908] SERVICE ALERT: porterpk;Device Status;OK;HARD;3;Genuine HP cartridge installed
[1457044213] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1457044274] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1457044334] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1457044423] HOST ALERT: marshallpk;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1457044487] HOST ALERT: marshallpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.47 ms
[1457044512] SERVICE ALERT: porterpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.133.4. Is the printer off?  Is the network down?
[1457044572] SERVICE ALERT: porterpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.133.4. Is the printer off?  Is the network down?
[1457044612] SERVICE ALERT: marshallpk;Black Cartridge HP CC364X;OK;HARD;3;Black Cartridge HP CC364X is at 78% - OK!
[1457044632] SERVICE ALERT: porterpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.133.4. Is the printer off?  Is the network down?
[1457044947] SERVICE ALERT: marshallpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1457044987] SERVICE ALERT: marshallpk;Device Status;OK;HARD;3;Ready
[1457045082] HOST ALERT: hallethirdsouth;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1457045172] HOST ALERT: hallethirdsouth;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1457045249] HOST ALERT: hallethirdsouth;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1457045339] HOST ALERT: hallethirdsouth;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1457045378] HOST ALERT: hallethirdsouth;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1457045468] HOST ALERT: hallethirdsouth;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1457045558] HOST ALERT: hallethirdsouth;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1457045648] HOST ALERT: hallethirdsouth;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1457045738] HOST ALERT: hallethirdsouth;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1457045827] HOST ALERT: hallethirdsouth;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1457048434] SERVICE ALERT: owen101lab1;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1457048494] SERVICE ALERT: owen101lab1;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1457048554] SERVICE ALERT: owen101lab1;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1457049513] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1457050354] SERVICE ALERT: owen101lab1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1457051887] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1457051947] SERVICE ALERT: halleground2;Device Status;OK;SOFT;2;CLOSE RIGHT DOOR  46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1457052134] SERVICE ALERT: crossroadspk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1457052547] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1457052607] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1457052668] SERVICE ALERT: halleground2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!