[1463200015] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 73% (15 min average)
[1463200615] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1463200735] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1463200855] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463202655] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463203255] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463203855] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463205403] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1463205523] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1463205643] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 72% (15 min average)
[1463205665] SERVICE ALERT: INTLDB;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1463205670] SERVICE ALERT: EMUFILES;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1463205785] SERVICE ALERT: INTLDB;CPU Load;OK;SOFT;2;CPU Load 0% (15 min average)
[1463205790] SERVICE ALERT: EMUFILES;CPU Load;OK;SOFT;2;CPU Load 4% (15 min average)
[1463206843] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1463206855] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463206963] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463207083] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1463207455] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463209878] SERVICE ALERT: INTLAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463209998] SERVICE ALERT: INTLAPP;CPU Load;OK;SOFT;2;CPU Load 1% (15 min average)
[1463211883] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 91% (15 min average)
[1463212483] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463212603] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463212723] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1463212855] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463213323] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1463213323] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1463213455] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463213923] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1463214043] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 77% (15 min average)
[1463215243] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1463215255] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 90% (15 min average)
[1463215363] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1463215483] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 89% (15 min average)
[1463216083] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1463216203] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1463216323] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 87% (15 min average)
[1463216923] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1463217043] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1463217163] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 91% (15 min average)
[1463217655] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463217763] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1463217775] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1463217883] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1463217895] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 41% (15 min average)
[1463218003] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 94% (15 min average)
[1463218495] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1463218603] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1463218615] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 90% (15 min average)
[1463218724] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463218843] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1463219443] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 89% (15 min average)
[1463219815] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463219935] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 73% (15 min average)
[1463220535] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463220655] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1463220775] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463225575] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 68% (15 min average)
[1463226175] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1463226295] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1463226415] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463226864] SERVICE ALERT: BBTS;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.227.3 and port 12489: Connection refused
[1463226864] SERVICE ALERT: BBTS;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463226864] SERVICE ALERT: BBTS;VMWare Tools;CRITICAL;SOFT;1;connect to address 164.76.227.3 and port 12489: Connection refused
[1463226903] SERVICE ALERT: BBTS_DB;Size: IFDATA03.DBF;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1463226964] SERVICE ALERT: BBTS;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1463226981] SERVICE ALERT: BBTS;C:\ Drive Space;OK;SOFT;2;c:\ - total: 246.90 Gb - used: 61.38 Gb (25%) - free 185.51 Gb (75%)
[1463226984] SERVICE ALERT: BBTS;VMWare Tools;OK;SOFT;2;VMTools: Started
[1463227023] SERVICE ALERT: BBTS_DB;Size: IFDATA03.DBF;OK;SOFT;2;OK: Size is  996.0078125  Mb
[1463228213] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 79% (15 min average)
[1463228323] SERVICE ALERT: BBTS;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463228338] SERVICE ALERT: BBTS;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.227.3 and port 12489: Connection refused
[1463228423] SERVICE ALERT: BBTS;Memory Usage;OK;SOFT;2;Memory usage: total:28669.41 MB - used: 1930.18 MB (7%) - free: 26739.23 MB (93%)
[1463228458] SERVICE ALERT: BBTS;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1463229973] SERVICE ALERT: BOEAPP;VMWare Tools;WARNING;SOFT;1;No data was received from host!
[1463229982] SERVICE ALERT: BOEAPP;Service: WMI;WARNING;SOFT;1;No data was received from host!
[1463230009] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;SOFT;1;No data was received from host!
[1463230022] SERVICE ALERT: BOEAPP;Service: NetLogin;WARNING;SOFT;1;No data was received from host!
[1463230037] SERVICE ALERT: BOEAPP;Service: Windows Time;WARNING;SOFT;1;No data was received from host!
[1463230093] SERVICE ALERT: BOEAPP;VMWare Tools;WARNING;SOFT;2;No data was received from host!
[1463230102] SERVICE ALERT: BOEAPP;Service: WMI;WARNING;SOFT;2;No data was received from host!
[1463230129] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;SOFT;2;No data was received from host!
[1463230142] SERVICE ALERT: BOEAPP;Service: NetLogin;WARNING;SOFT;2;No data was received from host!
[1463230152] SERVICE ALERT: BOEAPP;Service: DNS Client;WARNING;SOFT;1;No data was received from host!
[1463230157] SERVICE ALERT: BOEAPP;Service: Windows Time;WARNING;SOFT;2;No data was received from host!
[1463230212] SERVICE ALERT: BOEAPP;Service: Firewall;WARNING;SOFT;1;No data was received from host!
[1463230213] SERVICE ALERT: BOEAPP;VMWare Tools;WARNING;HARD;3;No data was received from host!
[1463230222] SERVICE ALERT: BOEAPP;Service: WMI;WARNING;HARD;3;No data was received from host!
[1463230249] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;No data was received from host!
[1463230260] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;No data was received from host!
[1463230262] SERVICE ALERT: BOEAPP;Service: NetLogin;WARNING;HARD;3;No data was received from host!
[1463230267] SERVICE ALERT: BOEAPP;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1463230267] SERVICE ALERT: BOEAPP;Service: Windows Time;OK;SOFT;3;W32Time: Started
[1463230313] SERVICE ALERT: BOEAPP;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1463230361] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 86% (15 min average)
[1463230613] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463230733] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 54% (15 min average)
[1463230794] SERVICE ALERT: BOEAPP;VMWare Tools;OK;HARD;3;VMTools: Started
[1463230803] SERVICE ALERT: BOEAPP;Service: WMI;OK;HARD;3;Winmgmt: Started
[1463230830] SERVICE ALERT: BOEAPP;C:\ Drive Space;OK;HARD;3;c:\ - total: 219.90 Gb - used: 162.09 Gb (74%) - free 57.81 Gb (26%)
[1463230844] SERVICE ALERT: BOEAPP;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1463230961] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1463231933] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463232053] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1463232161] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1463232173] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463232281] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 91% (15 min average)
[1463232773] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 32% (15 min average)
[1463233373] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1463233493] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1463233613] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 94% (15 min average)
[1463233624] SERVICE ALERT: BBTS_DB;VMWare Tools;CRITICAL;SOFT;1;VMTools: Stopped
[1463233735] SERVICE ALERT: ADDC5;AD Group Changes;OK;HARD;3;WARNING:  0  groups have changed!
[1463233744] SERVICE ALERT: BBTS_DB;VMWare Tools;CRITICAL;SOFT;2;VMTools: Stopped
[1463233864] SERVICE ALERT: BBTS_DB;VMWare Tools;CRITICAL;HARD;3;VMTools: Stopped
[1463234213] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463234333] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463234453] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463235881] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463236001] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 43% (15 min average)
[1463236253] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 86% (15 min average)
[1463236853] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1463236973] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463237093] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 78% (15 min average)
[1463238293] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1463238413] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 81% (15 min average)
[1463239013] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1463239133] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1463239253] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463239853] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 81% (15 min average)
[1463241053] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463241173] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 84% (15 min average)
[1463241773] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1463241893] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463242013] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 93% (15 min average)
[1463242613] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463242733] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463242853] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463244081] SERVICE ALERT: studentcenterpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1463244141] SERVICE ALERT: studentcenterpk;Device Status;OK;SOFT;2;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1463245942] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1463246453] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 73% (15 min average)
[1463247053] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463247173] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 79% (15 min average)
[1463247773] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1463247893] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463248013] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1463248601] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1463248721] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1463249213] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1463250413] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463250533] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 81% (15 min average)
[1463252933] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463253053] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463253173] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463254373] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 81% (15 min average)
[1463256173] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1463256293] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463256413] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1463258321] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1463258441] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1463259041] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1463259161] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1463259281] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 91% (15 min average)
[1463261213] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 80% (15 min average)
[1463263013] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463263133] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1463263253] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463265279] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1463265399] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463265519] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 92% (15 min average)
[1463268053] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463268653] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1463274653] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463275253] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463278253] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 84% (15 min average)
[1463278410] SERVICE ALERT: VUM;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463278510] SERVICE ALERT: VUM;Memory Usage;OK;SOFT;2;Memory usage: total:14333.42 MB - used: 1740.04 MB (12%) - free: 12593.37 MB (88%)
[1463278853] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463278973] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463279023] SERVICE ALERT: LANSWEEPER;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463279093] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463279123] SERVICE ALERT: LANSWEEPER;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1463280488] SERVICE ALERT: BBTS_DB;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463280588] SERVICE ALERT: BBTS_DB;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1463280893] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463281119] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1463281239] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 90% (15 min average)
[1463281493] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463281839] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1463281959] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1463282693] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 68% (15 min average)
[1463283293] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463283413] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463283533] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463284133] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)