[1463025659] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463025677] SERVICE ALERT: ADDC1;AD Replication;OK;HARD;3;OK: replication is up and running
[1463025775] SERVICE ALERT: ADDC3;AD Replication;OK;HARD;3;OK: replication is up and running
[1463026259] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 90% (15 min average)
[1463026859] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1463026979] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1463027099] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463027400] HOST DOWNTIME ALERT: SP2010;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: PNCAPP-2012;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: OFFICEKMS;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: TI22;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: ADCA2;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: ADAUTHDC1-OLD;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: ADCA1;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: R25-ZETA2012;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: OPIE;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: ADAUTH-RMAD;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: ADAUTHDC1;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: MBAMDB;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: ADARC;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: AD-RMAD;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: TABLEAU;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: LANSWEEPER;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: SCCM-PRIM01;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: SP2010DB;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: IMAGES;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: HDBACKUPS;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: Echo-01;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: FTK;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: SANDBOX-DB;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: EMU-VALUE01;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: EMUSCRIPT;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: EMUSQLEXPRESS;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: REFWIN;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: XDM;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: ADDC1;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: DPS-SECURE02;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: COE-WEBDAV;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: DIRECTACCESS;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: R25-WEBVIEWER;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: WINKMS;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: CEDARON;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: CBORD;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: R25-WEBSERVICES;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: NESSUS;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: ADDC55;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: ADWU1;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: UHSPHARM-2012;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: ADWSUSCLIENTS;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: R25-INTERFACE;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: R25-DELTA2012;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: ADDC44;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: ADDC5;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: ADDC4;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: MBAM;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: TI22-PSYCLINIC;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: NETENG-PA;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: ADDC11;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: R25-ALPHA2012;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: PNCDB-2012;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: ADDC33;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: ADDC3;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: PROCARE;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: ADDC2;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: PNCWEB-2012;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: ADDC22;STOPPED; Host has exited from a period of scheduled downtime
[1463027400] HOST DOWNTIME ALERT: TI22-COE;STOPPED; Host has exited from a period of scheduled downtime
[1463027639] SERVICE ALERT: SCCM-PRIM01;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463027669] HOST ALERT: SCCM-PRIM01;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1463027700] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 41% (15 min average)
[1463027704] HOST ALERT: SCCM-PRIM01;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.36 ms
[1463027739] SERVICE ALERT: SCCM-PRIM01;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1463028299] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463028419] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1463028539] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463028719] SERVICE ALERT: ADCA1;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1463028839] SERVICE ALERT: ADCA1;CPU Load;OK;SOFT;2;CPU Load 0% (15 min average)
[1463033938] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463034446] SERVICE ALERT: INTLAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1463034538] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463034566] SERVICE ALERT: INTLAPP;CPU Load;OK;SOFT;2;CPU Load 0% (15 min average)
[1463035210] SERVICE ALERT: FITNESSGRAM;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463035330] SERVICE ALERT: FITNESSGRAM;CPU Load;OK;SOFT;2;CPU Load 7% (15 min average)
[1463035950] SERVICE FLAPPING ALERT: XDM;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1463036938] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463037530] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463037538] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463037650] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 80% (15 min average)
[1463038250] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1463038370] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1463039570] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463039690] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 84% (15 min average)
[1463039938] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1463040538] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463040658] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463040778] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463041378] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463041978] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 80% (15 min average)
[1463043178] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1463043298] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1463043418] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1463044018] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 90% (15 min average)
[1463044618] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1463044738] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 90% (15 min average)
[1463045338] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1463045458] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 75% (15 min average)
[1463046058] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463046178] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463046298] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 93% (15 min average)
[1463046898] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1463047018] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 88% (15 min average)
[1463048218] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463048338] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463048458] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 79% (15 min average)
[1463049058] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1463049178] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1463049298] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463049898] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 80% (15 min average)
[1463050498] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463050618] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463050738] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1463051938] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 91% (15 min average)
[1463052538] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463052658] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1463052778] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 88% (15 min average)
[1463053489] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463053609] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1463053729] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1463054329] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463054383] HOST DOWNTIME ALERT: JSS9;STARTED; Host has entered a period of scheduled downtime
[1463054399] HOST DOWNTIME ALERT: CASPER;STARTED; Host has entered a period of scheduled downtime
[1463054500] HOST ALERT: JSS9;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.98)
[1463054528] SERVICE ALERT: JSS9;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.102.98 and port 12489: No route to host
[1463054531] SERVICE ALERT: JSS9;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.102.98 and port 12489: No route to host
[1463054531] HOST ALERT: JSS9;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.98)
[1463054535] SERVICE ALERT: JSS9;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.102.98 and port 12489: No route to host
[1463054542] SERVICE ALERT: JSS9;Uptime;CRITICAL;SOFT;1;connect to address 164.76.102.98 and port 12489: No route to host
[1463054547] SERVICE ALERT: JSS9;CPU Load;CRITICAL;SOFT;1;connect to address 164.76.102.98 and port 12489: No route to host
[1463054550] HOST ALERT: JSS9;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.102.98)
[1463054562] SERVICE ALERT: CASPER;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463054562] SERVICE ALERT: CASPER;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463054578] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463054591] SERVICE ALERT: JSS9;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.102.98 and port 12489: No route to host
[1463054592] HOST ALERT: CASPER;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1463054594] HOST ALERT: JSS9;DOWN;SOFT;4;CRITICAL - Host Unreachable (164.76.102.98)
[1463054603] SERVICE ALERT: CASPER;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463054609] SERVICE ALERT: CASPER;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463054633] HOST ALERT: CASPER;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1463054641] SERVICE ALERT: JSS9;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.102.98 and port 12489: No route to host
[1463054644] HOST ALERT: JSS9;DOWN;SOFT;5;CRITICAL - Host Unreachable (164.76.102.98)
[1463054648] SERVICE ALERT: JSS9;C:\ Drive Space;CRITICAL;SOFT;2;connect to address 164.76.102.98 and port 12489: No route to host
[1463054651] SERVICE ALERT: JSS9;Service: DNS Client;CRITICAL;SOFT;2;connect to address 164.76.102.98 and port 12489: No route to host
[1463054655] SERVICE ALERT: JSS9;NSClient++ Version;CRITICAL;SOFT;2;connect to address 164.76.102.98 and port 12489: No route to host
[1463054662] SERVICE ALERT: JSS9;Uptime;CRITICAL;SOFT;2;connect to address 164.76.102.98 and port 12489: No route to host
[1463054665] SERVICE ALERT: CASPER;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463054665] HOST ALERT: JSS9;DOWN;SOFT;6;CRITICAL - Host Unreachable (164.76.102.98)
[1463054668] SERVICE ALERT: JSS9;CPU Load;CRITICAL;SOFT;2;connect to address 164.76.102.98 and port 12489: No route to host
[1463054671] SERVICE ALERT: JSS9;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.102.98 and port 12489: No route to host
[1463054682] SERVICE ALERT: CASPER;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463054682] SERVICE ALERT: CASPER;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463054695] HOST ALERT: CASPER;DOWN;SOFT;3;(Host check timed out after 30.00 seconds)
[1463054698] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463054710] SERVICE ALERT: CASPER;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463054711] SERVICE ALERT: JSS9;Memory Usage;CRITICAL;SOFT;2;connect to address 164.76.102.98 and port 12489: No route to host
[1463054711] SERVICE ALERT: JSS9;Service: NetLogin;CRITICAL;SOFT;1;connect to address 164.76.102.98 and port 12489: No route to host
[1463054714] HOST ALERT: JSS9;DOWN;SOFT;7;CRITICAL - Host Unreachable (164.76.102.98)
[1463054719] SERVICE ALERT: CASPER;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463054723] SERVICE ALERT: CASPER;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463054729] SERVICE ALERT: CASPER;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463054732] SERVICE ALERT: JSS9;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.102.98 and port 12489: No route to host
[1463054734] SERVICE ALERT: CASPER;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463054735] HOST ALERT: JSS9;DOWN;SOFT;8;CRITICAL - Host Unreachable (164.76.102.98)
[1463054740] HOST ALERT: CASPER;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1463054749] SERVICE ALERT: CASPER;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463054761] SERVICE ALERT: JSS9;Service: Group Policy;CRITICAL;SOFT;2;connect to address 164.76.102.98 and port 12489: No route to host
[1463054764] HOST ALERT: JSS9;DOWN;SOFT;9;CRITICAL - Host Unreachable (164.76.102.98)
[1463054767] SERVICE ALERT: JSS9;C:\ Drive Space;CRITICAL;HARD;3;connect to address 164.76.102.98 and port 12489: No route to host
[1463054771] SERVICE ALERT: JSS9;Service: DNS Client;CRITICAL;HARD;3;connect to address 164.76.102.98 and port 12489: No route to host
[1463054774] SERVICE ALERT: JSS9;NSClient++ Version;CRITICAL;HARD;3;connect to address 164.76.102.98 and port 12489: No route to host
[1463054779] HOST ALERT: CASPER;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1463054782] SERVICE ALERT: JSS9;Uptime;CRITICAL;HARD;3;connect to address 164.76.102.98 and port 12489: No route to host
[1463054785] SERVICE ALERT: CASPER;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463054785] SERVICE ALERT: JSS9;CPU Load;CRITICAL;HARD;3;connect to address 164.76.102.98 and port 12489: No route to host
[1463054785] HOST ALERT: JSS9;DOWN;HARD;10;CRITICAL - Host Unreachable (164.76.102.98)
[1463054789] SERVICE ALERT: CASPER;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463054792] SERVICE ALERT: JSS9;Service: Windows Time;CRITICAL;SOFT;2;connect to address 164.76.102.98 and port 12489: No route to host
[1463054801] SERVICE ALERT: JSS9;Service: User Profile;CRITICAL;SOFT;1;connect to address 164.76.102.98 and port 12489: No route to host
[1463054801] SERVICE ALERT: JSS9;VMWare Tools;CRITICAL;SOFT;1;connect to address 164.76.102.98 and port 12489: No route to host
[1463054802] SERVICE ALERT: CASPER;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1463054802] SERVICE ALERT: CASPER;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1463054815] HOST ALERT: CASPER;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1463054818] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463054821] SERVICE ALERT: JSS9;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.102.98 and port 12489: No route to host
[1463054830] SERVICE ALERT: CASPER;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463054831] SERVICE ALERT: JSS9;Memory Usage;CRITICAL;HARD;3;connect to address 164.76.102.98 and port 12489: No route to host
[1463054831] SERVICE ALERT: JSS9;Service: NetLogin;CRITICAL;SOFT;2;connect to address 164.76.102.98 and port 12489: No route to host
[1463054839] SERVICE ALERT: CASPER;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463054843] SERVICE ALERT: CASPER;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1463054849] SERVICE ALERT: JSS9;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.102.98 and port 12489: No route to host
[1463054849] SERVICE ALERT: JSS9;Service: RPC;CRITICAL;SOFT;2;connect to address 164.76.102.98 and port 12489: No route to host
[1463054849] SERVICE ALERT: CASPER;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1463054854] SERVICE ALERT: CASPER;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463054860] HOST ALERT: CASPER;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1463054866] SERVICE ALERT: CASPER;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463054869] SERVICE ALERT: CASPER;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463054881] SERVICE ALERT: JSS9;Service: Group Policy;CRITICAL;HARD;3;connect to address 164.76.102.98 and port 12489: No route to host
[1463054890] SERVICE ALERT: CASPER;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463054896] HOST ALERT: CASPER;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1463054905] SERVICE ALERT: CASPER;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1463054909] SERVICE ALERT: CASPER;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463054912] SERVICE ALERT: JSS9;Service: Windows Time;CRITICAL;HARD;3;connect to address 164.76.102.98 and port 12489: No route to host
[1463054921] SERVICE ALERT: JSS9;Service: User Profile;CRITICAL;SOFT;2;connect to address 164.76.102.98 and port 12489: No route to host
[1463054921] SERVICE ALERT: JSS9;VMWare Tools;CRITICAL;SOFT;2;connect to address 164.76.102.98 and port 12489: No route to host
[1463054929] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 40% (15 min average)
[1463054935] HOST ALERT: CASPER;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1463054941] SERVICE ALERT: JSS9;Service: Firewall;CRITICAL;SOFT;2;connect to address 164.76.102.98 and port 12489: No route to host
[1463054950] SERVICE ALERT: CASPER;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1463054954] SERVICE ALERT: JSS9;Service: NetLogin;CRITICAL;HARD;3;connect to address 164.76.102.98 and port 12489: No route to host
[1463054959] SERVICE ALERT: CASPER;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1463054965] SERVICE ALERT: CASPER;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463054969] SERVICE ALERT: JSS9;Service: RPC;CRITICAL;HARD;3;connect to address 164.76.102.98 and port 12489: No route to host
[1463054969] SERVICE ALERT: JSS9;Service: WMI;CRITICAL;SOFT;2;connect to address 164.76.102.98 and port 12489: No route to host
[1463054974] SERVICE ALERT: CASPER;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1463054980] HOST ALERT: CASPER;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1463054986] SERVICE ALERT: CASPER;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463054989] SERVICE ALERT: CASPER;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1463055010] SERVICE ALERT: CASPER;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463055029] SERVICE ALERT: CASPER;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1463055041] SERVICE ALERT: JSS9;Service: User Profile;CRITICAL;HARD;3;connect to address 164.76.102.98 and port 12489: No route to host
[1463055041] SERVICE ALERT: JSS9;VMWare Tools;CRITICAL;HARD;3;connect to address 164.76.102.98 and port 12489: No route to host
[1463055061] SERVICE ALERT: JSS9;Service: Firewall;CRITICAL;HARD;3;connect to address 164.76.102.98 and port 12489: No route to host
[1463055085] SERVICE ALERT: CASPER;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463055089] SERVICE ALERT: JSS9;Service: WMI;CRITICAL;HARD;3;connect to address 164.76.102.98 and port 12489: No route to host
[1463055106] SERVICE ALERT: CASPER;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1463055130] SERVICE ALERT: CASPER;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1463055140] SERVICE ALERT: CASPER;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463055205] SERVICE ALERT: CASPER;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1463055260] SERVICE ALERT: CASPER;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463055380] SERVICE ALERT: CASPER;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1463055928] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 109482 found since last run. oneLevelSearchOps OK. searchOps is too high! 109483 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463056018] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1463056049] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 29634 found since last run. oneLevelSearchOps OK. searchOps is too high! 29635 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463056169] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 32909 found since last run. oneLevelSearchOps OK. searchOps is too high! 32910 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463056318] SERVICE ALERT: mckennypk;Device Status;OK;HARD;3;Warming up
[1463056729] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463056849] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463056919] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1463056969] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 72% (15 min average)
[1463056978] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1463057038] SERVICE ALERT: mckennypk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1463057968] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463058418] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463058538] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1463058658] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 72% (15 min average)
[1463058769] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1463058890] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463059009] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463059009] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1463059258] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463059378] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463059498] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 76% (15 min average)
[1463059609] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1463060209] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463060329] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1463060449] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1463060698] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1463060818] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463060938] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 76% (15 min average)
[1463061052] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463061907] SERVICE ALERT: BOEAPP;Service: NetLogin;WARNING;SOFT;1;No data was received from host!
[1463062008] SERVICE ALERT: BOEAPP;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1463062138] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463062258] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463062378] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463062978] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 70% (15 min average)
[1463063578] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1463063698] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463063818] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 91% (15 min average)
[1463064418] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1463064538] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463064658] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 94% (15 min average)
[1463065229] HOST ALERT: authtestdb;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1463065258] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463065319] HOST ALERT: authtestdb;DOWN;SOFT;2;(Host check timed out after 30.00 seconds)
[1463065378] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463065409] HOST ALERT: authtestdb;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1463065438] SERVICE ALERT: mckennypk;Device Status;OK;HARD;3;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1463065498] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1463065499] HOST ALERT: authtestdb;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1463065563] HOST ALERT: authtestdb;UP;SOFT;5;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1463065849] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1463066038] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1463066098] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1463066158] SERVICE ALERT: mckennypk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1463066158] SERVICE FLAPPING ALERT: mckennypk;Device Status;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1463066501] HOST ALERT: authtestdb;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1463066591] HOST ALERT: authtestdb;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1463066681] HOST ALERT: authtestdb;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1463066698] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 81% (15 min average)
[1463066771] HOST ALERT: authtestdb;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1463066861] HOST ALERT: authtestdb;DOWN;SOFT;5;(Host check timed out after 30.00 seconds)
[1463066951] HOST ALERT: authtestdb;DOWN;SOFT;6;(Host check timed out after 30.01 seconds)
[1463067041] HOST ALERT: authtestdb;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1463067131] HOST ALERT: authtestdb;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1463067221] HOST ALERT: authtestdb;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1463067311] HOST ALERT: authtestdb;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1463067898] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463068018] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463068138] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463068249] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463068738] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463069338] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463069449] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463070649] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 89% (15 min average)
[1463071138] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 86% (15 min average)
[1463071738] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463071849] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1463071858] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463071968] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1463071978] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463072088] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 92% (15 min average)
[1463072158] SERVICE ALERT: mckennypk;Device Status;OK;HARD;3;MENUS  RETRIEVE JOB  INFORMATION  PAPER HANDLING
[1463072578] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 58% (15 min average)
[1463072758] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1463072818] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1463072878] SERVICE ALERT: mckennypk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1463073178] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1463073288] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1463073298] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463073408] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1463073418] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463073528] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1463073640] SERVICE ALERT: halleground2;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1463073700] SERVICE ALERT: halleground2;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! All Trays Still Full
[1463073782] SERVICE ALERT: hallefirst1;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1463073842] SERVICE ALERT: hallefirst1;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1463073902] SERVICE ALERT: hallefirst1;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1463074394] HOST ALERT: kingkong;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1463074444] HOST ALERT: webstage;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1463074462] HOST ALERT: webdevl;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1463074484] HOST ALERT: kingkong;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1463074534] HOST ALERT: webstage;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1463074535] HOST ALERT: phred;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1463074552] HOST ALERT: webdevl;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1463074574] HOST ALERT: kingkong;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1463074599] HOST ALERT: phred;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.45 ms
[1463074616] HOST ALERT: webdevl;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1463074624] HOST ALERT: webstage;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1463074638] HOST ALERT: kingkong;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.44 ms
[1463074688] HOST ALERT: webstage;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.44 ms
[1463075278] SERVICE ALERT: mckennypk;Device Status;OK;HARD;3;Warming up
[1463075368] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8902 found since last run. oneLevelSearchOps OK. searchOps is too high! 8903 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463075489] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 41241 found since last run. oneLevelSearchOps OK. searchOps is too high! 41242 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463075608] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 30904 found since last run. oneLevelSearchOps OK. searchOps is too high! 30905 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463075878] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1463075928] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 72% (15 min average)
[1463075938] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1463075998] SERVICE ALERT: mckennypk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1463076302] SERVICE ALERT: hallefirst1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1463076917] SERVICE ALERT: sillpk;Black Cartridge HP CC364X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.126.102. Is the printer off?  Is the network down?
[1463076947] HOST ALERT: sillpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1463076955] SERVICE ALERT: sillpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.126.102. Is the printer off?  Is the network down?
[1463076977] SERVICE ALERT: sillpk;Black Cartridge HP CC364X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.126.102. Is the printer off?  Is the network down?
[1463076985] HOST ALERT: sillpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1463077015] SERVICE ALERT: sillpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.126.102. Is the printer off?  Is the network down?
[1463077018] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463077037] SERVICE ALERT: sillpk;Black Cartridge HP CC364X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.126.102. Is the printer off?  Is the network down?
[1463077045] HOST ALERT: sillpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1463077075] SERVICE ALERT: sillpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.126.102. Is the printer off?  Is the network down?
[1463077077] SERVICE ALERT: sillpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.126.102. Is the printer off?  Is the network down?
[1463077105] HOST ALERT: sillpk;DOWN;SOFT;4;(Host check timed out after 30.00 seconds)
[1463077128] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1463077137] SERVICE ALERT: sillpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.126.102. Is the printer off?  Is the network down?
[1463077167] HOST ALERT: sillpk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1463077197] SERVICE ALERT: sillpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.126.102. Is the printer off?  Is the network down?
[1463077227] HOST ALERT: sillpk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1463077248] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 91% (15 min average)
[1463077317] HOST ALERT: sillpk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1463077407] HOST ALERT: sillpk;DOWN;SOFT;8;(Host check timed out after 30.00 seconds)
[1463077408] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463077497] HOST ALERT: sillpk;DOWN;SOFT;9;(Host check timed out after 30.00 seconds)
[1463077587] HOST ALERT: sillpk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1463077618] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1463078072] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1463078131] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1463078151] HOST ALERT: porterpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1463078190] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1463078218] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463078240] HOST ALERT: porterpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1463078266] SERVICE ALERT: porterpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.133.4. Is the printer off?  Is the network down?
[1463078296] HOST ALERT: porterpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1463078304] SERVICE ALERT: porterpk;Black Cartridge HP CC364X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.133.4. Is the printer off?  Is the network down?
[1463078316] SERVICE ALERT: porterpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.133.4. Is the printer off?  Is the network down?
[1463078326] SERVICE ALERT: porterpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.133.4. Is the printer off?  Is the network down?
[1463078334] HOST ALERT: porterpk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1463078364] SERVICE ALERT: porterpk;Black Cartridge HP CC364X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.133.4. Is the printer off?  Is the network down?
[1463078376] SERVICE ALERT: porterpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.133.4. Is the printer off?  Is the network down?
[1463078386] SERVICE ALERT: porterpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.133.4. Is the printer off?  Is the network down?
[1463078394] HOST ALERT: porterpk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1463078398] SERVICE ALERT: mckennypk;Device Status;OK;HARD;3;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1463078424] SERVICE ALERT: porterpk;Black Cartridge HP CC364X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.133.4. Is the printer off?  Is the network down?
[1463078436] SERVICE ALERT: porterpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.133.4. Is the printer off?  Is the network down?
[1463078454] HOST ALERT: porterpk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1463078544] HOST ALERT: porterpk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1463078634] HOST ALERT: porterpk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1463078724] HOST ALERT: porterpk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1463078814] HOST ALERT: porterpk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1463079048] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1463079085] SERVICE ALERT: rooseveltpk;Black Cartridge HP CC364X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.239.9. Is the printer off?  Is the network down?
[1463079115] HOST ALERT: rooseveltpk;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1463079145] SERVICE ALERT: rooseveltpk;Black Cartridge HP CC364X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.239.9. Is the printer off?  Is the network down?
[1463079154] SERVICE ALERT: rooseveltpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.239.9. Is the printer off?  Is the network down?
[1463079155] SERVICE ALERT: rooseveltpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.239.9. Is the printer off?  Is the network down?
[1463079168] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 91% (15 min average)
[1463079175] HOST ALERT: rooseveltpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1463079205] SERVICE ALERT: rooseveltpk;Black Cartridge HP CC364X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.239.9. Is the printer off?  Is the network down?
[1463079214] SERVICE ALERT: rooseveltpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.239.9. Is the printer off?  Is the network down?
[1463079215] SERVICE ALERT: rooseveltpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.239.9. Is the printer off?  Is the network down?
[1463079235] HOST ALERT: rooseveltpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1463079274] SERVICE ALERT: rooseveltpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.239.9. Is the printer off?  Is the network down?
[1463079275] SERVICE ALERT: rooseveltpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.239.9. Is the printer off?  Is the network down?
[1463079304] HOST ALERT: rooseveltpk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1463079394] HOST ALERT: rooseveltpk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1463079418] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463079484] HOST ALERT: rooseveltpk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1463079574] HOST ALERT: rooseveltpk;DOWN;SOFT;7;(Host check timed out after 30.00 seconds)
[1463079664] HOST ALERT: rooseveltpk;DOWN;SOFT;8;(Host check timed out after 30.01 seconds)
[1463079753] HOST ALERT: rooseveltpk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1463079835] HOST ALERT: rooseveltpk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1463080164] SERVICE ALERT: INTLAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1463080204] SERVICE ALERT: mckennypk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1463080234] HOST ALERT: mckennypk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1463080248] SERVICE ALERT: mckennypk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1463080264] SERVICE ALERT: mckennypk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1463080274] HOST ALERT: mckennypk;UP;SOFT;2;PING WARNING - Packet loss = 81%, RTA = 0.83 ms
[1463080282] SERVICE ALERT: INTLAPP;CPU Load;OK;SOFT;2;CPU Load 0% (15 min average)
[1463080302] SERVICE ALERT: mckennypk;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! All Trays Still Full
[1463080318] SERVICE ALERT: mckennypk;Device Status;OK;SOFT;3;Warming up
[1463080590] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1463080968] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1463081088] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 91% (15 min average)
[1463081218] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 74% (15 min average)
[1463081688] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1463081808] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1463081912] SERVICE ALERT: VUM;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463081928] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 91% (15 min average)
[1463082033] SERVICE ALERT: VUM;CPU Load;OK;SOFT;2;CPU Load 69% (15 min average)
[1463082418] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1463082528] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1463082538] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 92% (15 min average)
[1463082648] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 76% (15 min average)
[1463083047] SERVICE ALERT: VCENTERSSO;VMWare Tools;CRITICAL;SOFT;1;connect to address 164.76.158.220 and port 12489: Connection refused
[1463083049] SERVICE ALERT: VCENTERSSO;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.158.220 and port 12489: Connection refused
[1463083050] SERVICE ALERT: VCENTERSSO;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.158.220 and port 12489: Connection refused
[1463083077] HOST ALERT: VCENTERSSO;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1463083138] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1463083141] HOST ALERT: VCENTERSSO;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.40 ms
[1463083160] SERVICE ALERT: VCENTERSSO;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1463083164] SERVICE ALERT: VCENTERSSO;VMWare Tools;OK;SOFT;2;VMTools: Started
[1463083170] SERVICE ALERT: VCENTERSSO;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1463083232] SERVICE ALERT: VUM;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1463083258] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1463083352] SERVICE ALERT: VUM;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1463083378] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463083465] SERVICE ALERT: VCENTERSSO;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1463083472] SERVICE ALERT: VUM;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463083585] SERVICE ALERT: VCENTERSSO;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463083705] SERVICE ALERT: VCENTERSSO;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463084072] SERVICE ALERT: VUM;CPU Load;OK;HARD;3;CPU Load 0% (15 min average)
[1463084305] SERVICE ALERT: VCENTERSSO;CPU Load;OK;HARD;3;CPU Load 2% (15 min average)
[1463085648] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1463085768] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1463085888] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463086488] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 82% (15 min average)
[1463088118] SERVICE FLAPPING ALERT: mckennypk;Device Status;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1463088778] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 91% (15 min average)
[1463089378] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1463089498] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1463089618] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 64% (15 min average)
[1463090218] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1463090338] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463090458] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463092258] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 90% (15 min average)
[1463092488] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1463092608] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1463092728] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 92% (15 min average)
[1463092728] HOST DOWNTIME ALERT: GOVGIS;STOPPED; Host has exited from a period of scheduled downtime
[1463092858] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1463092978] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1463094778] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1463094898] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1463095018] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463095128] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1463095248] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1463095368] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1463096218] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463096568] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 73% (15 min average)
[1463096818] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463100768] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1463100888] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 71% (15 min average)
[1463101228] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1463101289] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1463101348] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1463101588] SERVICE ALERT: AX-PROD;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463101609] SERVICE ALERT: AX-PROD;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463101688] SERVICE ALERT: AX-PROD;Service: RPC;OK;SOFT;2;RpcSs: Started
[1463101709] SERVICE ALERT: AX-PROD;Memory Usage;OK;SOFT;2;Memory usage: total:20199.66 MB - used: 1017.85 MB (5%) - free: 19181.81 MB (95%)
[1463102218] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1463102818] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463102938] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463103058] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463103658] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463105598] SERVICE ALERT: INTLTESTAPP;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463105698] SERVICE ALERT: INTLTESTAPP;Service: RPC;OK;SOFT;2;RpcSs: Started
[1463106288] SERVICE ALERT: VC;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463106388] SERVICE ALERT: VC;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1463106658] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1463107258] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463107378] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1463107488] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1463107498] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463108688] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1463108698] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463108808] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1463109298] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463109408] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1463109528] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1463109648] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 89% (15 min average)
[1463110248] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1463110368] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1463110968] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1463111088] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1463111208] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1463111808] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 91% (15 min average)