[1491537637] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1491537697] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 79% (15 min average)
[1491538297] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1491538357] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 88% (15 min average)
[1491539401] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1491540157] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1491540217] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1491540277] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1491541477] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 91% (15 min average)
[1491541994] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1491542053] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;2;Warning:  idm1  are more than 120 seconds behind!
[1491542077] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1491542114] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1491542137] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1491542197] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1491543997] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 64% (15 min average)
[1491544513] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1491544597] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1491544657] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1491544717] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1491545113] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1491545114] SERVICE ALERT: FITNESSGRAM;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1491545174] SERVICE ALERT: FITNESSGRAM;CPU Load;OK;SOFT;2;CPU Load 8% (15 min average)
[1491545463] SERVICE ALERT: ADDC5;AD Group Changes;UNKNOWN;SOFT;1;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1491545698] SERVICE ALERT: ADDC5;AD Group Changes;OK;SOFT;2;WARNING:  0  groups have changed!
[1491550117] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 62% (15 min average)
[1491550717] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1491550777] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1491550837] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1491551113] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1491551113] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1491551713] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1491552037] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1491552313] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1491552637] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1491553237] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1491553800] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service has exited from a period of scheduled downtime
[1491553803] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (46.0% change >= 20.0% threshold)
[1491553803] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (30.0% change >= 20.0% threshold)
[1491553803] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1491553803] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1491553803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1491553803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1491553803] HOST DOWNTIME ALERT: INTLAPP;STARTED; Host has entered a period of scheduled downtime
[1491553803] HOST DOWNTIME ALERT: INTLDB;STARTED; Host has entered a period of scheduled downtime
[1491553803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1491553803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1491553803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1491553803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1491553803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1491553803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1491553803] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1491553803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1491553803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1491553803] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1491553803] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1491555037] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1491555097] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1491555157] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 82% (15 min average)
[1491555602] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (35.6% change >= 20.0% threshold)
[1491555602] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (24.3% change >= 20.0% threshold)
[1491555602] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1491555602] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1491555602] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1491555602] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1491555602] HOST DOWNTIME ALERT: INTLAPP;STARTED; Host has entered a period of scheduled downtime
[1491555602] HOST DOWNTIME ALERT: INTLDB;STARTED; Host has entered a period of scheduled downtime
[1491555602] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1491555602] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1491555602] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1491555602] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1491555602] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1491555602] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1491555602] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1491555602] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1491555602] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1491555602] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1491555602] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1491555757] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1491555817] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 83% (15 min average)
[1491556514] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1491556573] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1491559363] SERVICE ALERT: cava2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1491559423] SERVICE ALERT: cava2;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1491559483] SERVICE ALERT: cava2;CPU Load;OK;SOFT;3;CPU Load 0% (15 min average)
[1491560017] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1491560077] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1491560083] SERVICE ALERT: cava2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1491560137] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1491560143] SERVICE ALERT: cava2;CPU Load;OK;SOFT;2;CPU Load 7% (15 min average)
[1491560737] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 85% (15 min average)
[1491561337] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1491561397] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 66% (15 min average)
[1491561997] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1491562057] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1491562117] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 93% (15 min average)
[1491563772] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1491563917] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1491563977] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1491564037] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1491564637] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 85% (15 min average)
[1491565237] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1491565297] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1491565357] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1491565957] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 49% (15 min average)
[1491567373] SERVICE ALERT: mckennypk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1491567403] HOST ALERT: mckennypk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1491567433] SERVICE ALERT: mckennypk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1491567463] HOST ALERT: mckennypk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1491567475] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1491567493] SERVICE ALERT: mckennypk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1491567505] HOST ALERT: mckennypk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1491567535] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1491567565] HOST ALERT: mckennypk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1491567595] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1491567625] HOST ALERT: mckennypk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1491567648] SERVICE ALERT: mckennypk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1491567678] HOST ALERT: mckennypk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1491567708] SERVICE ALERT: mckennypk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1491567738] HOST ALERT: mckennypk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1491567768] SERVICE ALERT: mckennypk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1491567798] HOST ALERT: mckennypk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1491567888] HOST ALERT: mckennypk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1491567978] HOST ALERT: mckennypk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1491568357] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1491568417] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1491568477] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1491569677] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 38% (15 min average)
[1491569716] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1491569777] SERVICE ALERT: marshallpk;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! 1 of 2 Trays Still Full
[1491570243] SERVICE ALERT: marshallpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1491570303] SERVICE ALERT: marshallpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1491570363] SERVICE ALERT: marshallpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1491570589] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 56% - OK!
[1491570593] HOST ALERT: mckennypk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.67 ms
[1491570762] SERVICE ALERT: mckennypk;Device Status;OK;HARD;3;Ready
[1491571087] SERVICE ALERT: mckennypk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1491571803] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (38.8% change >= 20.0% threshold)
[1491571803] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1491571803] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1491571803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1491571803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1491571803] HOST DOWNTIME ALERT: INTLAPP;STARTED; Host has entered a period of scheduled downtime
[1491571803] HOST DOWNTIME ALERT: INTLDB;STARTED; Host has entered a period of scheduled downtime
[1491571803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1491571803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1491571803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1491571803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1491571803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1491571803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1491571803] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1491571803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1491571803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1491571803] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1491571803] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1491573646] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1491573707] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1491573766] SERVICE ALERT: hallefirst1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1491573962] SERVICE ALERT: ADDC5;AD Group Changes;UNKNOWN;SOFT;1;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1491574201] SERVICE ALERT: ADDC5;AD Group Changes;OK;SOFT;2;WARNING:  0  groups have changed!
[1491574223] SERVICE ALERT: C4SAS-SE;Memory Usage;WARNING;SOFT;1;Memory usage: total:10237.83 MB - used: 9356.81 MB (91%) - free: 881.02 MB (9%)
[1491574283] SERVICE ALERT: C4SAS-SE;Memory Usage;WARNING;SOFT;2;Memory usage: total:10237.83 MB - used: 9353.77 MB (91%) - free: 884.05 MB (9%)
[1491574343] SERVICE ALERT: C4SAS-SE;Memory Usage;WARNING;HARD;3;Memory usage: total:10237.83 MB - used: 9326.98 MB (91%) - free: 910.85 MB (9%)
[1491574564] SERVICE ALERT: marshallpk;Device Status;OK;HARD;3;Clearing paper path
[1491574891] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5306 found since last run. oneLevelSearchOps OK. searchOps is too high! 5424 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491574951] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491575163] SERVICE ALERT: marshallpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1491575223] SERVICE ALERT: marshallpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1491575283] SERVICE ALERT: marshallpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1491575543] SERVICE ALERT: C4SAS-SE;Memory Usage;OK;HARD;3;Memory usage: total:10237.83 MB - used: 9118.27 MB (89%) - free: 1119.56 MB (11%)
[1491575676] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1491575736] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 76% (15 min average)
[1491578736] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1491578796] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1491578856] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1491579166] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;Clearing paper path
[1491579456] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 49% (15 min average)
[1491579554] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1491579613] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1491579673] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1491579878] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5147 found since last run. oneLevelSearchOps OK. searchOps is too high! 5240 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491579938] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491580273] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1491580366] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1491580426] SERVICE ALERT: hallefirst1;Device Status;OK;SOFT;2;Genuine HP cartridge installed
[1491581026] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1491581087] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1491581146] SERVICE ALERT: hallefirst1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1491581738] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8655 found since last run. oneLevelSearchOps OK. searchOps is too high! 8740 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491581798] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491581856] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1491581916] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1491581976] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1491582398] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5781 found since last run. oneLevelSearchOps OK. searchOps is too high! 5886 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491582458] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491582574] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1491582603] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1491582603] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1491582603] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1491582603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1491582603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1491582603] HOST DOWNTIME ALERT: INTLAPP;STARTED; Host has entered a period of scheduled downtime
[1491582603] HOST DOWNTIME ALERT: INTLDB;STARTED; Host has entered a period of scheduled downtime
[1491582603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1491582603] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1491582603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1491582603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1491582603] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1491582603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1491582603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1491582603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1491582603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1491582603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1491582603] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1491582946] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;MENUS  RETRIEVE JOB  INFORMATION  PAPER HANDLING
[1491582946] SERVICE FLAPPING ALERT: hallefirst1;Device Status;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1491583273] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1491583333] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 3% - OK!
[1491583341] SERVICE ALERT: C4SAS-SE;Memory Usage;WARNING;SOFT;1;Memory usage: total:10237.83 MB - used: 9241.06 MB (90%) - free: 996.77 MB (10%)
[1491583351] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7393 found since last run. oneLevelSearchOps OK. searchOps is too high! 7654 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491583401] SERVICE ALERT: C4SAS-SE;Memory Usage;WARNING;SOFT;2;Memory usage: total:10237.83 MB - used: 9242.12 MB (90%) - free: 995.71 MB (10%)
[1491583410] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491583461] SERVICE ALERT: C4SAS-SE;Memory Usage;WARNING;HARD;3;Memory usage: total:10237.83 MB - used: 9242.36 MB (90%) - free: 995.47 MB (10%)
[1491583895] HOST ALERT: hallefirst1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1491583939] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1491583969] HOST ALERT: hallefirst1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1491583999] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1491584010] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5045 found since last run. oneLevelSearchOps OK. searchOps is too high! 5063 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491584029] HOST ALERT: hallefirst1;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1491584059] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1491584070] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491584089] HOST ALERT: hallefirst1;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1491584152] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1491584179] HOST ALERT: hallefirst1;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1491584212] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1491584242] HOST ALERT: hallefirst1;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1491584269] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1491584272] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1491584299] HOST ALERT: hallefirst1;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1491584329] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1491584359] HOST ALERT: hallefirst1;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1491584374] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1491584389] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1491584419] HOST ALERT: hallefirst1;DOWN;SOFT;9;(Host check timed out after 30.01 seconds)
[1491584433] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1491584493] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1491584509] HOST ALERT: hallefirst1;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1491584857] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5014 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491584883] SERVICE ALERT: marshallpk;Device Status;OK;HARD;3;MENUS  RETRIEVE JOB  INFORMATION  PAPER HANDLING
[1491584916] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491585093] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 86% (15 min average)
[1491585261] SERVICE ALERT: C4SAS-SE;Memory Usage;OK;HARD;3;Memory usage: total:10237.83 MB - used: 9178.44 MB (90%) - free: 1059.38 MB (10%)
[1491585261] SERVICE FLAPPING ALERT: C4SAS-SE;Memory Usage;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1491585516] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5145 found since last run. oneLevelSearchOps OK. searchOps is too high! 5229 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491585575] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491585693] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1491585753] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 78% (15 min average)
[1491586461] SERVICE ALERT: C4SAS-SE;Memory Usage;WARNING;SOFT;1;Memory usage: total:10237.83 MB - used: 9260.21 MB (90%) - free: 977.61 MB (10%)
[1491586521] SERVICE ALERT: C4SAS-SE;Memory Usage;WARNING;SOFT;2;Memory usage: total:10237.83 MB - used: 9260.04 MB (90%) - free: 977.79 MB (10%)
[1491586581] SERVICE ALERT: C4SAS-SE;Memory Usage;WARNING;HARD;3;Memory usage: total:10237.83 MB - used: 9258.51 MB (90%) - free: 979.31 MB (10%)
[1491587781] SERVICE ALERT: C4SAS-SE;Memory Usage;OK;HARD;3;Memory usage: total:10237.83 MB - used: 9193.42 MB (90%) - free: 1044.40 MB (10%)
[1491589581] SERVICE ALERT: C4SAS-SE;Memory Usage;WARNING;SOFT;1;Memory usage: total:10237.83 MB - used: 9249.52 MB (90%) - free: 988.31 MB (10%)
[1491589641] SERVICE ALERT: C4SAS-SE;Memory Usage;WARNING;SOFT;2;Memory usage: total:10237.83 MB - used: 9251.48 MB (90%) - free: 986.34 MB (10%)
[1491589701] SERVICE ALERT: C4SAS-SE;Memory Usage;WARNING;HARD;3;Memory usage: total:10237.83 MB - used: 9253.73 MB (90%) - free: 984.09 MB (10%)
[1491590553] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1491590613] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1491590673] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 92% (15 min average)
[1491590901] SERVICE ALERT: C4SAS-SE;Memory Usage;OK;HARD;3;Memory usage: total:10237.83 MB - used: 9192.92 MB (90%) - free: 1044.90 MB (10%)
[1491590932] SERVICE ALERT: studentcenterpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.150.252. Is the printer off?  Is the network down?
[1491590949] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.150.252. Is the printer off?  Is the network down?
[1491590962] HOST ALERT: studentcenterpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1491590987] SERVICE ALERT: studentcenterpk;Device Status;OK;SOFT;2;Genuine HP cartridge installed
[1491590991] HOST ALERT: studentcenterpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.53 ms
[1491591003] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 56% - OK!
[1491591268] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5412 found since last run. oneLevelSearchOps OK. searchOps is too high! 5550 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491591328] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491592683] SERVICE ALERT: marshallpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1491592743] SERVICE ALERT: marshallpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1491592803] SERVICE ALERT: marshallpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1491593403] SERVICE ALERT: marshallpk;Device Status;OK;HARD;3;Ready
[1491593728] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5520 found since last run. oneLevelSearchOps OK. searchOps is too high! 5601 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491593788] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491593901] SERVICE ALERT: C4SAS-SE;Memory Usage;WARNING;SOFT;1;Memory usage: total:10237.83 MB - used: 9278.67 MB (91%) - free: 959.16 MB (9%)
[1491593961] SERVICE ALERT: C4SAS-SE;Memory Usage;WARNING;SOFT;2;Memory usage: total:10237.83 MB - used: 9278.08 MB (91%) - free: 959.74 MB (9%)
[1491594021] SERVICE ALERT: C4SAS-SE;Memory Usage;WARNING;HARD;3;Memory usage: total:10237.83 MB - used: 9279.96 MB (91%) - free: 957.87 MB (9%)
[1491595072] SERVICE FLAPPING ALERT: hallefirst1;Device Status;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1491595175] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5627 found since last run. oneLevelSearchOps OK. searchOps is too high! 5637 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491595221] SERVICE ALERT: C4SAS-SE;Memory Usage;OK;HARD;3;Memory usage: total:10237.83 MB - used: 9176.92 MB (90%) - free: 1060.91 MB (10%)
[1491595235] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491595835] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5015 found since last run. oneLevelSearchOps OK. searchOps is too high! 5020 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491595895] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491596073] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1491596133] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1491596193] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 92% (15 min average)
[1491596421] SERVICE ALERT: C4SAS-SE;Memory Usage;WARNING;SOFT;1;Memory usage: total:10237.83 MB - used: 9389.99 MB (92%) - free: 847.84 MB (8%)
[1491596481] SERVICE ALERT: C4SAS-SE;Memory Usage;WARNING;SOFT;2;Memory usage: total:10237.83 MB - used: 9382.70 MB (92%) - free: 855.12 MB (8%)
[1491596541] SERVICE ALERT: C4SAS-SE;Memory Usage;WARNING;HARD;3;Memory usage: total:10237.83 MB - used: 9384.69 MB (92%) - free: 853.14 MB (8%)
[1491596793] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1491597387] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5200 found since last run. oneLevelSearchOps OK. searchOps is too high! 5315 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491597447] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491597741] SERVICE ALERT: C4SAS-SE;Memory Usage;OK;HARD;3;Memory usage: total:10237.83 MB - used: 9180.11 MB (90%) - free: 1057.71 MB (10%)
[1491598941] SERVICE ALERT: C4SAS-SE;Memory Usage;WARNING;SOFT;1;Memory usage: total:10237.83 MB - used: 9364.50 MB (91%) - free: 873.33 MB (9%)
[1491599001] SERVICE ALERT: C4SAS-SE;Memory Usage;WARNING;SOFT;2;Memory usage: total:10237.83 MB - used: 9341.69 MB (91%) - free: 896.14 MB (9%)
[1491599061] SERVICE ALERT: C4SAS-SE;Memory Usage;WARNING;HARD;3;Memory usage: total:10237.83 MB - used: 9334.83 MB (91%) - free: 903.00 MB (9%)
[1491599193] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1491599253] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 76% (15 min average)
[1491599495] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5516 found since last run. oneLevelSearchOps OK. searchOps is too high! 5523 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491599555] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491599777] SERVICE ALERT: IT-2012;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1491599837] SERVICE ALERT: IT-2012;CPU Load;OK;SOFT;2;CPU Load 7% (15 min average)
[1491599838] SERVICE ALERT: hallegroundcolor;Yellow Cartridge HP CE262A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1491599892] SERVICE ALERT: hallegroundcolor;Yellow Cartridge HP CE262A;OK;SOFT;2;Yellow Cartridge HP CE262A is at 73% - OK!
[1491599931] HOST ALERT: IT-2012;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1491599995] HOST ALERT: IT-2012;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.40 ms
[1491600453] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1491600513] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1491601113] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1491601173] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1491601233] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1491602461] SERVICE FLAPPING ALERT: C4SAS-SE;Memory Usage;STARTED; Service appears to have started flapping (29.2% change >= 20.0% threshold)
[1491602461] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1491602461] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1491602461] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1491602461] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1491602461] HOST DOWNTIME ALERT: INTLAPP;STARTED; Host has entered a period of scheduled downtime
[1491602461] HOST DOWNTIME ALERT: INTLDB;STARTED; Host has entered a period of scheduled downtime
[1491602461] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1491602461] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1491602461] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1491602461] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1491602461] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1491602461] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1491602461] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1491602461] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1491602461] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1491602461] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1491602461] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1491602495] HOST ALERT: r25testdb;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1491602497] HOST ALERT: blogs;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1491602497] HOST ALERT: grit;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1491602510] HOST ALERT: r25db;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1491602561] HOST ALERT: grit;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1491602561] HOST ALERT: blogs;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1491602573] HOST ALERT: biapp;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.229)
[1491602585] HOST ALERT: r25testdb;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1491602600] HOST ALERT: r25db;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1491602631] HOST ALERT: irimdb;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1491602636] HOST ALERT: biapp;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.4.229)
[1491602667] HOST ALERT: eagledw;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1491602675] HOST ALERT: r25testdb;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1491602690] HOST ALERT: r25db;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1491602700] HOST ALERT: biapp;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1491602713] HOST ALERT: appdb;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1491602721] HOST ALERT: irimdb;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1491602729] HOST ALERT: irimtestdb;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1491602754] HOST ALERT: r25db;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1491602754] HOST ALERT: eagledw;UP;SOFT;2;PING WARNING - Packet loss = 82%, RTA = 0.29 ms
[1491602754] HOST ALERT: r25testdb;UP;SOFT;4;PING OK - Packet loss = 75%, RTA = 0.34 ms
[1491602777] HOST ALERT: appdb;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1491602785] HOST ALERT: irimdb;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1491602793] HOST ALERT: irimtestdb;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1491602793] HOST ALERT: sindex1;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.110)
[1491602821] SERVICE ALERT: idmrbs.emich.edu;LDAP Stats;CRITICAL;SOFT;1;CRITICAL: Failed to connect to LDAP!
[1491602856] HOST ALERT: sindex1;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.4.110)
[1491602863] SERVICE ALERT: idmrbs.emich.edu;eDirectory Service;CRITICAL;SOFT;1;CRITICAL: Failed to run the ndsd status command successfully!
[1491602866] HOST ALERT: sindex2;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.111)
[1491602881] SERVICE ALERT: idmrbs.emich.edu;LDAP Stats;CRITICAL;SOFT;2;CRITICAL: Failed to connect to LDAP!
[1491602919] HOST ALERT: sindex1;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.4.110)
[1491602929] HOST ALERT: sindex2;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.4.111)
[1491602933] SERVICE ALERT: idmrbs.emich.edu;eDirectory Service;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1491602951] SERVICE ALERT: idmrbs.emich.edu;LDAP Stats;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1491602963] HOST ALERT: idmrbs.emich.edu;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1491602974] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 73% (15 min average)
[1491602983] HOST ALERT: sindex1;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.62 ms
[1491602993] HOST ALERT: sindex2;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.95 ms
[1491602993] SERVICE ALERT: idmrbs.emich.edu;eDirectory Service;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1491603005] HOST ALERT: darkdb;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1491603023] HOST ALERT: idmrbs.emich.edu;DOWN;HARD;2;PING CRITICAL - Packet loss = 100%
[1491603059] SERVICE ALERT: ADDC5;AD Group Changes;UNKNOWN;SOFT;1;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1491603072] SERVICE ALERT: NETIDAD;Uptime;CRITICAL;SOFT;1;connect to address 164.76.4.232 and port 12489: Connection refused
[1491603095] HOST ALERT: darkdb;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1491603112] SERVICE ALERT: idm1.emich.edu;Google Driver;CRITICAL;SOFT;1;Unknown: Error reported?
[1491603123] SERVICE ALERT: NETIDAD;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 1 minute(s)
[1491603138] HOST ALERT: boetestdb;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1491603145] HOST ALERT: mydb;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1491603147] HOST ALERT: idmrbs.emich.edu;UP;HARD;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1491603172] SERVICE ALERT: idm1.emich.edu;Google Driver;CRITICAL;SOFT;2;Unknown: Error reported?
[1491603172] SERVICE ALERT: idm1.emich.edu;LDAP Stats;CRITICAL;SOFT;1;CRITICAL: Failed to connect to LDAP!
[1491603185] HOST ALERT: darkdb;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1491603186] HOST ALERT: litedb;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1491603207] SERVICE ALERT: NETIDADPROXY;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1491603209] SERVICE ALERT: idm1.emich.edu;Loopback Driver;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1491603209] HOST ALERT: mydb;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1491603214] HOST ALERT: banprod;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1491603215] HOST ALERT: boedb;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1491603228] HOST ALERT: boetestdb;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1491603234] SERVICE ALERT: idmrbs.emich.edu;Banner Driver;CRITICAL;SOFT;1;CRITICAL: cn=Ellucian Banner,cn=driverset1,o=system is stopped!
[1491603242] SERVICE ALERT: idm1.emich.edu;LDAP Stats;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1491603250] HOST ALERT: litedb;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1491603267] SERVICE ALERT: NETIDADPROXY;Service: Windows Time;CRITICAL;SOFT;2;W32Time: Stopped
[1491603272] HOST ALERT: idm1.emich.edu;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1491603274] SERVICE ALERT: idm1.emich.edu;eDirectory Service;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1491603275] HOST ALERT: darkdb;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1491603276] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 11
[1491603291] SERVICE ALERT: idm1.emich.edu;Google Driver;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 60 seconds.
[1491603293] SERVICE ALERT: idmrbs.emich.edu;Banner Driver;CRITICAL;SOFT;2;CRITICAL: cn=Ellucian Banner,cn=driverset1,o=system is stopped!
[1491603302] SERVICE ALERT: idm1.emich.edu;LDAP Stats;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1491603303] SERVICE ALERT: ADRODC-DET;Service: AD Web Service;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1491603304] HOST ALERT: idm1.emich.edu;DOWN;HARD;2;PING CRITICAL - Packet loss = 100%
[1491603304] HOST ALERT: banprod;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1491603305] HOST ALERT: boedb;UP;SOFT;2;PING WARNING - Packet loss = 86%, RTA = 0.35 ms
[1491603313] SERVICE ALERT: ADDC5;AD Group Changes;OK;SOFT;2;WARNING:  0  groups have changed!
[1491603318] HOST ALERT: boetestdb;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1491603327] SERVICE ALERT: NETIDADPROXY;Service: Windows Time;OK;SOFT;3;W32Time: Started
[1491603329] SERVICE ALERT: idm1.emich.edu;Loopback Driver;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 60 seconds.
[1491603334] SERVICE ALERT: idm1.emich.edu;eDirectory Service;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1491603336] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 11
[1491603339] HOST ALERT: darkdb;UP;SOFT;5;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1491603345] HOST ALERT: idm1.emich.edu;UP;HARD;2;PING OK - Packet loss = 70%, RTA = 0.37 ms
[1491603351] SERVICE ALERT: idm1.emich.edu;Loopback Driver;CRITICAL;HARD;3;Unknown: Error reported?
[1491603353] SERVICE ALERT: ADRODC-DET;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1491603354] SERVICE ALERT: idmrbs.emich.edu;Banner Driver;OK;SOFT;3;OK: cn=Ellucian Banner,cn=driverset1,o=system is running.
[1491603368] HOST ALERT: banprod;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1491603382] HOST ALERT: boetestdb;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.39 ms
[1491603383] HOST ALERT: ADRODC-DET;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1491603383] SERVICE ALERT: ADRODC-DET;Service: Remote Management;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1491603383] SERVICE ALERT: ADRODC-DET;Service: Key Distribution Center;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1491603383] SERVICE ALERT: ADRODC-DET;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1491603383] SERVICE ALERT: ADRODC-DET;Service: AD Web Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1491603383] SERVICE ALERT: ADRODC-DET;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1491603384] SERVICE ALERT: idm1.emich.edu;eDirectory Service;OK;SOFT;3;OK: ndsd status is good.
[1491603393] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 11
[1491603396] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 11
[1491603398] SERVICE ALERT: SNAP;Service: WMI;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1491603413] SERVICE ALERT: ADRODC-DET;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1491603413] SERVICE ALERT: ADRODC-DET;Service: Distributed File System Replication;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1491603413] HOST ALERT: ADRODC-DET;DOWN;HARD;2;PING CRITICAL - Packet loss = 100%
[1491603442] SERVICE ALERT: idm1.emich.edu;NSDS Errors;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 60 seconds.
[1491603443] SERVICE ALERT: ADRODC-DET;Service: Remote Management;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1491603443] SERVICE ALERT: ADRODC-DET;Service: Key Distribution Center;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1491603443] SERVICE ALERT: ADRODC-DET;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1491603443] SERVICE ALERT: ADRODC-DET;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1491603443] SERVICE ALERT: ADRODC-DET;Service: AD Web Service;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1491603451] HOST ALERT: rmandb;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1491603453] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 11
[1491603463] SERVICE ALERT: idm1.emich.edu;Oracle Driver;CRITICAL;SOFT;1;CRITICAL: cn=oradap,cn=driverset1,o=system is stopped!
[1491603470] SERVICE ALERT: SNAP;Service: WMI;CRITICAL;SOFT;2;connect to address 164.76.4.138 and port 12489: No route to host
[1491603470] SERVICE ALERT: SNAP;Service: NetLogin;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1491603473] SERVICE ALERT: ADRODC-DET;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1491603473] SERVICE ALERT: ADRODC-DET;Service: Distributed File System Replication;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1491603473] HOST ALERT: SNAP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.138)
[1491603479] SERVICE ALERT: SNAP;CPU Load;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1491603493] SERVICE ALERT: idm1.emich.edu;AD Driver;CRITICAL;SOFT;1;CRITICAL: cn=Campus AD,cn=driverset1,o=system is stopped!
[1491603502] SERVICE ALERT: idm1.emich.edu;NSDS Errors;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 60 seconds.
[1491603503] SERVICE ALERT: ADRODC-DET;Service: Remote Management;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1491603503] SERVICE ALERT: ADRODC-DET;Service: Key Distribution Center;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1491603503] SERVICE ALERT: ADRODC-DET;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1491603503] SERVICE ALERT: ADRODC-DET;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1491603507] HOST ALERT: ADRODC-DET;UP;HARD;2;PING OK - Packet loss = 0%, RTA = 9.29 ms
[1491603513] SERVICE ALERT: ADRODC-DET;Service: Distributed File System Replication;OK;SOFT;3;DFSR: Started
[1491603513] SERVICE ALERT: idm3.emich.edu;NSDS Errors;OK;SOFT;3;OK: Total errors: 0
[1491603521] SERVICE ALERT: SNAP;Service: WMI;CRITICAL;HARD;3;connect to address 164.76.4.138 and port 12489: No route to host
[1491603524] SERVICE ALERT: idm1.emich.edu;Oracle Driver;CRITICAL;SOFT;2;CRITICAL: cn=oradap,cn=driverset1,o=system is stopped!
[1491603524] SERVICE ALERT: SNAP;Service: NetLogin;CRITICAL;SOFT;2;connect to address 164.76.4.138 and port 12489: No route to host
[1491603524] HOST ALERT: SNAP;DOWN;HARD;2;CRITICAL - Host Unreachable (164.76.4.138)
[1491603539] SERVICE ALERT: SNAP;CPU Load;CRITICAL;SOFT;2;connect to address 164.76.4.138 and port 12489: No route to host
[1491603541] HOST ALERT: rmandb;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1491603541] SERVICE ALERT: idmrbs.emich.edu;LDAP Stats;OK;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491603554] SERVICE ALERT: idm1.emich.edu;AD Driver;CRITICAL;SOFT;2;CRITICAL: cn=Campus AD,cn=driverset1,o=system is stopped!
[1491603562] SERVICE ALERT: SNAP;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.4.138 and port 12489: No route to host
[1491603562] SERVICE ALERT: idm1.emich.edu;NSDS Errors;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 60 seconds.
[1491603566] HOST ALERT: SNAP;UP;HARD;2;PING OK - Packet loss = 0%, RTA = 601.16 ms
[1491603583] SERVICE ALERT: idmrbs.emich.edu;eDirectory Service;OK;HARD;3;OK: ndsd status is good.
[1491603584] SERVICE ALERT: idm1.emich.edu;Oracle Driver;CRITICAL;HARD;3;CRITICAL: cn=oradap,cn=driverset1,o=system is stopped!
[1491603584] SERVICE ALERT: SNAP;Service: NetLogin;OK;SOFT;3;Netlogon: Started
[1491603596] SERVICE ALERT: SNAP;CPU Load;OK;SOFT;3;CPU Load 1% (15 min average)
[1491603605] HOST ALERT: rmandb;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1491603614] SERVICE ALERT: idm1.emich.edu;AD Driver;CRITICAL;HARD;3;CRITICAL: cn=Campus AD,cn=driverset1,o=system is stopped!
[1491603619] SERVICE ALERT: SNAP;Service: RPC;OK;SOFT;2;RpcSs: Started
[1491603735] HOST ALERT: ADDC7;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1491603743] HOST ALERT: authdb;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1491603783] SERVICE ALERT: ADDC7;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1491603813] SERVICE ALERT: ADDC7;AD Replication;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 30 seconds.
[1491603813] HOST ALERT: ADDC7;DOWN;HARD;2;PING CRITICAL - Packet loss = 100%
[1491603818] SERVICE ALERT: ADDC7;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1491603830] HOST ALERT: authdb;UP;SOFT;2;PING WARNING - Packet loss = 82%, RTA = 0.30 ms
[1491603831] SERVICE ALERT: idm1.emich.edu;Google Driver;OK;HARD;3;OK: cn=gdap,cn=driverset1,o=system is running.
[1491603843] SERVICE ALERT: ADDC7;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1491603852] SERVICE ALERT: ADDC7;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1491603852] SERVICE ALERT: ADDC7;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1491603853] SERVICE ALERT: ADDC7;Service: Dell Backup Agent;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1491603862] SERVICE ALERT: ADDC7;Service: Key Distribution Center;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1491603872] HOST ALERT: ADDC7;UP;HARD;2;PING OK - Packet loss = 76%, RTA = 0.32 ms
[1491603873] SERVICE ALERT: ADDC7;AD Replication;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 30 seconds.
[1491603878] SERVICE ALERT: ADDC7;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1491603878] SERVICE ALERT: ADDC7;Service: Distributed File System Replication;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1491603878] SERVICE ALERT: ADDC7;Service: Intersite Messaging;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1491603879] SERVICE ALERT: ADDC7;Service: DFS Namespace;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1491603883] SERVICE ALERT: ADDC7;CPU Load;OK;SOFT;3;CPU Load 2% (15 min average)
[1491603884] SERVICE ALERT: ADDC7;Service: DNS Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1491603886] SERVICE ALERT: ADDC7;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1491603892] SERVICE ALERT: ADDC7;Memory Usage;OK;SOFT;2;Memory usage: total:56310.69 MB - used: 1605.06 MB (3%) - free: 54705.63 MB (97%)
[1491603892] SERVICE ALERT: ADDC7;Service: RPC;OK;SOFT;2;RpcSs: Started
[1491603892] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491603893] SERVICE ALERT: ADDC7;Service: Dell Backup Agent;CRITICAL;SOFT;2;ErdAgent: Starting
[1491603901] SERVICE ALERT: ADFS;Service: Group Policy;CRITICAL;SOFT;1;gpsvc: Stopped
[1491603902] SERVICE ALERT: ADDC7;Service: Key Distribution Center;OK;SOFT;2;KDC: Started
[1491603908] SERVICE ALERT: ADDC7;AD Replication;OK;SOFT;3;OK: replication is up and running
[1491603918] SERVICE ALERT: ADDC7;Service: Group Policy;OK;SOFT;3;gpsvc: Started
[1491603918] SERVICE ALERT: ADDC7;Service: Distributed File System Replication;OK;SOFT;2;DFSR: Started
[1491603918] SERVICE ALERT: ADDC7;Service: Intersite Messaging;OK;SOFT;2;IsmServ: Started
[1491603919] SERVICE ALERT: ADDC7;Service: DFS Namespace;OK;SOFT;2;Dfs: Started
[1491603919] SERVICE ALERT: idm2.emich.edu;LDAP Stats;UNKNOWN;SOFT;1;Difference is negative. LDAP recently rebooted.
[1491603924] SERVICE ALERT: ADDC7;Service: DNS Service;OK;SOFT;2;DNS: Started
[1491603926] SERVICE ALERT: ADDC7;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1491603929] SERVICE ALERT: idm1.emich.edu;Loopback Driver;OK;HARD;3;OK: cn=Generic Loopback,cn=driverset1,o=system is running.
[1491603953] SERVICE ALERT: ADDC7;Service: Dell Backup Agent;CRITICAL;HARD;3;ErdAgent: Starting
[1491603961] SERVICE ALERT: ADFS;Service: Group Policy;CRITICAL;SOFT;2;gpsvc: Stopped
[1491603979] SERVICE ALERT: idm2.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491603997] SERVICE ALERT: idm4.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1491604021] SERVICE ALERT: ADFS;Service: Group Policy;CRITICAL;HARD;3;gpsvc: Stopped
[1491604023] SERVICE ALERT: ADRODC-DET;Service: AD Web Service;OK;HARD;3;ADWS: Started
[1491604053] SERVICE ALERT: ADRODC-DET;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1491604058] HOST DOWNTIME ALERT: AD-RMAD;STARTED; Host has entered a period of scheduled downtime
[1491604058] HOST DOWNTIME ALERT: ADCA1;STARTED; Host has entered a period of scheduled downtime
[1491604058] HOST DOWNTIME ALERT: ADCA2;STARTED; Host has entered a period of scheduled downtime
[1491604058] HOST DOWNTIME ALERT: ADDC1;STARTED; Host has entered a period of scheduled downtime
[1491604058] HOST DOWNTIME ALERT: ADDC2;STARTED; Host has entered a period of scheduled downtime
[1491604058] HOST DOWNTIME ALERT: ADDC4;STARTED; Host has entered a period of scheduled downtime
[1491604058] HOST DOWNTIME ALERT: ADDC5;STARTED; Host has entered a period of scheduled downtime
[1491604058] HOST DOWNTIME ALERT: ADDC6;STARTED; Host has entered a period of scheduled downtime
[1491604058] HOST DOWNTIME ALERT: ADDC7;STARTED; Host has entered a period of scheduled downtime
[1491604058] HOST DOWNTIME ALERT: ADDFS01;STARTED; Host has entered a period of scheduled downtime
[1491604058] HOST DOWNTIME ALERT: ADRODC-DET;STARTED; Host has entered a period of scheduled downtime
[1491604058] HOST DOWNTIME ALERT: ADRODC-LIV;STARTED; Host has entered a period of scheduled downtime
[1491604058] HOST DOWNTIME ALERT: ADRODC-WEMU;STARTED; Host has entered a period of scheduled downtime
[1491604058] HOST DOWNTIME ALERT: ADSECURED01;STARTED; Host has entered a period of scheduled downtime
[1491604058] HOST DOWNTIME ALERT: ADWSUSCLIENTS;STARTED; Host has entered a period of scheduled downtime
[1491604058] HOST DOWNTIME ALERT: ADWU1;STARTED; Host has entered a period of scheduled downtime
[1491604058] HOST DOWNTIME ALERT: ALERTUS;STARTED; Host has entered a period of scheduled downtime
[1491604058] HOST DOWNTIME ALERT: AX-PROD;STARTED; Host has entered a period of scheduled downtime
[1491604058] HOST DOWNTIME ALERT: AXIOM-2012;STARTED; Host has entered a period of scheduled downtime
[1491604058] HOST DOWNTIME ALERT: AXLICENSE;STARTED; Host has entered a period of scheduled downtime
[1491604058] HOST DOWNTIME ALERT: BOEAPP;STARTED; Host has entered a period of scheduled downtime
[1491604058] HOST DOWNTIME ALERT: C4SAS-SE;STARTED; Host has entered a period of scheduled downtime
[1491604058] HOST DOWNTIME ALERT: CASPER;STARTED; Host has entered a period of scheduled downtime
[1491604058] HOST DOWNTIME ALERT: CLIENTLANSWEEP;STARTED; Host has entered a period of scheduled downtime
[1491604058] HOST DOWNTIME ALERT: CSCFILESERV;STARTED; Host has entered a period of scheduled downtime
[1491604058] HOST DOWNTIME ALERT: DB01-2012;STARTED; Host has entered a period of scheduled downtime
[1491604058] HOST DOWNTIME ALERT: EMU-VALUE01;STARTED; Host has entered a period of scheduled downtime
[1491604058] HOST DOWNTIME ALERT: EMUFILES;STARTED; Host has entered a period of scheduled downtime
[1491604058] HOST DOWNTIME ALERT: EMUPRINTPS;STARTED; Host has entered a period of scheduled downtime
[1491604058] HOST DOWNTIME ALERT: EVENTS-2012;STARTED; Host has entered a period of scheduled downtime
[1491604058] HOST DOWNTIME ALERT: FITNESSGRAM;STARTED; Host has entered a period of scheduled downtime
[1491604058] HOST DOWNTIME ALERT: GOVGIS;STARTED; Host has entered a period of scheduled downtime
[1491604058] HOST DOWNTIME ALERT: IMAGES;STARTED; Host has entered a period of scheduled downtime
[1491604058] HOST DOWNTIME ALERT: JSS9;STARTED; Host has entered a period of scheduled downtime
[1491604058] HOST DOWNTIME ALERT: LANSWEEPER;STARTED; Host has entered a period of scheduled downtime
[1491604058] HOST DOWNTIME ALERT: MAPS01;STARTED; Host has entered a period of scheduled downtime
[1491604058] HOST DOWNTIME ALERT: NETIDAD;STARTED; Host has entered a period of scheduled downtime
[1491604058] HOST DOWNTIME ALERT: NETIDADPROXY;STARTED; Host has entered a period of scheduled downtime
[1491604058] HOST DOWNTIME ALERT: PLANT-METASYS;STARTED; Host has entered a period of scheduled downtime
[1491604058] HOST DOWNTIME ALERT: SNAP;STARTED; Host has entered a period of scheduled downtime
[1491604058] HOST DOWNTIME ALERT: TARGETX-2012;STARTED; Host has entered a period of scheduled downtime
[1491604058] HOST DOWNTIME ALERT: emu-fndn-fe7;STARTED; Host has entered a period of scheduled downtime
[1491604058] HOST DOWNTIME ALERT: epprogsql2;STARTED; Host has entered a period of scheduled downtime
[1491604083] SERVICE ALERT: ADRODC-DET;Service: Remote Management;OK;HARD;3;WinRM: Started
[1491604083] SERVICE ALERT: ADRODC-DET;Service: Key Distribution Center;OK;HARD;3;KDC: Started
[1491604083] SERVICE ALERT: ADRODC-DET;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1491604083] SERVICE ALERT: ADRODC-DET;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1491604102] SERVICE ALERT: idm1.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 13
[1491604118] SERVICE ALERT: SNAP;Service: WMI;OK;HARD;3;Winmgmt: Started
[1491604172] SERVICE ALERT: CSCFILESERV;Uptime;CRITICAL;SOFT;1;connect to address 164.76.4.42 and port 12489: No route to host
[1491604175] HOST ALERT: CSCFILESERV;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.42)
[1491604183] SERVICE ALERT: idm1.emich.edu;Oracle Driver;OK;HARD;3;OK: cn=oradap,cn=driverset1,o=system is running.
[1491604213] SERVICE ALERT: idm1.emich.edu;AD Driver;OK;HARD;3;OK: cn=Campus AD,cn=driverset1,o=system is running.
[1491604214] SERVICE ALERT: idm2.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 3
[1491604216] SERVICE ALERT: CSCFILESERV;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.4.42 and port 12489: No route to host
[1491604219] HOST ALERT: CSCFILESERV;DOWN;HARD;2;CRITICAL - Host Unreachable (164.76.4.42)
[1491604223] SERVICE ALERT: CSCFILESERV;Uptime;CRITICAL;SOFT;2;connect to address 164.76.4.42 and port 12489: No route to host
[1491604274] SERVICE ALERT: idm2.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 2
[1491604276] SERVICE ALERT: CSCFILESERV;C:\ Drive Space;CRITICAL;SOFT;2;connect to address 164.76.4.42 and port 12489: No route to host
[1491604280] SERVICE ALERT: CSCFILESERV;Uptime;CRITICAL;HARD;3;connect to address 164.76.4.42 and port 12489: Connection refused
[1491604333] SERVICE ALERT: CSCFILESERV;C:\ Drive Space;OK;SOFT;3;c:\ - total: 58.50 Gb - used: 28.13 Gb (48%) - free 30.37 Gb (52%)
[1491604334] SERVICE ALERT: idm2.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 2
[1491604337] HOST ALERT: CSCFILESERV;UP;HARD;2;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1491604355] SERVICE ALERT: idm3.emich.edu;LDAP Stats;UNKNOWN;SOFT;1;Difference is negative. LDAP recently rebooted.
[1491604384] SERVICE ALERT: ADSECURED01;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1491604393] SERVICE ALERT: ADSECURED01;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1491604393] HOST ALERT: ADSECURED01;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.199)
[1491604415] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491604416] SERVICE ALERT: ADSECURED01;Service: NetLogin;CRITICAL;SOFT;1;connect to address 164.76.4.199 and port 12489: No route to host
[1491604419] HOST ALERT: ADSECURED01;DOWN;HARD;2;CRITICAL - Host Unreachable (164.76.4.199)
[1491604424] SERVICE ALERT: ADDC5;Service: Windows Time;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1491604427] SERVICE ALERT: ADSECURED01;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.4.199 and port 12489: No route to host
[1491604427] SERVICE ALERT: ADSECURED01;CPU Load;CRITICAL;SOFT;2;connect to address 164.76.4.199 and port 12489: No route to host
[1491604430] SERVICE ALERT: ADSECURED01;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.4.199 and port 12489: No route to host
[1491604431] SERVICE ALERT: ADDC5;Service: Dell Backup Agent;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1491604433] SERVICE ALERT: ADSECURED01;Service: Windows Time;CRITICAL;SOFT;2;connect to address 164.76.4.199 and port 12489: No route to host
[1491604438] SERVICE ALERT: ADDC5;AD Group Changes;WARNING;SOFT;1;WARNING:  0  groups have changed!
[1491604439] SERVICE ALERT: ADDC5;AD Group Changes;UNKNOWN;SOFT;2;The command (winteampowershell) returned an invalid return code: -65536
[1491604439] SERVICE ALERT: ADDC5;AD Replication;UNKNOWN;SOFT;1;The command (addcreplicationcheck) returned an invalid return code: -65536
[1491604445] SERVICE ALERT: ADSECURED01;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.4.199 and port 12489: No route to host
[1491604457] SERVICE ALERT: ADSECURED01;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.4.199 and port 12489: No route to host
[1491604460] SERVICE ALERT: ADSECURED01;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.4.199 and port 12489: No route to host
[1491604461] SERVICE ALERT: ADDC5;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.7.205 and port 12489: Connection refused
[1491604464] SERVICE ALERT: ADDC5;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.7.205 and port 12489: Connection refused
[1491604465] SERVICE ALERT: ADDC5;Uptime;CRITICAL;SOFT;1;connect to address 164.76.7.205 and port 12489: Connection refused
[1491604467] SERVICE ALERT: ADSECURED01;Uptime;CRITICAL;SOFT;1;connect to address 164.76.4.199 and port 12489: No route to host
[1491604470] SERVICE ALERT: ADSECURED01;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.4.199 and port 12489: No route to host
[1491604476] SERVICE ALERT: ADSECURED01;Service: NetLogin;CRITICAL;SOFT;2;connect to address 164.76.4.199 and port 12489: No route to host
[1491604484] SERVICE ALERT: ADDC5;Service: Windows Time;CRITICAL;SOFT;2;W32Time: Stopped
[1491604487] SERVICE ALERT: ADSECURED01;CPU Load;CRITICAL;HARD;3;connect to address 164.76.4.199 and port 12489: No route to host
[1491604487] SERVICE ALERT: ADSECURED01;C:\ Drive Space;CRITICAL;SOFT;2;connect to address 164.76.4.199 and port 12489: No route to host
[1491604490] SERVICE ALERT: ADSECURED01;Service: Group Policy;CRITICAL;SOFT;2;connect to address 164.76.4.199 and port 12489: No route to host
[1491604491] SERVICE ALERT: ADDC5;Service: Dell Backup Agent;CRITICAL;SOFT;2;ErdAgent: Starting
[1491604496] SERVICE ALERT: ADSECURED01;Service: Windows Time;CRITICAL;HARD;3;connect to address 164.76.4.199 and port 12489: No route to host
[1491604496] SERVICE ALERT: ADSECURED01;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.4.199 and port 12489: No route to host
[1491604499] SERVICE ALERT: ADDC5;AD Replication;OK;SOFT;2;OK: replication is up and running
[1491604505] SERVICE ALERT: ADSECURED01;Service: Firewall;CRITICAL;SOFT;2;connect to address 164.76.4.199 and port 12489: No route to host
[1491604516] SERVICE ALERT: ADDC5;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 1 minute(s)
[1491604517] SERVICE ALERT: ADSECURED01;Service: RPC;CRITICAL;SOFT;2;connect to address 164.76.4.199 and port 12489: No route to host
[1491604518] SERVICE ALERT: ADDC5;C:\ Drive Space;OK;SOFT;2;c:\ - total: 79.66 Gb - used: 42.26 Gb (53%) - free 37.39 Gb (47%)
[1491604520] SERVICE ALERT: ADSECURED01;Service: WMI;CRITICAL;SOFT;2;connect to address 164.76.4.199 and port 12489: No route to host
[1491604524] SERVICE ALERT: ADDC5;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1491604527] SERVICE ALERT: ADSECURED01;Uptime;CRITICAL;SOFT;2;connect to address 164.76.4.199 and port 12489: No route to host
[1491604530] SERVICE ALERT: ADSECURED01;Service: DNS Client;CRITICAL;SOFT;2;connect to address 164.76.4.199 and port 12489: No route to host
[1491604536] SERVICE ALERT: ADSECURED01;Service: NetLogin;CRITICAL;HARD;3;connect to address 164.76.4.199 and port 12489: No route to host
[1491604544] SERVICE ALERT: ADDC5;Service: Windows Time;CRITICAL;HARD;3;W32Time: Stopped
[1491604547] SERVICE ALERT: ADSECURED01;C:\ Drive Space;CRITICAL;HARD;3;connect to address 164.76.4.199 and port 12489: No route to host
[1491604550] SERVICE ALERT: ADSECURED01;Service: Group Policy;CRITICAL;HARD;3;connect to address 164.76.4.199 and port 12489: No route to host
[1491604551] SERVICE ALERT: ADDC5;Service: Dell Backup Agent;CRITICAL;HARD;3;ErdAgent: Starting
[1491604553] SERVICE ALERT: ADDC7;Service: Dell Backup Agent;OK;HARD;3;ErdAgent: Started
[1491604558] SERVICE ALERT: ADSECURED01;NSClient++ Version;CRITICAL;SOFT;2;connect to address 164.76.4.199 and port 12489: No route to host
[1491604565] SERVICE ALERT: ADSECURED01;Service: Firewall;CRITICAL;HARD;3;connect to address 164.76.4.199 and port 12489: No route to host
[1491604577] SERVICE ALERT: ADSECURED01;Service: RPC;CRITICAL;HARD;3;connect to address 164.76.4.199 and port 12489: No route to host
[1491604580] SERVICE ALERT: ADSECURED01;Service: WMI;CRITICAL;HARD;3;connect to address 164.76.4.199 and port 12489: No route to host
[1491604587] SERVICE ALERT: ADSECURED01;Uptime;CRITICAL;HARD;3;connect to address 164.76.4.199 and port 12489: No route to host
[1491604590] SERVICE ALERT: ADSECURED01;Service: DNS Client;CRITICAL;HARD;3;connect to address 164.76.4.199 and port 12489: No route to host
[1491604596] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 2
[1491604596] HOST ALERT: EMU-VALUE01;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.85)
[1491604615] SERVICE ALERT: ADSECURED01;NSClient++ Version;OK;SOFT;3;NSClient++ 0.4.4.15 2015-11-25
[1491604619] HOST ALERT: ADSECURED01;UP;HARD;2;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1491604619] SERVICE ALERT: EMU-VALUE01;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.4.85 and port 12489: No route to host
[1491604621] SERVICE ALERT: ADFS;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1491604622] HOST ALERT: EMU-VALUE01;DOWN;HARD;2;CRITICAL - Host Unreachable (164.76.4.85)
[1491604656] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 2
[1491604679] SERVICE ALERT: EMU-VALUE01;C:\ Drive Space;CRITICAL;SOFT;2;connect to address 164.76.4.85 and port 12489: No route to host
[1491604697] SERVICE ALERT: ADDC5;AD Group Changes;OK;SOFT;3;WARNING:  0  groups have changed!
[1491604714] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 2
[1491604716] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 2
[1491604739] SERVICE ALERT: EMU-VALUE01;C:\ Drive Space;CRITICAL;HARD;3;connect to address 164.76.4.85 and port 12489: No route to host
[1491604760] SERVICE ALERT: EMU-VALUE01;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.4.85 and port 12489: No route to host
[1491604773] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 2
[1491604773] SERVICE ALERT: EMU-VALUE01;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.4.85 and port 12489: No route to host
[1491604774] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1491604820] SERVICE ALERT: EMU-VALUE01;Service: Firewall;CRITICAL;SOFT;2;connect to address 164.76.4.85 and port 12489: No route to host
[1491604833] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 2
[1491604833] SERVICE ALERT: EMU-VALUE01;Memory Usage;CRITICAL;SOFT;2;connect to address 164.76.4.85 and port 12489: No route to host
[1491604834] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1491604866] SERVICE ALERT: EMU-VALUE01;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.4.85 and port 12489: No route to host
[1491604866] SERVICE ALERT: EMU-VALUE01;Service: NetLogin;CRITICAL;SOFT;1;connect to address 164.76.4.85 and port 12489: No route to host
[1491604866] SERVICE ALERT: EMU-VALUE01;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.4.85 and port 12489: No route to host
[1491604878] SERVICE ALERT: EMU-VALUE01;Service: Firewall;CRITICAL;HARD;3;connect to address 164.76.4.85 and port 12489: Connection refused
[1491604880] SERVICE ALERT: CSCFILESERV;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 10 minute(s)
[1491604890] SERVICE ALERT: EMU-VALUE01;Memory Usage;CRITICAL;HARD;3;connect to address 164.76.4.85 and port 12489: Connection refused
[1491604893] SERVICE ALERT: EMU-VALUE01;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.4.85 and port 12489: Connection refused
[1491604894] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1491604915] SERVICE ALERT: EMU-VALUE01;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1491604919] SERVICE ALERT: EMU-VALUE01;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1491604919] SERVICE ALERT: EMU-VALUE01;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1491604921] SERVICE ALERT: EMU-VALUE01;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1491604924] SERVICE ALERT: EMU-VALUE01;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1491604924] SERVICE ALERT: EMU-VALUE01;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1491604924] SERVICE ALERT: EMU-VALUE01;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1491604934] SERVICE ALERT: idm2.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1491604935] HOST ALERT: EMU-VALUE01;UP;HARD;2;PING OK - Packet loss = 66%, RTA = 0.33 ms
[1491604953] SERVICE ALERT: EMU-VALUE01;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1491604955] SERVICE ALERT: EMU-VALUE01;Service: RPC;OK;SOFT;2;RpcSs: Started
[1491604959] SERVICE ALERT: EMU-VALUE01;CPU Load;OK;SOFT;2;CPU Load 21% (15 min average)
[1491604959] SERVICE ALERT: EMU-VALUE01;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1491604961] SERVICE ALERT: EMU-VALUE01;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1491605015] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6608 found since last run. oneLevelSearchOps OK. searchOps is too high! 6789 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491605024] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1491605075] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491605084] SERVICE ALERT: ADSECURED01;CPU Load;OK;HARD;3;CPU Load 7% (15 min average)
[1491605084] SERVICE ALERT: BOEAPPTEST;CPU Load;OK;SOFT;2;CPU Load 9% (15 min average)
[1491605093] SERVICE ALERT: ADSECURED01;Service: Windows Time;OK;HARD;3;W32Time: Started
[1491605133] SERVICE ALERT: ADSECURED01;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1491605144] SERVICE ALERT: ADSECURED01;C:\ Drive Space;OK;HARD;3;c:\ - total: 278.86 Gb - used: 16.46 Gb (6%) - free 262.40 Gb (94%)
[1491605144] SERVICE ALERT: ADDC5;Service: Windows Time;OK;HARD;3;W32Time: Started
[1491605147] SERVICE ALERT: ADSECURED01;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1491605151] SERVICE ALERT: ADDC5;Service: Dell Backup Agent;OK;HARD;3;ErdAgent: Started
[1491605162] SERVICE ALERT: ADSECURED01;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1491605174] SERVICE ALERT: ADSECURED01;Service: RPC;OK;HARD;3;RpcSs: Started
[1491605179] SERVICE ALERT: ADSECURED01;Service: WMI;OK;HARD;3;Winmgmt: Started
[1491605184] SERVICE ALERT: ADSECURED01;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 11 minute(s)
[1491605187] SERVICE ALERT: ADSECURED01;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1491605247] SERVICE ALERT: idm4.emich.edu;LDAP Stats;UNKNOWN;SOFT;1;Difference is negative. LDAP recently rebooted.
[1491605307] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491605336] SERVICE ALERT: EMU-VALUE01;C:\ Drive Space;OK;HARD;3;c:\ - total: 465.42 Gb - used: 17.42 Gb (4%) - free 448.00 Gb (96%)
[1491605477] SERVICE ALERT: EMU-VALUE01;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1491605490] SERVICE ALERT: EMU-VALUE01;Memory Usage;OK;HARD;3;Memory usage: total:9444.43 MB - used: 2926.68 MB (31%) - free: 6517.75 MB (69%)
[1491605494] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1491605494] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1491605686] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1491605717] SERVICE ALERT: TARGETX-2012;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1491605745] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1491605777] SERVICE ALERT: TARGETX-2012;CPU Load;OK;SOFT;2;CPU Load 85% (15 min average)
[1491605804] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1491605902] SERVICE ALERT: idm1.emich.edu;NSDS Errors;OK;HARD;3;OK: 0 Errors Reported
[1491605916] SERVICE ALERT: idm4.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1491605916] SERVICE FLAPPING ALERT: idm4.emich.edu;NSDS Errors;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1491606033] SERVICE ALERT: idm3.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1491606094] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 80% (15 min average)
[1491606363] SERVICE ALERT: ADCA2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1491606377] SERVICE ALERT: TARGETX-2012;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1491606404] SERVICE ALERT: BOEAPPTEST;CPU Load;OK;HARD;3;CPU Load 14% (15 min average)
[1491606423] SERVICE ALERT: ADCA2;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1491606437] SERVICE ALERT: TARGETX-2012;CPU Load;OK;SOFT;2;CPU Load 53% (15 min average)
[1491606445] SERVICE ALERT: TARGETX-2012;Service: NetLogin;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1491606458] SERVICE ALERT: ADCA1;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1491606479] SERVICE ALERT: TARGETX-2012;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1491606483] SERVICE ALERT: ADCA2;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1491606506] SERVICE ALERT: TARGETX-2012;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1491606509] HOST ALERT: TARGETX-2012;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1491606518] SERVICE ALERT: ADCA1;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1491606519] SERVICE ALERT: TARGETX-2012;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1491606523] HOST ALERT: TARGETX-2012;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1491606597] HOST ALERT: MAPS01;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.17)
[1491606613] SERVICE ALERT: JSS9;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1491606627] HOST ALERT: MAPS01;UP;SOFT;2;PING OK - Packet loss = 44%, RTA = 0.38 ms
[1491606653] SERVICE ALERT: JSS9;VMWare Tools;OK;SOFT;2;VMTools: Started
[1491607083] SERVICE ALERT: ADCA2;CPU Load;OK;HARD;3;CPU Load 5% (15 min average)
[1491607205] SERVICE ALERT: ADDC1;Service: Dell Backup Agent;CRITICAL;SOFT;1;ErdAgent: Starting
[1491607211] SERVICE ALERT: ADDC6;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1491607265] SERVICE ALERT: ADDC1;Service: Dell Backup Agent;OK;SOFT;2;ErdAgent: Started
[1491607271] SERVICE ALERT: ADDC6;Service: Windows Time;CRITICAL;SOFT;2;W32Time: Stopped
[1491607286] SERVICE ALERT: ADDC6;Service: Dell Backup Agent;CRITICAL;SOFT;1;ErdAgent: Starting
[1491607294] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1491607294] SERVICE ALERT: ADDC4;AD Replication;WARNING;SOFT;1;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1491607331] SERVICE ALERT: ADDC6;Service: Windows Time;OK;SOFT;3;W32Time: Started
[1491607336] SERVICE ALERT: ADDC2;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1491607346] SERVICE ALERT: ADDC6;Service: Dell Backup Agent;OK;SOFT;2;ErdAgent: Started
[1491607354] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1491607355] SERVICE ALERT: ADDC4;AD Replication;WARNING;SOFT;2;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1491607376] SERVICE ALERT: ADDC2;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1491607414] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 93% (15 min average)
[1491607415] SERVICE ALERT: ADDC4;AD Replication;WARNING;HARD;3;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1491607458] SERVICE ALERT: ADDC2;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1491607461] SERVICE ALERT: ADDC4;Service: Key Distribution Center;CRITICAL;SOFT;1;connect to address 164.76.7.204 and port 12489: Connection refused
[1491607463] SERVICE ALERT: ADDC4;Service: DFS Namespace;CRITICAL;SOFT;1;connect to address 164.76.7.204 and port 12489: Connection refused
[1491607495] SERVICE ALERT: ADDC5;AD Replication;WARNING;SOFT;1;WARNING: Replication with ADDC2 failed!
[1491607518] SERVICE ALERT: ADDC2;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1491607521] SERVICE ALERT: ADDC4;Service: Key Distribution Center;OK;SOFT;2;KDC: Started
[1491607523] SERVICE ALERT: ADDC4;Service: DFS Namespace;OK;SOFT;2;Dfs: Started
[1491607553] SERVICE ALERT: ADDC6;AD Replication;WARNING;SOFT;1;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1491607560] SERVICE ALERT: ADDC5;AD Replication;OK;SOFT;2;OK: replication is up and running
[1491607613] SERVICE ALERT: ADDC6;AD Replication;WARNING;SOFT;2;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1491607678] SERVICE ALERT: ADDC6;AD Replication;OK;SOFT;3;OK: replication is up and running
[1491607707] SERVICE ALERT: cava2;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1491607708] SERVICE ALERT: epprogsql2;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.4.204 and port 12489: Connection refused
[1491607708] SERVICE ALERT: epprogsql2;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.4.204 and port 12489: Connection refused
[1491607715] SERVICE ALERT: cava2;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1491607737] HOST ALERT: cava2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1491607755] SERVICE ALERT: cava2;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1491607759] HOST ALERT: cava2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1491607767] SERVICE ALERT: cava2;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1491607768] SERVICE ALERT: epprogsql2;Service: RPC;OK;SOFT;2;RpcSs: Started
[1491607768] SERVICE ALERT: epprogsql2;Memory Usage;OK;SOFT;2;Memory usage: total:20479.73 MB - used: 2439.76 MB (12%) - free: 18039.97 MB (88%)
[1491607904] SERVICE ALERT: cava2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1491607964] SERVICE ALERT: cava2;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1491608018] SERVICE ALERT: ADDC4;AD Replication;OK;HARD;3;OK: replication is up and running
[1491608024] SERVICE ALERT: cava2;CPU Load;OK;SOFT;3;CPU Load 91% (15 min average)
[1491608037] SERVICE ALERT: EVENTS-2012;Uptime;CRITICAL;SOFT;1;connect to address 164.76.4.144 and port 12489: Connection refused
[1491608047] SERVICE ALERT: EVENTS-2012;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1491608060] SERVICE ALERT: C4SAS-SE;Memory Usage;CRITICAL;HARD;3;Memory usage: total:10237.83 MB - used: 9832.58 MB (96%) - free: 405.25 MB (4%)
[1491608087] SERVICE ALERT: EVENTS-2012;C:\ Drive Space;OK;SOFT;2;c:\ - total: 82.66 Gb - used: 35.90 Gb (43%) - free 46.76 Gb (57%)
[1491608088] SERVICE ALERT: EVENTS-2012;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 1 minute(s)
[1491608505] SERVICE ALERT: AX-PROD;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1491608518] SERVICE ALERT: epprogsql2;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.4.204 and port 12489: No route to host
[1491608518] SERVICE ALERT: epprogsql2;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.4.204 and port 12489: No route to host
[1491608565] SERVICE ALERT: AX-PROD;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1491608566] SERVICE ALERT: epprogsql2;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1491608566] SERVICE ALERT: epprogsql2;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1491608602] SERVICE ALERT: CASPER;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1491608614] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1491608653] SERVICE ALERT: CASPER;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1491608660] SERVICE ALERT: C4SAS-SE;Memory Usage;WARNING;HARD;3;Memory usage: total:10237.83 MB - used: 9704.42 MB (95%) - free: 533.40 MB (5%)
[1491608674] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1491608698] SERVICE ALERT: AXLICENSE;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1491608714] SERVICE ALERT: CASPER;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1491608734] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1491608758] SERVICE ALERT: AXLICENSE;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1491608774] SERVICE ALERT: CASPER;CPU Load;OK;SOFT;2;CPU Load 25% (15 min average)
[1491608818] SERVICE ALERT: AXLICENSE;CPU Load;OK;SOFT;3;CPU Load 52% (15 min average)
[1491608977] SERVICE ALERT: ADRODC-WEMU;Service: AD Domain Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1491609007] HOST ALERT: ADRODC-WEMU;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1491609037] SERVICE ALERT: ADRODC-WEMU;Service: AD Domain Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1491609067] HOST ALERT: ADRODC-WEMU;DOWN;HARD;2;PING CRITICAL - Packet loss = 100%
[1491609072] SERVICE ALERT: ADRODC-WEMU;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.241.195 and port 12489: Connection refused
[1491609073] SERVICE ALERT: ADRODC-WEMU;Service: DFS Namespace;CRITICAL;SOFT;1;connect to address 164.76.241.195 and port 12489: Connection refused
[1491609077] SERVICE ALERT: ADRODC-WEMU;Service: AD Domain Service;CRITICAL;HARD;3;connect to address 164.76.241.195 and port 12489: Connection refused
[1491609085] HOST ALERT: ADRODC-WEMU;UP;HARD;2;PING OK - Packet loss = 64%, RTA = 0.47 ms
[1491609123] SERVICE ALERT: ADRODC-WEMU;Memory Usage;OK;SOFT;2;Memory usage: total:21654.72 MB - used: 1239.54 MB (6%) - free: 20415.19 MB (94%)
[1491609127] SERVICE ALERT: EMUFILES;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1491609133] SERVICE ALERT: ADRODC-WEMU;Service: DFS Namespace;OK;SOFT;2;Dfs: Started
[1491609187] SERVICE ALERT: EMUFILES;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1491609247] SERVICE ALERT: EMUFILES;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1491609260] SERVICE ALERT: C4SAS-SE;Memory Usage;OK;HARD;3;Memory usage: total:10237.83 MB - used: 1965.19 MB (19%) - free: 8272.63 MB (81%)
[1491609334] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1491609603] SERVICE FLAPPING ALERT: C4SAS-SE;Memory Usage;STARTED; Service appears to have started flapping (26.1% change >= 20.0% threshold)
[1491609603] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (38.4% change >= 20.0% threshold)
[1491609603] SERVICE FLAPPING ALERT: idm4.emich.edu;NSDS Errors;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1491609603] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1491609603] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1491609603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1491609603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1491609603] HOST DOWNTIME ALERT: INTLAPP;STARTED; Host has entered a period of scheduled downtime
[1491609603] HOST DOWNTIME ALERT: ADWSUSCLIENTS;STARTED; Host has entered a period of scheduled downtime
[1491609603] HOST DOWNTIME ALERT: ADWU1;STARTED; Host has entered a period of scheduled downtime
[1491609603] HOST DOWNTIME ALERT: INTLDB;STARTED; Host has entered a period of scheduled downtime
[1491609603] HOST DOWNTIME ALERT: AX-PROD;STARTED; Host has entered a period of scheduled downtime
[1491609603] HOST DOWNTIME ALERT: AXIOM-2012;STARTED; Host has entered a period of scheduled downtime
[1491609603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1491609603] HOST DOWNTIME ALERT: AD-RMAD;STARTED; Host has entered a period of scheduled downtime
[1491609603] HOST DOWNTIME ALERT: AXLICENSE;STARTED; Host has entered a period of scheduled downtime
[1491609603] HOST DOWNTIME ALERT: BOEAPP;STARTED; Host has entered a period of scheduled downtime
[1491609603] HOST DOWNTIME ALERT: ADCA1;STARTED; Host has entered a period of scheduled downtime
[1491609603] HOST DOWNTIME ALERT: C4SAS-SE;STARTED; Host has entered a period of scheduled downtime
[1491609603] HOST DOWNTIME ALERT: CASPER;STARTED; Host has entered a period of scheduled downtime
[1491609603] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1491609603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1491609603] HOST DOWNTIME ALERT: ADCA2;STARTED; Host has entered a period of scheduled downtime
[1491609603] HOST DOWNTIME ALERT: CLIENTLANSWEEP;STARTED; Host has entered a period of scheduled downtime
[1491609603] HOST DOWNTIME ALERT: CSCFILESERV;STARTED; Host has entered a period of scheduled downtime
[1491609603] HOST DOWNTIME ALERT: ADDC1;STARTED; Host has entered a period of scheduled downtime
[1491609603] HOST DOWNTIME ALERT: DB01-2012;STARTED; Host has entered a period of scheduled downtime
[1491609603] HOST DOWNTIME ALERT: EMU-VALUE01;STARTED; Host has entered a period of scheduled downtime
[1491609603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1491609603] HOST DOWNTIME ALERT: ADDC2;STARTED; Host has entered a period of scheduled downtime
[1491609603] HOST DOWNTIME ALERT: EMUFILES;STARTED; Host has entered a period of scheduled downtime
[1491609603] HOST DOWNTIME ALERT: EMUPRINTPS;STARTED; Host has entered a period of scheduled downtime
[1491609603] HOST DOWNTIME ALERT: ADDC4;STARTED; Host has entered a period of scheduled downtime
[1491609603] HOST DOWNTIME ALERT: EVENTS-2012;STARTED; Host has entered a period of scheduled downtime
[1491609603] HOST DOWNTIME ALERT: FITNESSGRAM;STARTED; Host has entered a period of scheduled downtime
[1491609603] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1491609603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1491609603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1491609603] HOST DOWNTIME ALERT: ADDC5;STARTED; Host has entered a period of scheduled downtime
[1491609603] HOST DOWNTIME ALERT: GOVGIS;STARTED; Host has entered a period of scheduled downtime
[1491609603] HOST DOWNTIME ALERT: IMAGES;STARTED; Host has entered a period of scheduled downtime
[1491609603] HOST DOWNTIME ALERT: ADDC6;STARTED; Host has entered a period of scheduled downtime
[1491609603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1491609603] HOST DOWNTIME ALERT: ADDC7;STARTED; Host has entered a period of scheduled downtime
[1491609603] HOST DOWNTIME ALERT: JSS9;STARTED; Host has entered a period of scheduled downtime
[1491609603] HOST DOWNTIME ALERT: LANSWEEPER;STARTED; Host has entered a period of scheduled downtime
[1491609603] HOST DOWNTIME ALERT: ADDFS01;STARTED; Host has entered a period of scheduled downtime
[1491609603] HOST DOWNTIME ALERT: MAPS01;STARTED; Host has entered a period of scheduled downtime
[1491609603] HOST DOWNTIME ALERT: NETIDAD;STARTED; Host has entered a period of scheduled downtime
[1491609603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1491609603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1491609603] HOST DOWNTIME ALERT: ADRODC-DET;STARTED; Host has entered a period of scheduled downtime
[1491609603] HOST DOWNTIME ALERT: NETIDADPROXY;STARTED; Host has entered a period of scheduled downtime
[1491609603] HOST DOWNTIME ALERT: PLANT-METASYS;STARTED; Host has entered a period of scheduled downtime
[1491609603] HOST DOWNTIME ALERT: ADRODC-LIV;STARTED; Host has entered a period of scheduled downtime
[1491609603] HOST DOWNTIME ALERT: SNAP;STARTED; Host has entered a period of scheduled downtime
[1491609603] HOST DOWNTIME ALERT: TARGETX-2012;STARTED; Host has entered a period of scheduled downtime
[1491609603] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1491609603] HOST DOWNTIME ALERT: ADRODC-WEMU;STARTED; Host has entered a period of scheduled downtime
[1491609603] HOST DOWNTIME ALERT: emu-fndn-fe7;STARTED; Host has entered a period of scheduled downtime
[1491609603] HOST DOWNTIME ALERT: epprogsql2;STARTED; Host has entered a period of scheduled downtime
[1491609603] HOST DOWNTIME ALERT: ADSECURED01;STARTED; Host has entered a period of scheduled downtime
[1491609677] SERVICE ALERT: ADRODC-WEMU;Service: AD Domain Service;OK;HARD;3;NTDS: Started
[1491609799] SERVICE ALERT: EMUFILES;Service: User Profile;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1491609818] SERVICE ALERT: EMUFILES;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.4.5 and port 12489: Connection refused
[1491609847] SERVICE ALERT: EMUFILES;CPU Load;OK;HARD;3;CPU Load 69% (15 min average)
[1491609859] SERVICE ALERT: EMUFILES;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1491609869] SERVICE ALERT: EMUFILES;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1491609934] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 31% (15 min average)
[1491611734] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1491611794] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1491611854] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1491613054] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 83% (15 min average)
[1491613654] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1491613714] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 89% (15 min average)
[1491614314] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1491614374] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1491615574] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1491615634] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1491615694] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1491616001] SERVICE ALERT: IMAGES;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1491616011] SERVICE ALERT: IMAGES;MDT Monitor;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1491616022] HOST ALERT: IMAGES;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1491616061] SERVICE ALERT: IMAGES;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1491616091] HOST ALERT: IMAGES;DOWN;HARD;2;PING CRITICAL - Packet loss = 100%
[1491616101] SERVICE ALERT: IMAGES;CPU Load;OK;SOFT;3;CPU Load 2% (15 min average)
[1491616103] SERVICE ALERT: IMAGES;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1491616104] SERVICE ALERT: IMAGES;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1491616105] HOST ALERT: IMAGES;UP;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1491616143] SERVICE ALERT: IMAGES;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1491616144] SERVICE ALERT: IMAGES;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1491616294] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 53% (15 min average)
[1491616591] SERVICE ALERT: IMAGES;MDT Monitor;UNKNOWN;HARD;3;Failed to open service MDT_Monitor: 1060: The specified service does not exist as an installed service.
[1491616716] SERVICE FLAPPING ALERT: idm4.emich.edu;NSDS Errors;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1491616894] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1491616954] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1491617014] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1491617396] SERVICE ALERT: ADDC5;AD Group Changes;UNKNOWN;SOFT;1;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1491617614] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (15 min average)
[1491617618] SERVICE ALERT: ADDC5;AD Group Changes;OK;SOFT;2;WARNING:  0  groups have changed!
[1491619413] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1491619473] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1491619533] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1491620133] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 54% (15 min average)
[1491620659] SERVICE FLAPPING ALERT: C4SAS-SE;Memory Usage;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1491620733] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1491620793] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1491620853] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1491622049] HOST DOWNTIME ALERT: ADCA1;STOPPED; Host has exited from a period of scheduled downtime
[1491622049] HOST DOWNTIME ALERT: ADWU1;STOPPED; Host has exited from a period of scheduled downtime
[1491622049] HOST DOWNTIME ALERT: ADWSUSCLIENTS;STOPPED; Host has exited from a period of scheduled downtime
[1491622049] HOST DOWNTIME ALERT: AXIOM-2012;STOPPED; Host has exited from a period of scheduled downtime
[1491622049] HOST DOWNTIME ALERT: AX-PROD;STOPPED; Host has exited from a period of scheduled downtime
[1491622049] HOST DOWNTIME ALERT: BOEAPP;STOPPED; Host has exited from a period of scheduled downtime
[1491622049] HOST DOWNTIME ALERT: AXLICENSE;STOPPED; Host has exited from a period of scheduled downtime
[1491622049] HOST DOWNTIME ALERT: AD-RMAD;STOPPED; Host has exited from a period of scheduled downtime
[1491622049] HOST DOWNTIME ALERT: C4SAS-SE;STOPPED; Host has exited from a period of scheduled downtime
[1491622049] HOST DOWNTIME ALERT: CSCFILESERV;STOPPED; Host has exited from a period of scheduled downtime
[1491622049] HOST DOWNTIME ALERT: CASPER;STOPPED; Host has exited from a period of scheduled downtime
[1491622049] HOST DOWNTIME ALERT: ADDC4;STOPPED; Host has exited from a period of scheduled downtime
[1491622049] HOST DOWNTIME ALERT: EVENTS-2012;STOPPED; Host has exited from a period of scheduled downtime
[1491622049] HOST DOWNTIME ALERT: CLIENTLANSWEEP;STOPPED; Host has exited from a period of scheduled downtime
[1491622049] HOST DOWNTIME ALERT: ADRODC-LIV;STOPPED; Host has exited from a period of scheduled downtime
[1491622049] HOST DOWNTIME ALERT: ADCA2;STOPPED; Host has exited from a period of scheduled downtime
[1491622049] HOST DOWNTIME ALERT: ADDC1;STOPPED; Host has exited from a period of scheduled downtime
[1491622049] HOST DOWNTIME ALERT: ADRODC-WEMU;STOPPED; Host has exited from a period of scheduled downtime
[1491622049] HOST DOWNTIME ALERT: EMU-VALUE01;STOPPED; Host has exited from a period of scheduled downtime
[1491622049] HOST DOWNTIME ALERT: DB01-2012;STOPPED; Host has exited from a period of scheduled downtime
[1491622049] HOST DOWNTIME ALERT: EMUPRINTPS;STOPPED; Host has exited from a period of scheduled downtime
[1491622049] HOST DOWNTIME ALERT: EMUFILES;STOPPED; Host has exited from a period of scheduled downtime
[1491622049] HOST DOWNTIME ALERT: ADDC2;STOPPED; Host has exited from a period of scheduled downtime
[1491622049] HOST DOWNTIME ALERT: FITNESSGRAM;STOPPED; Host has exited from a period of scheduled downtime
[1491622049] HOST DOWNTIME ALERT: GOVGIS;STOPPED; Host has exited from a period of scheduled downtime
[1491622049] HOST DOWNTIME ALERT: ADDC5;STOPPED; Host has exited from a period of scheduled downtime
[1491622049] HOST DOWNTIME ALERT: ADDFS01;STOPPED; Host has exited from a period of scheduled downtime
[1491622049] HOST DOWNTIME ALERT: ADDC6;STOPPED; Host has exited from a period of scheduled downtime
[1491622049] HOST DOWNTIME ALERT: emu-fndn-fe7;STOPPED; Host has exited from a period of scheduled downtime
[1491622049] HOST DOWNTIME ALERT: LANSWEEPER;STOPPED; Host has exited from a period of scheduled downtime
[1491622049] HOST DOWNTIME ALERT: JSS9;STOPPED; Host has exited from a period of scheduled downtime
[1491622049] HOST DOWNTIME ALERT: ADDC7;STOPPED; Host has exited from a period of scheduled downtime
[1491622049] HOST DOWNTIME ALERT: ADSECURED01;STOPPED; Host has exited from a period of scheduled downtime
[1491622049] HOST DOWNTIME ALERT: epprogsql2;STOPPED; Host has exited from a period of scheduled downtime
[1491622049] HOST DOWNTIME ALERT: MAPS01;STOPPED; Host has exited from a period of scheduled downtime
[1491622049] HOST DOWNTIME ALERT: SNAP;STOPPED; Host has exited from a period of scheduled downtime
[1491622049] HOST DOWNTIME ALERT: NETIDAD;STOPPED; Host has exited from a period of scheduled downtime
[1491622049] HOST DOWNTIME ALERT: NETIDADPROXY;STOPPED; Host has exited from a period of scheduled downtime
[1491622049] HOST DOWNTIME ALERT: ADRODC-DET;STOPPED; Host has exited from a period of scheduled downtime
[1491622049] HOST DOWNTIME ALERT: PLANT-METASYS;STOPPED; Host has exited from a period of scheduled downtime
[1491622049] HOST DOWNTIME ALERT: TARGETX-2012;STOPPED; Host has exited from a period of scheduled downtime
[1491622200] HOST DOWNTIME ALERT: IMAGES;STOPPED; Host has exited from a period of scheduled downtime
[1491623251] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 48% (15 min average)