[1477628737] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477628866] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 101 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477628986] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477629336] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 81% (15 min average)
[1477629934] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1477630054] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1477630174] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 80% (15 min average)
[1477630774] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1477630894] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1477631014] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 94% (15 min average)
[1477631394] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1477631515] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;2;Warning:  idm1  are more than 120 seconds behind!
[1477631634] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;3;OK. All replicas within 120 seconds.
[1477632214] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1477632236] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1477632334] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1477632355] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1477632454] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477632600] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1477633054] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1477633554] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1477633654] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1477633674] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1477633774] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1477633894] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477634274] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1477634395] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1477634994] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1477635113] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1477635233] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;3;OK. All replicas within 120 seconds.
[1477636294] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477636433] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1477636553] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;2;Warning:  idm1  are more than 120 seconds behind!
[1477636673] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1477636894] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477637494] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1477638094] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1477638213] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1477638333] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477638688] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1477638808] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1477638928] SERVICE ALERT: epprogsql2;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477638932] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 86% (15 min average)
[1477639528] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477639532] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1477639653] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1477639673] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1477639772] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1477640273] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1477640393] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1477640514] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;3;OK. All replicas within 120 seconds.
[1477640728] SERVICE ALERT: epprogsql2;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477641328] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477641328] SERVICE FLAPPING ALERT: epprogsql2;CPU Load;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1477641713] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1477641833] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1477641953] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1477642172] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 86% (15 min average)
[1477642553] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1477642553] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1477643128] SERVICE ALERT: epprogsql2;CPU Load;OK;HARD;3;CPU Load 30% (15 min average)
[1477643403] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (49.2% change >= 20.0% threshold)
[1477643403] SERVICE FLAPPING ALERT: epprogsql2;CPU Load;STARTED; Service appears to have started flapping (28.6% change >= 20.0% threshold)
[1477643403] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1477643403] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1477643403] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1477643403] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1477643403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1477643403] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1477643403] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1477643403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1477643403] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1477643403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1477643403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1477643403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1477643403] HOST DOWNTIME ALERT: VUM;STARTED; Host has entered a period of scheduled downtime
[1477643403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1477643403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1477643403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1477643403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1477643403] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1477643403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1477643753] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1477645172] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1477645203] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (27.0% change >= 20.0% threshold)
[1477645203] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1477645203] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1477645203] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1477645203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1477645203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1477645203] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1477645203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1477645203] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1477645203] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1477645203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1477645203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1477645203] HOST DOWNTIME ALERT: VUM;STARTED; Host has entered a period of scheduled downtime
[1477645203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1477645203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1477645203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1477645203] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1477645203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1477645203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1477645292] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1477645292] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (41.6% change >= 20.0% threshold)
[1477647692] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1477647812] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 51% (15 min average)
[1477649012] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1477649132] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1477649252] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477649852] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477650600] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service has exited from a period of scheduled downtime
[1477651051] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477652584] HOST ALERT: hallefirst1;UP;HARD;10;PING OK - Packet loss = 61%, RTA = 0.54 ms
[1477652740] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 53% - OK!
[1477652749] SERVICE ALERT: hallefirst1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1477652836] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;Ready
[1477653451] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1477654051] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1477654171] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1477654291] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477654553] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1477654985] SERVICE ALERT: marshallpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1477655491] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1477656091] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1477656211] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 85% (15 min average)
[1477658873] SERVICE ALERT: strongpk;Device Status;OK;HARD;3;LOAD TRAY 3 PLAIN LETTER 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1477658877] SERVICE ALERT: strongpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1477658936] SERVICE ALERT: strongpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1477658997] SERVICE ALERT: strongpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1477659586] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5039 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1477659706] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477659811] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1477659931] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 55% (15 min average)
[1477660906] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5042 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1477661026] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477661041] HOST DOWNTIME ALERT: ADAUTHDC1;STARTED; Host has entered a period of scheduled downtime
[1477661054] HOST DOWNTIME ALERT: ADAUTHDC1-OLD;STARTED; Host has entered a period of scheduled downtime
[1477661060] HOST DOWNTIME ALERT: ADAUTHDC3-NEW;STARTED; Host has entered a period of scheduled downtime
[1477661114] HOST ALERT: ADAUTHDC1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1477661131] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1477661136] SERVICE ALERT: ADAUTHDC1;VMWare Tools;CRITICAL;SOFT;1;connect to address 164.76.4.207 and port 12489: No route to host
[1477661139] HOST ALERT: ADAUTHDC1;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.4.207)
[1477661182] SERVICE ALERT: ADAUTHDC1;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.4.207 and port 12489: No route to host
[1477661185] HOST ALERT: ADAUTHDC1;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.4.207)
[1477661246] HOST ALERT: ADAUTHDC1;DOWN;SOFT;4;CRITICAL - Host Unreachable (164.76.4.207)
[1477661246] SERVICE ALERT: ADAUTHDC1;Service: Dell Backup Agent;CRITICAL;SOFT;1;connect to address 164.76.4.207 and port 12489: No route to host
[1477661251] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 84% (15 min average)
[1477661256] SERVICE ALERT: ADAUTHDC1;VMWare Tools;CRITICAL;SOFT;2;connect to address 164.76.4.207 and port 12489: No route to host
[1477661265] SERVICE ALERT: ADAUTHDC1-OLD;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.4.206 and port 12489: No route to host
[1477661268] SERVICE ALERT: ADAUTHDC1-OLD;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.4.206 and port 12489: No route to host
[1477661268] HOST ALERT: ADAUTHDC1-OLD;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.206)
[1477661271] SERVICE ALERT: ADAUTHDC1-OLD;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.4.206 and port 12489: No route to host
[1477661279] SERVICE ALERT: ADAUTHDC1;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.4.207 and port 12489: No route to host
[1477661282] HOST ALERT: ADAUTHDC1;DOWN;SOFT;5;CRITICAL - Host Unreachable (164.76.4.207)
[1477661284] SERVICE ALERT: ADAUTHDC1-OLD;Service: Key Distribution Center;CRITICAL;SOFT;1;connect to address 164.76.4.206 and port 12489: No route to host
[1477661287] HOST ALERT: ADAUTHDC1-OLD;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.4.206)
[1477661302] SERVICE ALERT: ADAUTHDC1;NSClient++ Version;CRITICAL;SOFT;2;connect to address 164.76.4.207 and port 12489: No route to host
[1477661305] HOST ALERT: ADAUTHDC1;DOWN;SOFT;6;CRITICAL - Host Unreachable (164.76.4.207)
[1477661314] SERVICE ALERT: ADAUTHDC1;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.4.207 and port 12489: No route to host
[1477661327] SERVICE ALERT: ADAUTHDC3-NEW;VMWare Tools;CRITICAL;SOFT;1;connect to address 164.76.4.208 and port 12489: No route to host
[1477661330] HOST ALERT: ADAUTHDC3-NEW;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.208)
[1477661333] SERVICE ALERT: ADAUTHDC1;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.4.207 and port 12489: No route to host
[1477661336] HOST ALERT: ADAUTHDC1;DOWN;SOFT;7;CRITICAL - Host Unreachable (164.76.4.207)
[1477661340] SERVICE ALERT: ADAUTHDC3-NEW;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1477661350] HOST ALERT: ADAUTHDC1-OLD;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.4.206)
[1477661366] SERVICE ALERT: ADAUTHDC1;Service: Dell Backup Agent;CRITICAL;SOFT;2;connect to address 164.76.4.207 and port 12489: No route to host
[1477661369] HOST ALERT: ADAUTHDC1;DOWN;SOFT;8;CRITICAL - Host Unreachable (164.76.4.207)
[1477661372] SERVICE ALERT: ADAUTHDC3-NEW;Service: NetLogin;CRITICAL;SOFT;1;connect to address 164.76.4.208 and port 12489: No route to host
[1477661375] HOST ALERT: ADAUTHDC3-NEW;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.4.208)
[1477661376] SERVICE ALERT: ADAUTHDC1;VMWare Tools;CRITICAL;HARD;3;connect to address 164.76.4.207 and port 12489: No route to host
[1477661376] SERVICE ALERT: ADAUTHDC1-OLD;Service: Group Policy;CRITICAL;SOFT;2;connect to address 164.76.4.206 and port 12489: No route to host
[1477661379] HOST ALERT: ADAUTHDC1-OLD;DOWN;SOFT;4;CRITICAL - Host Unreachable (164.76.4.206)
[1477661390] SERVICE ALERT: ADAUTHDC1-OLD;C:\ Drive Space;CRITICAL;SOFT;2;connect to address 164.76.4.206 and port 12489: No route to host
[1477661393] SERVICE ALERT: ADAUTHDC1;AD Replication;WARNING;SOFT;1;(No output on stdout) stderr: Could not resolve hostname ADAUTHDC1: Name or service not known
[1477661394] SERVICE ALERT: ADAUTHDC1-OLD;Service: WMI;CRITICAL;SOFT;2;connect to address 164.76.4.206 and port 12489: No route to host
[1477661396] SERVICE ALERT: ADAUTHDC1;C:\ Drive Space;CRITICAL;SOFT;2;connect to address 164.76.4.207 and port 12489: No route to host
[1477661396] HOST ALERT: ADAUTHDC1;DOWN;SOFT;9;CRITICAL - Host Unreachable (164.76.4.207)
[1477661397] HOST ALERT: ADAUTHDC1-OLD;DOWN;SOFT;5;CRITICAL - Host Unreachable (164.76.4.206)
[1477661402] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (26.3% change >= 20.0% threshold)
[1477661403] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1477661403] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1477661403] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1477661403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1477661403] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1477661403] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1477661403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1477661403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1477661403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1477661403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1477661403] HOST DOWNTIME ALERT: VUM;STARTED; Host has entered a period of scheduled downtime
[1477661403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1477661403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1477661403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1477661403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1477661403] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1477661403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1477663195] HOST ALERT: dirtycow5.emich.edu;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1477663285] HOST ALERT: dirtycow5.emich.edu;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1477663375] HOST ALERT: dirtycow5.emich.edu;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1477663465] HOST ALERT: dirtycow5.emich.edu;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1477663555] HOST ALERT: dirtycow5.emich.edu;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1477663645] HOST ALERT: dirtycow5.emich.edu;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1477663735] HOST ALERT: dirtycow5.emich.edu;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1477663825] HOST ALERT: dirtycow5.emich.edu;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1477663915] HOST ALERT: dirtycow5.emich.edu;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1477664005] HOST ALERT: dirtycow5.emich.edu;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1477664026] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6360 found since last run. oneLevelSearchOps OK. searchOps is too high! 6414 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1477664146] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477664251] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1477664371] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1477664491] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477665161] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5373 found since last run. oneLevelSearchOps OK. searchOps is too high! 5664 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1477665281] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477665596] SERVICE ALERT: strongpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1477665691] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 22% (15 min average)
[1477665881] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6310 found since last run. oneLevelSearchOps OK. searchOps is too high! 6662 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1477666001] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477666601] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5114 found since last run. oneLevelSearchOps OK. searchOps is too high! 5525 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1477666617] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;OK;HARD;3;
[1477666721] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477666803] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1477666803] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1477666803] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1477666803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1477666803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1477666803] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1477666803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1477666803] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1477666803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1477666803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1477666803] HOST DOWNTIME ALERT: VUM;STARTED; Host has entered a period of scheduled downtime
[1477666803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1477666803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1477666803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1477666803] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1477666803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1477666803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1477667146] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6232 found since last run. oneLevelSearchOps OK. searchOps is too high! 6291 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1477667244] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1477667266] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477668521] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7229 found since last run. oneLevelSearchOps OK. searchOps is too high! 7610 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1477668641] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477669066] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5776 found since last run. oneLevelSearchOps OK. searchOps is too high! 5829 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1477669186] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477669291] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1477669411] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1477669531] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1477669786] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5109 found since last run. oneLevelSearchOps OK. searchOps is too high! 5174 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1477669906] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477670441] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5150 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1477670506] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6865 found since last run. oneLevelSearchOps OK. searchOps is too high! 6921 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1477670560] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477670626] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477670731] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1477670731] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1477671160] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8673 found since last run. oneLevelSearchOps OK. searchOps is too high! 8914 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1477671280] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477671880] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5952 found since last run. oneLevelSearchOps OK. searchOps is too high! 6191 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1477671931] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1477672000] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477672051] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1477672171] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 87% (15 min average)
[1477672203] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1477672203] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1477672203] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1477672203] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1477672203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1477672203] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1477672203] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1477672203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1477672203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1477672203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1477672203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1477672203] HOST DOWNTIME ALERT: VUM;STARTED; Host has entered a period of scheduled downtime
[1477672203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1477672203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1477672203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1477672203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1477672203] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1477672203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1477672426] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6699 found since last run. oneLevelSearchOps OK. searchOps is too high! 6754 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1477672546] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477673800] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5155 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1477673920] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477673971] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1477674091] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1477674211] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1477674346] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5145 found since last run. oneLevelSearchOps OK. searchOps is too high! 5338 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1477674466] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477675066] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8016 found since last run. oneLevelSearchOps OK. searchOps is too high! 8140 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1477675186] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477675786] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7312 found since last run. oneLevelSearchOps OK. searchOps is too high! 7438 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1477675906] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477676011] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 32% (15 min average)
[1477677435] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1477677446] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1477677495] SERVICE ALERT: hallefirst1;Device Status;OK;SOFT;2;Turn green knobs.
[1477677505] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! All Trays Still Full
[1477677706] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5275 found since last run. oneLevelSearchOps OK. searchOps is too high! 5479 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1477677811] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1477677826] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477677930] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1477677941] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1477678006] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1477678036] HOST ALERT: hallefirst1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1477678050] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1477678065] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1477678095] HOST ALERT: hallefirst1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1477678101] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1477678131] HOST ALERT: hallefirst1;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1477678161] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1477678191] HOST ALERT: hallefirst1;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1477678221] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1477678251] HOST ALERT: hallefirst1;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1477678341] HOST ALERT: hallefirst1;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1477678431] HOST ALERT: hallefirst1;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1477678521] HOST ALERT: hallefirst1;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1477678554] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1477678584] HOST ALERT: hallefirst1;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1477678614] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1477678644] HOST ALERT: hallefirst1;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1477678674] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1477678720] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6311 found since last run. oneLevelSearchOps OK. searchOps is too high! 6520 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1477678840] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477679251] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477679849] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1477680040] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5150 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1477680160] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477680760] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5965 found since last run. oneLevelSearchOps OK. searchOps is too high! 6281 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1477680880] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477681480] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7288 found since last run. oneLevelSearchOps OK. searchOps is too high! 7589 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1477681600] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477682200] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5451 found since last run. oneLevelSearchOps OK. searchOps is too high! 5818 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1477682249] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 47% (15 min average)
[1477682320] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477682626] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5328 found since last run. oneLevelSearchOps OK. searchOps is too high! 5441 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1477682746] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477683449] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1477683569] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 75% (15 min average)
[1477684168] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1477684286] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1477684407] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477684546] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 173 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477684666] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477684720] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 153 found since last run. wholeSubtreeSearchOps is too high! 6692 found since last run. oneLevelSearchOps OK. searchOps is too high! 6973 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1477684840] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477685006] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 83% (15 min average)
[1477685266] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5171 found since last run. oneLevelSearchOps OK. searchOps is too high! 5269 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1477685386] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477686479] SERVICE ALERT: hallegroundcolor;Image Transfer Kit HP CE249A;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Image Transfer Kit HP CE249A'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1477686532] SERVICE ALERT: hallegroundcolor;Image Transfer Kit HP CE249A;OK;SOFT;2;Image Transfer Kit HP CE249A is at 55% - OK!
[1477687405] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1477687524] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1477687645] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477688243] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 64% (15 min average)
[1477689443] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1477689563] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1477689631] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1477689691] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1477689751] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1477689756] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1477689816] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1477689876] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1477690049] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1477690109] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1477690163] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1477690169] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1477690283] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1477690325] SERVICE ALERT: idmrbs.emich.edu;eDirectory Service;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1477690335] SERVICE ALERT: hallefirst3;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1477690355] HOST ALERT: idmrbs.emich.edu;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1477690366] SERVICE ALERT: hallefirst3;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1477690384] SERVICE ALERT: idmrbs.emich.edu;Banner Driver;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 60 seconds.
[1477690395] SERVICE ALERT: hallefirst3;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1477690403] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 90% (15 min average)
[1477690414] HOST ALERT: idmrbs.emich.edu;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1477690426] SERVICE ALERT: hallefirst3;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1477690436] SERVICE ALERT: idmrbs.emich.edu;eDirectory Service;OK;SOFT;2;OK: ndsd status is good.
[1477690440] HOST ALERT: idmrbs.emich.edu;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1477690445] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1477690447] SERVICE ALERT: idmrbs.emich.edu;Banner Driver;CRITICAL;SOFT;2;Unknown: Error reported?
[1477690455] SERVICE ALERT: hallefirst3;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1477690486] SERVICE ALERT: hallefirst3;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1477690505] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1477690565] SERVICE ALERT: idmrbs.emich.edu;Banner Driver;CRITICAL;HARD;3;CRITICAL: cn=Ellucian Banner,cn=driverset1,o=system is stopped!
[1477690565] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1477690573] SERVICE ALERT: idmrbs.emich.edu;LDAP Stats;UNKNOWN;SOFT;1;Difference is negative. LDAP recently rebooted.
[1477690693] SERVICE ALERT: idmrbs.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477691003] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1477691123] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 92% (15 min average)
[1477691165] SERVICE ALERT: idmrbs.emich.edu;Banner Driver;OK;HARD;3;OK: cn=Ellucian Banner,cn=driverset1,o=system is running.
[1477691177] SERVICE ALERT: hallefirstcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1477691207] HOST ALERT: hallefirstcolor;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1477691222] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1477691236] SERVICE ALERT: hallefirstcolor;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1477691237] SERVICE ALERT: hallefirstcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1477691246] SERVICE ALERT: hallefirstcolor;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1477691252] HOST ALERT: hallefirstcolor;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1477691282] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1477691296] SERVICE ALERT: hallefirstcolor;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1477691297] SERVICE ALERT: hallefirstcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1477691306] SERVICE ALERT: hallefirstcolor;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1477691312] HOST ALERT: hallefirstcolor;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1477691343] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1477691356] SERVICE ALERT: hallefirstcolor;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1477691366] SERVICE ALERT: hallefirstcolor;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1477691373] HOST ALERT: hallefirstcolor;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1477691436] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1477691462] HOST ALERT: hallefirstcolor;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1477691496] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1477691526] HOST ALERT: hallefirstcolor;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1477691556] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1477691574] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1477691586] HOST ALERT: hallefirstcolor;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1477691596] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1477691626] HOST ALERT: hallefirstcolor;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1477691634] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1477691642] SERVICE ALERT: hallefirstcolor;Toner Collection Unit HP CE265A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1477691656] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1477691664] HOST ALERT: hallefirstcolor;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1477691672] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1477691694] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1477691702] SERVICE ALERT: hallefirstcolor;Toner Collection Unit HP CE265A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1477691702] HOST ALERT: hallefirstcolor;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1477691716] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1477691723] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1477691732] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1477691762] SERVICE ALERT: hallefirstcolor;Toner Collection Unit HP CE265A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1477691792] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1477691843] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 57% (15 min average)
[1477691892] SERVICE ALERT: hallefirstcolor;Image Fuser Kit HP 110V-CE246A;OK;HARD;3;Image Fuser Kit HP 110V-CE246A, 220V-CE247A is at 100% - OK!
[1477691896] HOST ALERT: hallefirstcolor;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1477691936] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;OK;HARD;3;Cyan Cartridge HP CE261A is at 100% - OK!
[1477691950] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1477691960] SERVICE ALERT: hallefirstcolor;Device Status;OK;HARD;3;LOAD TRAY 1 PLAIN LETTER 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1477692058] HOST ALERT: tiny;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.7)
[1477692122] HOST ALERT: tiny;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1477692151] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;OK;HARD;3;Magenta Cartridge HP CE263A is at 94% - OK!
[1477692184] HOST ALERT: caine;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1477692255] HOST ALERT: caine;UP;SOFT;2;PING OK - Packet loss = 58%, RTA = 2.95 ms
[1477692269] HOST ALERT: r25db;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1477692288] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;OK;HARD;3;Image Transfer Kit HP CE249A is at 72% - OK!
[1477692310] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;OK;HARD;3;Black Cartridge HP CE260X is at 78% - OK!
[1477692333] HOST ALERT: r25db;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.39 ms
[1477692341] HOST ALERT: appdb;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1477692356] SERVICE ALERT: hallefirstcolor;Toner Collection Unit HP CE265A;OK;HARD;3;
[1477692375] HOST ALERT: mydb;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1477692386] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;OK;HARD;3;Yellow Cartridge HP CE262A is at 36% - OK!
[1477692405] HOST ALERT: appdb;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1477692443] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1477692465] HOST ALERT: mydb;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1477692529] HOST ALERT: mydb;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1477692563] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 85% (15 min average)
[1477692894] HOST ALERT: boedb;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1477692921] SERVICE ALERT: idm1.emich.edu;OfficeAD Driver;CRITICAL;SOFT;1;Unknown: Error reported?
[1477692958] HOST ALERT: boedb;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1477693036] HOST DOWNTIME ALERT: idm1.emich.edu;STARTED; Host has entered a period of scheduled downtime
[1477693036] HOST DOWNTIME ALERT: idm2.emich.edu;STARTED; Host has entered a period of scheduled downtime
[1477693036] HOST DOWNTIME ALERT: idm3.emich.edu;STARTED; Host has entered a period of scheduled downtime
[1477693036] HOST DOWNTIME ALERT: idm4.emich.edu;STARTED; Host has entered a period of scheduled downtime
[1477693036] HOST DOWNTIME ALERT: idmrbs.emich.edu;STARTED; Host has entered a period of scheduled downtime
[1477693036] HOST DOWNTIME ALERT: idmrbstest.emich.edu;STARTED; Host has entered a period of scheduled downtime
[1477693036] HOST DOWNTIME ALERT: ldap.emich.edu;STARTED; Host has entered a period of scheduled downtime
[1477693086] SERVICE ALERT: idm1.emich.edu;OfficeAD Driver;CRITICAL;SOFT;2;Unknown: Error reported?
[1477693109] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 11
[1477693160] SERVICE ALERT: idm1.emich.edu;OfficeAD Driver;CRITICAL;HARD;3;Unknown: Error reported?
[1477693163] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1477693229] SERVICE ALERT: idm4.emich.edu;NSDS Errors;OK;SOFT;2;OK: Total errors: 0
[1477693237] SERVICE ALERT: idm1.emich.edu;AD Driver;CRITICAL;SOFT;1;CRITICAL: cn=Campus AD,cn=driverset1,o=system is stopped!
[1477693264] SERVICE ALERT: idm1.emich.edu;Oracle Driver;CRITICAL;SOFT;1;CRITICAL: cn=oradap,cn=driverset1,o=system is stopped!
[1477693283] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1477693290] SERVICE ALERT: idm1.emich.edu;LDAP Stats;UNKNOWN;SOFT;1;Difference is negative. LDAP recently rebooted.
[1477693307] SERVICE ALERT: idm1.emich.edu;Loopback Driver;CRITICAL;SOFT;1;CRITICAL: cn=Generic Loopback,cn=driverset1,o=system is stopped!
[1477693320] SERVICE ALERT: idm1.emich.edu;NSDS Errors;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 60 seconds.
[1477693358] SERVICE ALERT: idm1.emich.edu;AD Driver;CRITICAL;SOFT;2;CRITICAL: cn=Campus AD,cn=driverset1,o=system is stopped!
[1477693380] SERVICE ALERT: idm1.emich.edu;NSDS Errors;OK;SOFT;2;OK: 0 Errors Reported
[1477693385] SERVICE ALERT: idm1.emich.edu;Oracle Driver;OK;SOFT;2;OK: cn=oradap,cn=driverset1,o=system is running.
[1477693403] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 71% (15 min average)
[1477693406] HOST ALERT: authdb;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1477693410] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477693426] SERVICE ALERT: idm1.emich.edu;Loopback Driver;OK;SOFT;2;OK: cn=Generic Loopback,cn=driverset1,o=system is running.
[1477693432] SERVICE ALERT: idm2.emich.edu;eDirectory Service;CRITICAL;SOFT;1;CRITICAL: Failed to run the ndsd status command successfully!
[1477693478] SERVICE ALERT: idm1.emich.edu;AD Driver;OK;SOFT;3;OK: cn=Campus AD,cn=driverset1,o=system is running.
[1477693496] HOST ALERT: authdb;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1477693551] SERVICE ALERT: idm2.emich.edu;eDirectory Service;CRITICAL;SOFT;2;CRITICAL: Failed to run the ndsd status command successfully!
[1477693586] HOST ALERT: authdb;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1477693610] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 60 seconds.
[1477693650] HOST ALERT: authdb;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1477693671] SERVICE ALERT: idm2.emich.edu;eDirectory Service;CRITICAL;HARD;3;CRITICAL: Failed to run the ndsd status command successfully!
[1477693730] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 60 seconds.
[1477693761] SERVICE ALERT: idm1.emich.edu;OfficeAD Driver;OK;HARD;3;OK: cn=OfficeAD,cn=driverset1,o=system is running.
[1477693806] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;3;OK. All replicas within 120 seconds.
[1477693829] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 13
[1477693863] SERVICE ALERT: idm2.emich.edu;LDAP Stats;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1477693871] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 13
[1477693949] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 1
[1477693980] SERVICE ALERT: idm1.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 1
[1477693983] SERVICE ALERT: idm2.emich.edu;LDAP Stats;UNKNOWN;SOFT;2;Difference is negative. LDAP recently rebooted.
[1477693991] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 1
[1477694003] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1477694004] SERVICE ALERT: idm2.emich.edu;NSDS Errors;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 60 seconds.
[1477694069] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 1
[1477694100] SERVICE ALERT: idm1.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 1
[1477694103] SERVICE ALERT: idm2.emich.edu;LDAP Stats;OK;SOFT;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477694111] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 1
[1477694114] SERVICE ALERT: idm2.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 1
[1477694123] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1477694183] SERVICE ALERT: idm2.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 1
[1477694220] SERVICE ALERT: idm1.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 1
[1477694243] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477694270] SERVICE ALERT: idm2.emich.edu;eDirectory Service;OK;HARD;3;OK: ndsd status is good.
[1477694313] HOST ALERT: liteapp;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1477694377] HOST ALERT: liteapp;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.41 ms
[1477694732] SERVICE ALERT: idm3.emich.edu;NSDS Errors;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1477694783] SERVICE ALERT: idm2.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1477694820] SERVICE ALERT: idm1.emich.edu;NSDS Errors;OK;HARD;3;OK: 0 Errors Reported
[1477694986] SERVICE ALERT: idm3.emich.edu;LDAP Stats;UNKNOWN;SOFT;1;Difference is negative. LDAP recently rebooted.
[1477695106] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477695247] SERVICE ALERT: ldap.emich.edu;LDAP Stats;CRITICAL;SOFT;1;18:54:07  - (FAILURE!) LDAP Failed
[1477695269] SERVICE ALERT: idm4.emich.edu;NSDS Errors;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1477695301] SERVICE ALERT: ldap.emich.edu;LDAP Stats;OK;SOFT;2;18:55:01  - Destination LDAP online and available.
[1477695311] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 2
[1477695443] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1477695640] SERVICE ALERT: idm4.emich.edu;LDAP Stats;UNKNOWN;SOFT;1;Difference is negative. LDAP recently rebooted.
[1477695760] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477695870] SERVICE ALERT: idm4.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1477695911] SERVICE ALERT: idm3.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1477695911] SERVICE FLAPPING ALERT: idm3.emich.edu;NSDS Errors;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1477696043] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1477696163] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 90% (15 min average)
[1477697963] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1477698083] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1477698203] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 82% (15 min average)
[1477698803] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1477698923] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1477699043] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477699643] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 72% (15 min average)
[1477699906] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5120 found since last run. oneLevelSearchOps OK. searchOps is too high! 5380 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1477700026] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477700843] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1477700963] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1477701083] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1477702992] SERVICE ALERT: ldap.emich.edu;LDAP Stats;CRITICAL;SOFT;1;21:03:12  - (FAILURE!) LDAP Failed
[1477702995] HOST ALERT: ldap.emich.edu;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.194)
[1477703043] SERVICE ALERT: ldap.emich.edu;LDAP Stats;CRITICAL;HARD;2;21:04:03  - (FAILURE!) LDAP Failed
[1477703046] HOST ALERT: ldap.emich.edu;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.4.194)
[1477703109] HOST ALERT: ldap.emich.edu;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.4.194)
[1477703166] HOST ALERT: ldap.emich.edu;DOWN;SOFT;4;CRITICAL - Host Unreachable (164.76.4.194)
[1477703229] HOST ALERT: ldap.emich.edu;DOWN;SOFT;5;CRITICAL - Host Unreachable (164.76.4.194)
[1477703285] HOST ALERT: ldap.emich.edu;UP;SOFT;6;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1477703829] HOST DOWNTIME ALERT: idmrbs.emich.edu;STOPPED; Host has exited from a period of scheduled downtime
[1477703829] HOST DOWNTIME ALERT: idm1.emich.edu;STOPPED; Host has exited from a period of scheduled downtime
[1477703829] HOST DOWNTIME ALERT: idmrbstest.emich.edu;STOPPED; Host has exited from a period of scheduled downtime
[1477703829] HOST DOWNTIME ALERT: ldap.emich.edu;STOPPED; Host has exited from a period of scheduled downtime
[1477703829] HOST DOWNTIME ALERT: idm4.emich.edu;STOPPED; Host has exited from a period of scheduled downtime
[1477703829] HOST DOWNTIME ALERT: idm3.emich.edu;STOPPED; Host has exited from a period of scheduled downtime
[1477703829] HOST DOWNTIME ALERT: idm2.emich.edu;STOPPED; Host has exited from a period of scheduled downtime
[1477704160] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 116 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477704280] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1477704360] SERVICE ALERT: ldap.emich.edu;LDAP Stats;OK;HARD;2;21:26:00  - Destination LDAP online and available.
[1477706483] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 55% (15 min average)
[1477707083] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1477707203] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1477707311] SERVICE FLAPPING ALERT: idm3.emich.edu;NSDS Errors;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1477707323] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1477707923] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477708523] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1477710323] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477710923] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1477711523] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1477712123] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)