[1454755318] SERVICE ALERT: ADWU1;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1454755438] SERVICE ALERT: ADWU1;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1454776039] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1454776099] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1454776159] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1454776381] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1454776441] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1454776501] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1454780359] SERVICE ALERT: studentcenterpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1454780697] SERVICE ALERT: studentcenterpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1454788862] SERVICE ALERT: studentcenterpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1454788922] SERVICE ALERT: studentcenterpk;Device Status;OK;SOFT;2;Move left accessory, if attached.
[1454789097] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1454789158] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1454789218] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1454789522] SERVICE ALERT: studentcenterpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1454789582] SERVICE ALERT: studentcenterpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1454789642] SERVICE ALERT: studentcenterpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1454789664] SERVICE ALERT: studentcenterpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1454789724] SERVICE ALERT: studentcenterpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1454789785] SERVICE ALERT: studentcenterpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1454789817] SERVICE ALERT: studentcenterpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1454790842] SERVICE ALERT: studentcenterpk;Device Status;OK;HARD;3;Push down on green tab.
[1454790984] SERVICE ALERT: studentcenterpk;Device Status;OK;HARD;3;REINSERT DUPLEXER
[1454793757] HOST ALERT: studentcenterpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1454793798] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1454793803] HOST ALERT: studentcenterpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.68 ms
[1454793858] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 37% - OK!
[1454796770] SERVICE ALERT: DB01-2012;CPU Load;WARNING;SOFT;1;CPU Load 82% (5 min average)
[1454796890] SERVICE ALERT: DB01-2012;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1454810690] SERVICE ALERT: DB01-2012;CPU Load;WARNING;SOFT;1;CPU Load 82% (5 min average)
[1454810810] SERVICE ALERT: DB01-2012;CPU Load;OK;SOFT;2;CPU Load 72% (5 min average)
[1454811410] SERVICE ALERT: DB01-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 91% (5 min average)
[1454811530] SERVICE ALERT: DB01-2012;CPU Load;WARNING;SOFT;2;CPU Load 89% (5 min average)
[1454811546] SERVICE ALERT: SANDBOX-DB;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1454811650] SERVICE ALERT: DB01-2012;CPU Load;WARNING;HARD;3;CPU Load 84% (5 min average)
[1454811666] SERVICE ALERT: SANDBOX-DB;CPU Load;OK;SOFT;2;CPU Load 81% (5 min average)
[1454812250] SERVICE ALERT: DB01-2012;CPU Load;OK;HARD;3;CPU Load 61% (5 min average)
[1454814339] SERVICE ALERT: ADWSUSSERVERS;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1454814439] SERVICE ALERT: ADWSUSSERVERS;Service: DNS Client;OK;SOFT;2;Dnscache: Started