[1462766471] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1462766591] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1462766711] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462769111] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462769711] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462770911] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 83% (15 min average)
[1462771511] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462771631] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1462771751] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1462772351] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 74% (15 min average)
[1462772951] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462773071] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462773191] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1462773785] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1462773905] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1462774025] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 74% (15 min average)
[1462774991] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462775591] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462777624] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462777744] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 85% (15 min average)
[1462777991] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462778591] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462782191] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 74% (15 min average)
[1462786990] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1462787110] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1462787230] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 46% (15 min average)
[1462787830] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1462787950] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 63% (15 min average)
[1462788550] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462788670] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1462788790] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462789390] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462789990] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 62% (15 min average)
[1462790590] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462790710] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462790830] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462791544] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1462791664] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 89% (15 min average)
[1462793230] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 84% (15 min average)
[1462794430] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462794550] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1462794670] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1462797670] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 74% (15 min average)
[1462799611] HOST ALERT: portal;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.51)
[1462799675] HOST ALERT: portal;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.78 ms
[1462799803] SERVICE ALERT: halleground1;Device Status;OK;HARD;3;Pull then lift green tab.
[1462800070] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1462800190] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 47% (15 min average)
[1462800403] SERVICE ALERT: halleground1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462800464] SERVICE ALERT: halleground1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462800523] SERVICE ALERT: halleground1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462800790] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462800910] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462801030] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462801864] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462801984] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1462802104] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 93% (15 min average)
[1462802408] SERVICE ALERT: BOEAPP;Memory Usage;WARNING;SOFT;1;Memory usage: total:18429.84 MB - used: 16837.27 MB (91%) - free: 1592.57 MB (9%)
[1462802528] SERVICE ALERT: BOEAPP;Memory Usage;WARNING;SOFT;2;Memory usage: total:18874.42 MB - used: 17395.70 MB (92%) - free: 1478.72 MB (8%)
[1462802648] SERVICE ALERT: BOEAPP;Memory Usage;OK;SOFT;3;Memory usage: total:18890.42 MB - used: 16864.84 MB (89%) - free: 2025.59 MB (11%)
[1462802704] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462802824] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462802830] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 43% (15 min average)
[1462802944] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1462803544] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462803818] HOST ALERT: c137.emich.edu;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462803908] HOST ALERT: c137.emich.edu;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1462803998] HOST ALERT: c137.emich.edu;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1462804030] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462804042] SERVICE ALERT: c137.emich.edu;Service: HAProxy Status;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1462804072] HOST ALERT: c137.emich.edu;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1462804150] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 81% (15 min average)
[1462804162] SERVICE ALERT: c137.emich.edu;Service: HAProxy Status;CRITICAL;HARD;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1462804162] HOST ALERT: c137.emich.edu;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1462804251] HOST ALERT: c137.emich.edu;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1462804341] HOST ALERT: c137.emich.edu;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1462804431] HOST ALERT: c137.emich.edu;DOWN;SOFT;8;(Host check timed out after 30.00 seconds)
[1462804448] SERVICE ALERT: BOEAPP;Memory Usage;WARNING;SOFT;1;Memory usage: total:18890.42 MB - used: 17012.67 MB (90%) - free: 1877.75 MB (10%)
[1462804492] HOST ALERT: c137.emich.edu;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1462804568] SERVICE ALERT: BOEAPP;Memory Usage;WARNING;SOFT;2;Memory usage: total:18890.42 MB - used: 17023.97 MB (90%) - free: 1866.45 MB (10%)
[1462804582] HOST ALERT: c137.emich.edu;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1462804688] SERVICE ALERT: BOEAPP;Memory Usage;OK;SOFT;3;Memory usage: total:18890.42 MB - used: 16956.99 MB (90%) - free: 1933.43 MB (10%)
[1462804750] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462804870] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462804990] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462805073] SERVICE ALERT: halleground2;Device Status;OK;HARD;3;Ready
[1462805344] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462808590] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1462809190] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1462809310] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 82% (15 min average)
[1462809910] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462810030] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462810144] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462810144] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1462810150] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 94% (15 min average)
[1462810744] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1462810750] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462810870] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462810990] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462811344] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 74% (15 min average)
[1462811923] SERVICE ALERT: halleground1;Device Status;OK;HARD;3;MENUS  RETRIEVE JOB  INFORMATION  PAPER HANDLING
[1462812190] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 42% (15 min average)
[1462812239] SERVICE ALERT: markjeffersonpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1462812293] SERVICE ALERT: markjeffersonpk;Device Status;OK;SOFT;2;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1462812524] SERVICE ALERT: halleground1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462812583] SERVICE ALERT: halleground1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462812643] SERVICE ALERT: halleground1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462812790] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462812910] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1462813030] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462813243] SERVICE ALERT: halleground1;Device Status;OK;HARD;3;59.20 ERROR  To continue turn off then on 59.20 ERROR  46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1462813384] HOST DOWNTIME ALERT: DIRECTACCESS;STARTED; Host has entered a period of scheduled downtime
[1462813648] HOST ALERT: DIRECTACCESS;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.41)
[1462813711] HOST ALERT: DIRECTACCESS;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.41)
[1462813774] HOST ALERT: DIRECTACCESS;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.102.41)
[1462813837] HOST ALERT: DIRECTACCESS;DOWN;SOFT;4;CRITICAL - Host Unreachable (164.76.102.41)
[1462813900] HOST ALERT: DIRECTACCESS;DOWN;SOFT;5;CRITICAL - Host Unreachable (164.76.102.41)
[1462813963] HOST ALERT: DIRECTACCESS;DOWN;SOFT;6;CRITICAL - Host Unreachable (164.76.102.41)
[1462814026] HOST ALERT: DIRECTACCESS;DOWN;SOFT;7;CRITICAL - Host Unreachable (164.76.102.41)
[1462814089] HOST ALERT: DIRECTACCESS;DOWN;SOFT;8;CRITICAL - Host Unreachable (164.76.102.41)
[1462814152] HOST ALERT: DIRECTACCESS;DOWN;SOFT;9;CRITICAL - Host Unreachable (164.76.102.41)
[1462814215] HOST ALERT: DIRECTACCESS;DOWN;HARD;10;CRITICAL - Host Unreachable (164.76.102.41)
[1462814414] SERVICE ALERT: ADDC5;GPO Replication;WARNING;SOFT;1;Warning! 2 policies are not replicating properly!
[1462814710] SERVICE ALERT: ADDC5;GPO Replication;WARNING;SOFT;2;Warning! 2 policies are not replicating properly!
[1462815011] SERVICE ALERT: ADDC5;GPO Replication;WARNING;HARD;3;Warning! 2 policies are not replicating properly!
[1462815855] HOST ALERT: c137.emich.edu;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1462815950] HOST DOWNTIME ALERT: FITNESSGRAM;STARTED; Host has entered a period of scheduled downtime
[1462816019] SERVICE ALERT: FITNESSGRAM;Service: User Profile;CRITICAL;SOFT;1;connect to address 164.76.102.19 and port 12489: No route to host
[1462816022] HOST ALERT: FITNESSGRAM;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.19)
[1462816030] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1462816056] SERVICE ALERT: FITNESSGRAM;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.102.19 and port 12489: No route to host
[1462816059] HOST ALERT: FITNESSGRAM;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.19)
[1462816096] SERVICE ALERT: FITNESSGRAM;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.102.19 and port 12489: No route to host
[1462816099] HOST ALERT: FITNESSGRAM;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.102.19)
[1462816130] SERVICE ALERT: FITNESSGRAM;Service: User Profile;CRITICAL;SOFT;2;connect to address 164.76.102.19 and port 12489: No route to host
[1462816133] HOST ALERT: FITNESSGRAM;DOWN;SOFT;4;CRITICAL - Host Unreachable (164.76.102.19)
[1462816176] SERVICE ALERT: FITNESSGRAM;Service: WMI;CRITICAL;SOFT;2;connect to address 164.76.102.19 and port 12489: No route to host
[1462816179] HOST ALERT: FITNESSGRAM;DOWN;SOFT;5;CRITICAL - Host Unreachable (164.76.102.19)
[1462816216] SERVICE ALERT: FITNESSGRAM;Service: Windows Time;CRITICAL;SOFT;2;connect to address 164.76.102.19 and port 12489: No route to host
[1462816219] HOST ALERT: FITNESSGRAM;DOWN;SOFT;6;CRITICAL - Host Unreachable (164.76.102.19)
[1462816250] SERVICE ALERT: FITNESSGRAM;Service: User Profile;CRITICAL;HARD;3;connect to address 164.76.102.19 and port 12489: No route to host
[1462816253] HOST ALERT: FITNESSGRAM;DOWN;SOFT;7;CRITICAL - Host Unreachable (164.76.102.19)
[1462816266] SERVICE ALERT: FITNESSGRAM;VMWare Tools;CRITICAL;SOFT;1;connect to address 164.76.102.19 and port 12489: No route to host
[1462816269] HOST ALERT: FITNESSGRAM;DOWN;SOFT;8;CRITICAL - Host Unreachable (164.76.102.19)
[1462816281] SERVICE ALERT: FITNESSGRAM;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.102.19 and port 12489: No route to host
[1462816284] HOST ALERT: FITNESSGRAM;DOWN;SOFT;9;CRITICAL - Host Unreachable (164.76.102.19)
[1462816296] SERVICE ALERT: FITNESSGRAM;Service: WMI;CRITICAL;HARD;3;connect to address 164.76.102.19 and port 12489: No route to host
[1462816299] HOST ALERT: FITNESSGRAM;DOWN;HARD;10;CRITICAL - Host Unreachable (164.76.102.19)
[1462816336] SERVICE ALERT: FITNESSGRAM;Service: Windows Time;CRITICAL;HARD;3;connect to address 164.76.102.19 and port 12489: No route to host
[1462816342] SERVICE ALERT: FITNESSGRAM;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.102.19 and port 12489: No route to host
[1462816357] SERVICE ALERT: FITNESSGRAM;CPU Load;CRITICAL;SOFT;1;connect to address 164.76.102.19 and port 12489: No route to host
[1462816386] SERVICE ALERT: FITNESSGRAM;VMWare Tools;CRITICAL;SOFT;2;connect to address 164.76.102.19 and port 12489: No route to host
[1462816396] SERVICE ALERT: FITNESSGRAM;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.102.19 and port 12489: No route to host
[1462816399] SERVICE ALERT: FITNESSGRAM;Memory Usage;CRITICAL;SOFT;2;connect to address 164.76.102.19 and port 12489: No route to host
[1462816421] SERVICE ALERT: FITNESSGRAM;Uptime;CRITICAL;SOFT;1;connect to address 164.76.102.19 and port 12489: No route to host
[1462816462] SERVICE ALERT: FITNESSGRAM;Service: DNS Client;CRITICAL;SOFT;2;connect to address 164.76.102.19 and port 12489: No route to host
[1462816462] SERVICE ALERT: FITNESSGRAM;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.102.19 and port 12489: No route to host
[1462816477] SERVICE ALERT: FITNESSGRAM;CPU Load;CRITICAL;SOFT;2;connect to address 164.76.102.19 and port 12489: No route to host
[1462816506] SERVICE ALERT: FITNESSGRAM;VMWare Tools;CRITICAL;HARD;3;connect to address 164.76.102.19 and port 12489: No route to host
[1462816516] SERVICE ALERT: FITNESSGRAM;Service: Firewall;CRITICAL;SOFT;2;connect to address 164.76.102.19 and port 12489: No route to host
[1462816519] SERVICE ALERT: FITNESSGRAM;Memory Usage;CRITICAL;HARD;3;connect to address 164.76.102.19 and port 12489: No route to host
[1462816526] SERVICE ALERT: FITNESSGRAM;Service: NetLogin;CRITICAL;SOFT;1;connect to address 164.76.102.19 and port 12489: No route to host
[1462816526] SERVICE ALERT: FITNESSGRAM;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.102.19 and port 12489: No route to host
[1462816541] SERVICE ALERT: FITNESSGRAM;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.102.19 and port 12489: No route to host
[1462816541] SERVICE ALERT: FITNESSGRAM;Uptime;CRITICAL;SOFT;2;connect to address 164.76.102.19 and port 12489: No route to host
[1462816548] SERVICE ALERT: FITNESSGRAM;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.102.19 and port 12489: No route to host
[1462816582] SERVICE ALERT: FITNESSGRAM;Service: DNS Client;CRITICAL;HARD;3;connect to address 164.76.102.19 and port 12489: No route to host
[1462816582] SERVICE ALERT: FITNESSGRAM;Service: Group Policy;CRITICAL;SOFT;2;connect to address 164.76.102.19 and port 12489: No route to host
[1462816597] SERVICE ALERT: FITNESSGRAM;CPU Load;CRITICAL;HARD;3;connect to address 164.76.102.19 and port 12489: No route to host
[1462816630] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1462816636] SERVICE ALERT: FITNESSGRAM;Service: Firewall;CRITICAL;HARD;3;connect to address 164.76.102.19 and port 12489: No route to host
[1462816646] SERVICE ALERT: FITNESSGRAM;NSClient++ Version;CRITICAL;SOFT;2;connect to address 164.76.102.19 and port 12489: No route to host
[1462816646] SERVICE ALERT: FITNESSGRAM;Service: NetLogin;CRITICAL;SOFT;2;connect to address 164.76.102.19 and port 12489: No route to host
[1462816661] SERVICE ALERT: FITNESSGRAM;Uptime;CRITICAL;HARD;3;connect to address 164.76.102.19 and port 12489: No route to host
[1462816661] SERVICE ALERT: FITNESSGRAM;Service: RPC;CRITICAL;SOFT;2;connect to address 164.76.102.19 and port 12489: No route to host
[1462816667] SERVICE ALERT: FITNESSGRAM;C:\ Drive Space;CRITICAL;SOFT;2;connect to address 164.76.102.19 and port 12489: No route to host
[1462816702] SERVICE ALERT: FITNESSGRAM;Service: Group Policy;CRITICAL;HARD;3;connect to address 164.76.102.19 and port 12489: No route to host
[1462816744] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462816750] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462816766] SERVICE ALERT: FITNESSGRAM;NSClient++ Version;CRITICAL;HARD;3;connect to address 164.76.102.19 and port 12489: No route to host
[1462816775] SERVICE ALERT: FITNESSGRAM;Service: NetLogin;CRITICAL;HARD;3;connect to address 164.76.102.19 and port 12489: No route to host
[1462816782] SERVICE ALERT: FITNESSGRAM;Service: RPC;CRITICAL;HARD;3;connect to address 164.76.102.19 and port 12489: No route to host
[1462816788] SERVICE ALERT: FITNESSGRAM;C:\ Drive Space;CRITICAL;HARD;3;connect to address 164.76.102.19 and port 12489: No route to host
[1462816864] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1462816870] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1462816984] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462817433] SERVICE ALERT: CEDARON;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462817457] SERVICE ALERT: CEDARON;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462817463] HOST ALERT: CEDARON;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462817470] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 81% (15 min average)
[1462817553] SERVICE ALERT: CEDARON;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462817553] HOST ALERT: CEDARON;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1462817575] SERVICE ALERT: CEDARON;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462817577] SERVICE ALERT: CEDARON;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462817584] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 91% (15 min average)
[1462817605] HOST ALERT: CEDARON;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1462817613] SERVICE ALERT: CEDARON;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462817613] SERVICE ALERT: CEDARON;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462817623] SERVICE ALERT: CEDARON;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462817643] HOST ALERT: CEDARON;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1462817673] SERVICE ALERT: CEDARON;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462817677] SERVICE ALERT: CEDARON;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462817694] SERVICE ALERT: CEDARON;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462817697] SERVICE ALERT: CEDARON;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462817703] HOST ALERT: CEDARON;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1462817733] SERVICE ALERT: CEDARON;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462817733] SERVICE ALERT: CEDARON;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462817738] SERVICE ALERT: CEDARON;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462817743] SERVICE ALERT: CEDARON;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462817759] SERVICE ALERT: CEDARON;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462817763] HOST ALERT: CEDARON;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1462817797] SERVICE ALERT: CEDARON;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462817803] SERVICE ALERT: CEDARON;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462817813] SERVICE ALERT: CEDARON;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462817827] HOST ALERT: CEDARON;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1462817829] SERVICE ALERT: CEDARON;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462817853] SERVICE ALERT: CEDARON;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462817853] SERVICE ALERT: CEDARON;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462817858] SERVICE ALERT: CEDARON;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462817859] HOST ALERT: CEDARON;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1462817863] SERVICE ALERT: CEDARON;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462817873] SERVICE ALERT: CEDARON;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462817873] SERVICE ALERT: CEDARON;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462817878] SERVICE ALERT: CEDARON;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462817893] HOST ALERT: CEDARON;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1462817906] SERVICE ALERT: CEDARON;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462817917] SERVICE ALERT: CEDARON;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462817923] SERVICE ALERT: CEDARON;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462817936] HOST ALERT: CEDARON;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1462817949] SERVICE ALERT: CEDARON;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462817978] SERVICE ALERT: CEDARON;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462817993] SERVICE ALERT: CEDARON;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462817993] SERVICE ALERT: CEDARON;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462817998] SERVICE ALERT: CEDARON;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462818026] SERVICE ALERT: CEDARON;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462818043] SERVICE ALERT: CEDARON;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462818069] SERVICE ALERT: CEDARON;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462818113] SERVICE ALERT: CEDARON;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462818113] SERVICE ALERT: CEDARON;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462818146] SERVICE ALERT: CEDARON;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462818670] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462818790] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462818910] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 94% (15 min average)
[1462819724] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;CRITICAL;SOFT;1;CRITICAL: cn=ADAUTH Driver,cn=driverset1,o=system is stopped!
[1462819843] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;CRITICAL;SOFT;2;CRITICAL: cn=ADAUTH Driver,cn=driverset1,o=system is stopped!
[1462819963] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;CRITICAL;HARD;3;CRITICAL: cn=ADAUTH Driver,cn=driverset1,o=system is stopped!
[1462819977] SERVICE ALERT: ADAUTHDC3;AD Replication;WARNING;SOFT;1;No output available from command (addcreplicationcheck).
[1462820023] SERVICE ALERT: ADAUTHDC3;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462820099] SERVICE ALERT: ADAUTHDC3;AD Replication;OK;SOFT;2;OK: replication is up and running
[1462820110] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462820123] SERVICE ALERT: ADAUTHDC3;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1462820123] SERVICE ALERT: ADAUTHDC3;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1462820230] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 71% (15 min average)
[1462820243] SERVICE ALERT: ADAUTHDC3;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1462820272] HOST ALERT: halleground1;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1462820336] HOST ALERT: halleground1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.55 ms
[1462820404] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1462820404] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (34.7% change >= 20.0% threshold)
[1462820404] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1462820404] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1462820404] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1462820404] HOST DOWNTIME ALERT: R25-BETA2012;STARTED; Host has entered a period of scheduled downtime
[1462820404] HOST DOWNTIME ALERT: DIRECTACCESS;STARTED; Host has entered a period of scheduled downtime
[1462820404] SERVICE DOWNTIME ALERT: NETIDADPROXY;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1462820404] HOST DOWNTIME ALERT: FITNESSGRAM;STARTED; Host has entered a period of scheduled downtime
[1462820404] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1462820404] HOST DOWNTIME ALERT: R25-GAMMA2012;STARTED; Host has entered a period of scheduled downtime
[1462820404] HOST DOWNTIME ALERT: R25-EPSILON2012;STARTED; Host has entered a period of scheduled downtime
[1462820423] SERVICE ALERT: CEDARON;Service: RPC;OK;HARD;3;RpcSs: Started
[1462820425] SERVICE ALERT: ADAUTHDC3;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.102.173 and port 12489: No route to host
[1462820427] HOST ALERT: CEDARON;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1462820428] SERVICE ALERT: ADAUTHDC3;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.102.173 and port 12489: No route to host
[1462820428] HOST ALERT: ADAUTHDC3;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.173)
[1462820445] SERVICE ALERT: ADAUTHDC3;Service: Dell Backup Agent;CRITICAL;SOFT;1;connect to address 164.76.102.173 and port 12489: No route to host
[1462820447] SERVICE ALERT: FITNESSGRAM;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1462820448] HOST ALERT: ADAUTHDC3;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.173)
[1462820449] SERVICE ALERT: CEDARON;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1462820449] SERVICE ALERT: halleground1;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1462820451] HOST ALERT: FITNESSGRAM;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1462820472] SERVICE ALERT: ADAUTHDC3;Service: AD Domain Service;CRITICAL;SOFT;1;connect to address 164.76.102.173 and port 12489: No route to host
[1462820475] HOST ALERT: ADAUTHDC3;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.102.173)
[1462820479] HOST ALERT: halleground1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462820481] SERVICE ALERT: ADAUTHDC3;Service: Remote Management;CRITICAL;SOFT;1;connect to address 164.76.102.173 and port 12489: No route to host
[1462820493] SERVICE ALERT: FITNESSGRAM;Service: WMI;OK;HARD;3;Winmgmt: Started
[1462820493] SERVICE ALERT: CEDARON;Service: Windows Time;OK;HARD;3;W32Time: Started
[1462820493] SERVICE ALERT: CEDARON;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1462820496] SERVICE ALERT: ADAUTHDC3;Service: User Profile;CRITICAL;SOFT;1;connect to address 164.76.102.173 and port 12489: No route to host
[1462820496] SERVICE ALERT: ADAUTHDC3;Service: AD Web Service;CRITICAL;SOFT;1;connect to address 164.76.102.173 and port 12489: No route to host
[1462820499] HOST ALERT: ADAUTHDC3;DOWN;SOFT;4;CRITICAL - Host Unreachable (164.76.102.173)
[1462820509] SERVICE ALERT: halleground1;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1462820526] SERVICE ALERT: CEDARON;Service: WMI;OK;HARD;3;Winmgmt: Started
[1462820533] SERVICE ALERT: FITNESSGRAM;Service: Windows Time;OK;HARD;3;W32Time: Started
[1462820536] SERVICE ALERT: ADAUTHDC3;Service: RPC;CRITICAL;SOFT;2;connect to address 164.76.102.173 and port 12489: No route to host
[1462820539] SERVICE ALERT: ADAUTHDC3;Service: DFS Namespace;CRITICAL;SOFT;1;connect to address 164.76.102.173 and port 12489: No route to host
[1462820539] HOST ALERT: ADAUTHDC3;DOWN;SOFT;5;CRITICAL - Host Unreachable (164.76.102.173)
[1462820539] HOST ALERT: halleground1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1462820543] SERVICE ALERT: ADAUTHDC3;NSClient++ Version;CRITICAL;SOFT;2;connect to address 164.76.102.173 and port 12489: No route to host
[1462820557] SERVICE ALERT: ADAUTHDC3;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.102.173 and port 12489: No route to host
[1462820560] HOST ALERT: ADAUTHDC3;DOWN;SOFT;6;CRITICAL - Host Unreachable (164.76.102.173)
[1462820565] SERVICE ALERT: ADAUTHDC3;Service: Dell Backup Agent;CRITICAL;SOFT;2;connect to address 164.76.102.173 and port 12489: No route to host
[1462820569] SERVICE ALERT: halleground1;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1462820578] HOST ALERT: DIRECTACCESS;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.53 ms
[1462820583] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462820592] SERVICE ALERT: ADAUTHDC3;Service: AD Domain Service;CRITICAL;SOFT;2;connect to address 164.76.102.173 and port 12489: No route to host
[1462820595] HOST ALERT: ADAUTHDC3;DOWN;SOFT;7;CRITICAL - Host Unreachable (164.76.102.173)
[1462820599] HOST ALERT: halleground1;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1462820601] SERVICE ALERT: ADAUTHDC3;Service: Remote Management;CRITICAL;SOFT;2;connect to address 164.76.102.173 and port 12489: No route to host
[1462820616] SERVICE ALERT: ADAUTHDC3;Service: User Profile;CRITICAL;SOFT;2;connect to address 164.76.102.173 and port 12489: No route to host
[1462820616] SERVICE ALERT: ADAUTHDC3;Service: AD Web Service;CRITICAL;SOFT;2;connect to address 164.76.102.173 and port 12489: No route to host
[1462820619] HOST ALERT: ADAUTHDC3;DOWN;SOFT;8;CRITICAL - Host Unreachable (164.76.102.173)
[1462820629] SERVICE ALERT: halleground1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1462820653] SERVICE ALERT: CEDARON;C:\ Drive Space;OK;HARD;3;c:\ - total: 249.66 Gb - used: 34.81 Gb (14%) - free 214.85 Gb (86%)
[1462820656] SERVICE ALERT: ADAUTHDC3;Service: RPC;CRITICAL;HARD;3;connect to address 164.76.102.173 and port 12489: No route to host
[1462820659] SERVICE ALERT: ADAUTHDC3;Service: DFS Namespace;CRITICAL;SOFT;2;connect to address 164.76.102.173 and port 12489: No route to host
[1462820659] HOST ALERT: ADAUTHDC3;DOWN;SOFT;9;CRITICAL - Host Unreachable (164.76.102.173)
[1462820659] HOST ALERT: halleground1;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1462820662] SERVICE ALERT: ADAUTHDC3;NSClient++ Version;CRITICAL;HARD;3;connect to address 164.76.102.173 and port 12489: No route to host
[1462820662] SERVICE ALERT: ADAUTHDC3;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.102.173 and port 12489: No route to host
[1462820677] SERVICE ALERT: ADAUTHDC3;Service: WMI;CRITICAL;SOFT;2;connect to address 164.76.102.173 and port 12489: No route to host
[1462820677] SERVICE ALERT: CEDARON;VMWare Tools;OK;HARD;3;VMTools: Started
[1462820680] HOST ALERT: ADAUTHDC3;DOWN;HARD;10;CRITICAL - Host Unreachable (164.76.102.173)
[1462820685] SERVICE ALERT: ADAUTHDC3;Service: Dell Backup Agent;CRITICAL;HARD;3;connect to address 164.76.102.173 and port 12489: No route to host
[1462820685] SERVICE ALERT: ADAUTHDC3;VMWare Tools;CRITICAL;SOFT;1;connect to address 164.76.102.173 and port 12489: No route to host
[1462820689] SERVICE ALERT: halleground1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1462820693] SERVICE ALERT: ADAUTHDC3;AD Replication;WARNING;SOFT;1;(No output on stdout) stderr: Could not resolve hostname ADAUTHDC3: Name or service not known
[1462820703] SERVICE ALERT: FITNESSGRAM;VMWare Tools;OK;HARD;3;VMTools: Started
[1462820703] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 77% (15 min average)
[1462820706] SERVICE ALERT: ADAUTHDC3;Service: NetLogin;CRITICAL;SOFT;1;connect to address 164.76.102.173 and port 12489: No route to host
[1462820709] SERVICE ALERT: ADAUTHDC3;Service: AD Domain Service;CRITICAL;HARD;3;connect to address 164.76.102.173 and port 12489: No route to host
[1462820712] SERVICE ALERT: ADAUTHDC3;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.102.173 and port 12489: No route to host
[1462820717] SERVICE ALERT: FITNESSGRAM;Memory Usage;OK;HARD;3;Memory usage: total:10237.85 MB - used: 1623.58 MB (16%) - free: 8614.27 MB (84%)
[1462820719] HOST ALERT: halleground1;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1462820721] SERVICE ALERT: ADAUTHDC3;Service: Remote Management;CRITICAL;HARD;3;connect to address 164.76.102.173 and port 12489: No route to host
[1462820724] SERVICE ALERT: ADAUTHDC3;Service: Distributed File System Replication;CRITICAL;SOFT;1;connect to address 164.76.102.173 and port 12489: No route to host
[1462820724] SERVICE ALERT: ADAUTHDC3;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.102.173 and port 12489: No route to host
[1462820736] SERVICE ALERT: ADAUTHDC3;Service: User Profile;CRITICAL;HARD;3;connect to address 164.76.102.173 and port 12489: No route to host
[1462820736] SERVICE ALERT: ADAUTHDC3;Service: AD Web Service;CRITICAL;HARD;3;connect to address 164.76.102.173 and port 12489: No route to host
[1462820749] SERVICE ALERT: halleground1;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1462820770] SERVICE ALERT: ADAUTHDC3;Uptime;CRITICAL;SOFT;1;connect to address 164.76.102.173 and port 12489: No route to host
[1462820770] SERVICE ALERT: ADAUTHDC3;Service: DNS Service;CRITICAL;SOFT;1;connect to address 164.76.102.173 and port 12489: No route to host
[1462820779] SERVICE ALERT: ADAUTHDC3;Service: DNS Client;CRITICAL;SOFT;2;connect to address 164.76.102.173 and port 12489: No route to host
[1462820779] SERVICE ALERT: ADAUTHDC3;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.102.173 and port 12489: No route to host
[1462820779] SERVICE ALERT: ADAUTHDC3;Service: DFS Namespace;CRITICAL;HARD;3;connect to address 164.76.102.173 and port 12489: No route to host
[1462820779] HOST ALERT: halleground1;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1462820780] SERVICE ALERT: FITNESSGRAM;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1462820793] SERVICE ALERT: CEDARON;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1462820794] SERVICE ALERT: FITNESSGRAM;CPU Load;OK;HARD;3;CPU Load 0% (15 min average)
[1462820797] SERVICE ALERT: ADAUTHDC3;Service: WMI;CRITICAL;HARD;3;connect to address 164.76.102.173 and port 12489: No route to host
[1462820806] SERVICE ALERT: ADAUTHDC3;VMWare Tools;CRITICAL;SOFT;2;connect to address 164.76.102.173 and port 12489: No route to host
[1462820813] SERVICE ALERT: ADAUTHDC3;AD Replication;WARNING;SOFT;2;(No output on stdout) stderr: Could not resolve hostname ADAUTHDC3: Name or service not known
[1462820826] SERVICE ALERT: ADAUTHDC3;Service: NetLogin;CRITICAL;SOFT;2;connect to address 164.76.102.173 and port 12489: No route to host
[1462820830] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462820833] SERVICE ALERT: FITNESSGRAM;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1462820833] SERVICE ALERT: CEDARON;CPU Load;OK;HARD;3;CPU Load 3% (15 min average)
[1462820833] SERVICE ALERT: CEDARON;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1462820834] SERVICE ALERT: ADAUTHDC3;Memory Usage;CRITICAL;SOFT;2;connect to address 164.76.102.173 and port 12489: No route to host
[1462820841] SERVICE ALERT: CEDARON;Uptime;OK;HARD;3;System Uptime - 17 day(s) 1 hour(s) 41 minute(s)
[1462820846] SERVICE ALERT: ADAUTHDC3;Service: Group Policy;CRITICAL;SOFT;2;connect to address 164.76.102.173 and port 12489: No route to host
[1462820849] SERVICE ALERT: ADAUTHDC3;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.102.173 and port 12489: No route to host
[1462820855] SERVICE ALERT: ADAUTHDC3;Service: Distributed File System Replication;CRITICAL;SOFT;2;connect to address 164.76.102.173 and port 12489: No route to host
[1462820858] SERVICE ALERT: FITNESSGRAM;Uptime;OK;HARD;3;System Uptime - 107 day(s) 17 hour(s) 19 minute(s)
[1462820869] HOST ALERT: halleground1;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1462820872] SERVICE ALERT: ADAUTHDC3;Service: Intersite Messaging;CRITICAL;SOFT;1;connect to address 164.76.102.173 and port 12489: No route to host
[1462820886] SERVICE ALERT: ADAUTHDC3;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.102.173 and port 12489: No route to host
[1462820889] SERVICE ALERT: ADAUTHDC3;Uptime;CRITICAL;SOFT;2;connect to address 164.76.102.173 and port 12489: No route to host
[1462820889] SERVICE ALERT: ADAUTHDC3;Service: DNS Service;CRITICAL;SOFT;2;connect to address 164.76.102.173 and port 12489: No route to host
[1462820896] SERVICE ALERT: CEDARON;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1462820899] SERVICE ALERT: ADAUTHDC3;Service: DNS Client;CRITICAL;HARD;3;connect to address 164.76.102.173 and port 12489: No route to host
[1462820899] SERVICE ALERT: ADAUTHDC3;Service: Firewall;CRITICAL;SOFT;2;connect to address 164.76.102.173 and port 12489: No route to host
[1462820899] SERVICE ALERT: FITNESSGRAM;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1462820907] HOST DOWNTIME ALERT: ADAUTHDC3;STARTED; Host has entered a period of scheduled downtime
[1462820912] SERVICE ALERT: ADAUTHDC3;CPU Load;CRITICAL;SOFT;1;connect to address 164.76.102.173 and port 12489: No route to host
[1462820912] SERVICE ALERT: ADAUTHDC3;Service: Key Distribution Center;CRITICAL;SOFT;1;connect to address 164.76.102.173 and port 12489: No route to host
[1462820926] SERVICE ALERT: ADAUTHDC3;VMWare Tools;CRITICAL;HARD;3;connect to address 164.76.102.173 and port 12489: No route to host
[1462820933] SERVICE ALERT: ADAUTHDC3;AD Replication;WARNING;HARD;3;(No output on stdout) stderr: Could not resolve hostname ADAUTHDC3: Name or service not known
[1462820939] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1462820946] SERVICE ALERT: ADAUTHDC3;Service: NetLogin;CRITICAL;HARD;3;connect to address 164.76.102.173 and port 12489: No route to host
[1462820950] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462820954] SERVICE ALERT: ADAUTHDC3;Memory Usage;CRITICAL;HARD;3;connect to address 164.76.102.173 and port 12489: No route to host
[1462820958] SERVICE ALERT: CEDARON;Memory Usage;OK;HARD;3;Memory usage: total:21759.74 MB - used: 2468.24 MB (11%) - free: 19291.49 MB (89%)
[1462820959] HOST ALERT: halleground1;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1462820963] SERVICE ALERT: FITNESSGRAM;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1462820963] SERVICE ALERT: FITNESSGRAM;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1462820966] SERVICE ALERT: ADAUTHDC3;Service: Windows Time;CRITICAL;SOFT;2;connect to address 164.76.102.173 and port 12489: No route to host
[1462820966] SERVICE ALERT: ADAUTHDC3;Service: Group Policy;CRITICAL;HARD;3;connect to address 164.76.102.173 and port 12489: No route to host
[1462820966] SERVICE ALERT: ADAUTHDC3;Service: Distributed File System Replication;CRITICAL;HARD;3;connect to address 164.76.102.173 and port 12489: No route to host
[1462820978] SERVICE ALERT: CEDARON;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1462820979] SERVICE ALERT: FITNESSGRAM;Service: RPC;OK;HARD;3;RpcSs: Started
[1462820985] SERVICE ALERT: FITNESSGRAM;C:\ Drive Space;OK;HARD;3;c:\ - total: 59.90 Gb - used: 27.41 Gb (46%) - free 32.49 Gb (54%)
[1462820992] SERVICE ALERT: ADAUTHDC3;Service: Intersite Messaging;CRITICAL;SOFT;2;connect to address 164.76.102.173 and port 12489: No route to host
[1462820999] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1462821006] SERVICE ALERT: ADAUTHDC3;C:\ Drive Space;CRITICAL;SOFT;2;connect to address 164.76.102.173 and port 12489: No route to host
[1462821010] SERVICE ALERT: ADAUTHDC3;Service: DNS Service;CRITICAL;HARD;3;connect to address 164.76.102.173 and port 12489: No route to host
[1462821010] SERVICE ALERT: ADAUTHDC3;Uptime;CRITICAL;HARD;3;connect to address 164.76.102.173 and port 12489: No route to host
[1462821021] SERVICE ALERT: ADAUTHDC3;Service: Firewall;CRITICAL;HARD;3;connect to address 164.76.102.173 and port 12489: No route to host
[1462821029] HOST ALERT: halleground1;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1462821032] SERVICE ALERT: ADAUTHDC3;CPU Load;CRITICAL;SOFT;2;connect to address 164.76.102.173 and port 12489: No route to host
[1462821032] SERVICE ALERT: ADAUTHDC3;Service: Key Distribution Center;CRITICAL;SOFT;2;connect to address 164.76.102.173 and port 12489: No route to host
[1462821059] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1462821070] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1462821086] SERVICE ALERT: ADAUTHDC3;Service: Windows Time;CRITICAL;HARD;3;connect to address 164.76.102.173 and port 12489: No route to host
[1462821089] HOST ALERT: halleground1;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1462821112] SERVICE ALERT: ADAUTHDC3;Service: Intersite Messaging;CRITICAL;HARD;3;connect to address 164.76.102.173 and port 12489: No route to host
[1462821126] SERVICE ALERT: ADAUTHDC3;C:\ Drive Space;CRITICAL;HARD;3;connect to address 164.76.102.173 and port 12489: No route to host
[1462821152] SERVICE ALERT: ADAUTHDC3;Service: Key Distribution Center;CRITICAL;HARD;3;connect to address 164.76.102.173 and port 12489: No route to host
[1462821152] SERVICE ALERT: ADAUTHDC3;CPU Load;CRITICAL;HARD;3;connect to address 164.76.102.173 and port 12489: No route to host
[1462821164] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;OK;HARD;3;0
[1462821859] SERVICE ALERT: PROCARE;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1462821979] SERVICE ALERT: PROCARE;CPU Load;OK;SOFT;2;CPU Load 1% (15 min average)
[1462822171] SERVICE ALERT: AX-QUAL;Service: User Profile;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1462822237] SERVICE ALERT: ADDC5;GPO Replication;OK;HARD;3;OK! All GPOs are properly replicated.
[1462822290] SERVICE ALERT: AX-QUAL;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1462823469] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (15 min average)
[1462824069] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462824159] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462824189] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1462824218] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462824278] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462824309] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462825509] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 76% (15 min average)
[1462826103] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1462826109] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1462826223] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1462826229] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462826343] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1462826349] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462826943] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 75% (15 min average)
[1462827278] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Processing job from tray 3
[1462827487] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462827547] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462827608] SERVICE ALERT: mckennypk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462827878] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462827938] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462827998] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462829897] SERVICE ALERT: TABLEAU;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1462829941] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1462830017] SERVICE ALERT: TABLEAU;CPU Load;OK;SOFT;2;CPU Load 4% (15 min average)
[1462830061] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 88% (15 min average)
[1462830549] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (15 min average)
[1462831598] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Clearing paper path
[1462831598] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1462831749] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462831869] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1462831989] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462832198] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462832258] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462832318] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462832461] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1462832581] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 92% (15 min average)
[1462833181] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1462833301] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1462833421] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1462833789] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1462834021] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1462834389] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462834509] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1462834611] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1462834629] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462834671] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 45% - OK!
[1462835229] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 80% (15 min average)
[1462835821] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1462835829] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462835941] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1462835949] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462836061] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 79% (15 min average)
[1462836069] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462836669] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462837269] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462838461] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462838581] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 91% (15 min average)
[1462839069] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 90% (15 min average)
[1462839669] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1462839789] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 91% (15 min average)
[1462840989] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462841109] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1462841229] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462842429] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462843718] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1462844229] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462845429] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462845781] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1462846029] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462848031] SERVICE ALERT: BBTS_DB;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462848131] SERVICE ALERT: BBTS_DB;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1462848784] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1462848901] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1462849021] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1462849621] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462850221] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1462850821] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1462850821] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1462851421] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1462851541] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1462852627] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 84% (15 min average)