[1493353801] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1493355183] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;1;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1493355303] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;2;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1493355423] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;HARD;3;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1493356088] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1493356147] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;2;Warning:  idm1  are more than 120 seconds behind!
[1493356207] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1493356811] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1493357408] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1493357467] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1493357527] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1493357996] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1493358056] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1493358116] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 89% (15 min average)
[1493358716] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1493358776] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1493358836] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 93% (15 min average)
[1493359436] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1493359496] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1493359556] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1493360156] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1493360755] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1493362554] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1493362554] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1493362623] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;OK;HARD;3;OK: No entries found
[1493363154] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1493363214] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1493363274] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1493363874] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1493364127] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1493364127] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1493365074] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1493365134] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1493365194] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1493365327] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1493365387] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1493365447] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1493365794] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 67% (15 min average)
[1493366994] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1493367054] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1493367114] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 77% (15 min average)
[1493368200] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service has exited from a period of scheduled downtime
[1493368203] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (41.6% change >= 20.0% threshold)
[1493368203] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1493368203] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1493368203] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1493368203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1493368203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1493368203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1493368203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1493368203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1493368203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1493368203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1493368203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1493368203] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1493368203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1493368203] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1493368203] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1493368203] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1493368447] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1493370002] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (38.4% change >= 20.0% threshold)
[1493370003] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1493370003] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1493370003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1493370003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1493370003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1493370003] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1493370003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1493370003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1493370003] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1493370003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1493370003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1493370003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1493370003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1493370003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1493370003] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1493372046] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1493372106] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;2;Warning:  idm1  are more than 120 seconds behind!
[1493372166] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;3;OK. All replicas within 120 seconds.
[1493374315] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1493374374] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1493374434] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1493375034] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 82% (15 min average)
[1493376834] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1493376894] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 54% (15 min average)
[1493377494] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1493377554] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1493377614] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1493377824] HOST ALERT: marshallpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1493377914] HOST ALERT: marshallpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1493378004] HOST ALERT: marshallpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1493378005] SERVICE ALERT: marshallpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.137.49. Is the printer off?  Is the network down?
[1493378035] HOST ALERT: marshallpk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1493378049] SERVICE ALERT: marshallpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.137.49. Is the printer off?  Is the network down?
[1493378065] SERVICE ALERT: marshallpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.137.49. Is the printer off?  Is the network down?
[1493378079] HOST ALERT: marshallpk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1493378109] SERVICE ALERT: marshallpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.137.49. Is the printer off?  Is the network down?
[1493378125] SERVICE ALERT: marshallpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.137.49. Is the printer off?  Is the network down?
[1493378139] HOST ALERT: marshallpk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1493378169] SERVICE ALERT: marshallpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.137.49. Is the printer off?  Is the network down?
[1493378199] HOST ALERT: marshallpk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1493378234] SERVICE ALERT: marshallpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.137.49. Is the printer off?  Is the network down?
[1493378264] HOST ALERT: marshallpk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1493378294] SERVICE ALERT: marshallpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.137.49. Is the printer off?  Is the network down?
[1493378324] HOST ALERT: marshallpk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1493378354] SERVICE ALERT: marshallpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.137.49. Is the printer off?  Is the network down?
[1493378384] HOST ALERT: marshallpk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1493380488] HOST ALERT: marshallpk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.52 ms
[1493380519] SERVICE ALERT: marshallpk;Device Status;OK;HARD;3;Ready
[1493380563] SERVICE ALERT: marshallpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 9% - OK!
[1493380748] SERVICE ALERT: marshallpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1493381767] SERVICE ALERT: ADRODC-LIV;Service: Windows Time;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1493381812] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 26% (15 min average)
[1493381847] SERVICE ALERT: ADRODC-LIV;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1493381877] HOST ALERT: ADRODC-LIV;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1493381890] SERVICE ALERT: ADRODC-LIV;Service: Windows Time;CRITICAL;HARD;3;connect to address 207.74.163.140 and port 12489: Connection refused
[1493381904] HOST ALERT: ADRODC-LIV;UP;SOFT;2;PING OK - Packet loss = 66%, RTA = 6.03 ms
[1493381950] SERVICE ALERT: ADRODC-LIV;Service: MDT Monitor;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1493381990] SERVICE ALERT: ADRODC-LIV;Service: MDT Monitor;OK;SOFT;2;MDT_Monitor: Started
[1493382487] SERVICE ALERT: ADRODC-LIV;Service: Windows Time;OK;HARD;3;W32Time: Started
[1493384730] SERVICE ALERT: hallegroundcolor;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1493384785] SERVICE ALERT: hallegroundcolor;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! All Trays Still Full
[1493385711] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1493387109] HOST ALERT: sega;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1493387173] HOST ALERT: sega;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.25 ms
[1493387188] SERVICE ALERT: MBAM;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1493387248] SERVICE ALERT: MBAM;CPU Load;OK;SOFT;2;CPU Load 7% (15 min average)
[1493387574] SERVICE ALERT: alexanderpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.142.94. Is the printer off?  Is the network down?
[1493387604] HOST ALERT: alexanderpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1493387669] SERVICE ALERT: alexanderpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.142.94. Is the printer off?  Is the network down?
[1493387694] HOST ALERT: alexanderpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1493387729] SERVICE ALERT: alexanderpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.142.94. Is the printer off?  Is the network down?
[1493387759] HOST ALERT: alexanderpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1493387781] SERVICE ALERT: alexanderpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.142.94. Is the printer off?  Is the network down?
[1493387789] SERVICE ALERT: alexanderpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.142.94. Is the printer off?  Is the network down?
[1493387811] HOST ALERT: alexanderpk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1493387841] SERVICE ALERT: alexanderpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.142.94. Is the printer off?  Is the network down?
[1493387871] HOST ALERT: alexanderpk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1493387901] SERVICE ALERT: alexanderpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.142.94. Is the printer off?  Is the network down?
[1493387931] HOST ALERT: alexanderpk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1493387995] HOST ALERT: alexanderpk;UP;SOFT;7;PING OK - Packet loss = 0%, RTA = 0.55 ms
[1493388169] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1493388312] SERVICE ALERT: hallegroundcolor;Black Cartridge HP CE260X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Cartridge HP CE260X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1493388367] SERVICE ALERT: hallegroundcolor;Black Cartridge HP CE260X;OK;SOFT;2;Black Cartridge HP CE260X is at 47% - OK!
[1493388383] SERVICE ALERT: alexanderpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 3% - OK!
[1493388423] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;1;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1493388495] SERVICE ALERT: alexanderpk;Device Status;OK;HARD;3;LOAD TRAY 3 PLAIN LETTER To use another 74 72 61 79 20 70 72 65 73 73 20 1E  LOAD TRAY 3 PLAIN LETTER 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1493388543] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;2;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1493388663] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;HARD;3;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1493388981] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1493388994] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1493389008] SERVICE ALERT: hallefirstcolor;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1493389011] HOST ALERT: hallefirstcolor;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1493389049] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;OK;SOFT;2;Magenta Cartridge HP CE263A is at 97% - OK!
[1493389053] HOST ALERT: hallefirstcolor;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1493389062] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! All Trays Still Full
[1493389576] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;CRITICAL;HARD;3;Yellow Cartridge HP CE262A is at CRITICAL level!
[1493390163] SERVICE ALERT: hallegroundcolor;Image Transfer Kit HP CE249A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1493390167] SERVICE ALERT: ADWU1;Windows Update Services Event 364;UNKNOWN;SOFT;1;Failed to process command : EvtQuery failed: 1717: The interface is unknown.
[1493390217] SERVICE ALERT: hallegroundcolor;Image Transfer Kit HP CE249A;OK;SOFT;2;Image Transfer Kit HP CE249A is at 42% - OK!
[1493390287] SERVICE ALERT: ADWU1;Windows Update Services Event 364;OK;SOFT;2;OK: No entries found
[1493390336] SERVICE ALERT: hallefirst3;Device Status;OK;HARD;3;Warming up
[1493392872] HOST ALERT: rackhampk;UP;HARD;10;PING OK - Packet loss = 50%, RTA = 0.42 ms
[1493393467] HOST ALERT: rackhampk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1493393557] HOST ALERT: rackhampk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1493393647] HOST ALERT: rackhampk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1493393712] HOST ALERT: rackhampk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1493393763] HOST ALERT: rackhampk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1493393853] HOST ALERT: rackhampk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1493393943] HOST ALERT: rackhampk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1493394033] HOST ALERT: rackhampk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1493394097] HOST ALERT: rackhampk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1493394187] HOST ALERT: rackhampk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1493394190] SERVICE ALERT: halleground1;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1493394250] SERVICE ALERT: halleground1;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1493394310] SERVICE ALERT: halleground1;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1493394886] HOST ALERT: rackhampk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1493395229] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 99.90 Gb - used: 90.35 Gb (90%) - free 9.55 Gb (10%)
[1493395263] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;OK;HARD;3;OK: No entries found
[1493395267] HOST ALERT: rackhampk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1493395289] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 99.90 Gb - used: 90.35 Gb (90%) - free 9.55 Gb (10%)
[1493395349] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 99.90 Gb - used: 90.35 Gb (90%) - free 9.55 Gb (10%)
[1493395357] HOST ALERT: rackhampk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1493395447] HOST ALERT: rackhampk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1493395483] SERVICE ALERT: CAVA2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1493395512] HOST ALERT: rackhampk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1493395543] SERVICE ALERT: CAVA2;CPU Load;OK;SOFT;2;CPU Load 1% (15 min average)
[1493395563] HOST ALERT: rackhampk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1493395653] HOST ALERT: rackhampk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1493395743] HOST ALERT: rackhampk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1493395833] HOST ALERT: rackhampk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1493395897] HOST ALERT: rackhampk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1493395987] HOST ALERT: rackhampk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1493397003] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1493397310] SERVICE ALERT: halleground1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1493400797] HOST ALERT: HDBACKUPS;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1493400816] SERVICE ALERT: HDBACKUPS;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1493400842] SERVICE ALERT: HDBACKUPS;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1493400846] HOST ALERT: HDBACKUPS;DOWN;HARD;2;PING CRITICAL - Packet loss = 100%
[1493400847] SERVICE ALERT: HDBACKUPS;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1493400849] SERVICE ALERT: HDBACKUPS;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1493400856] SERVICE ALERT: HDBACKUPS;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1493400857] SERVICE ALERT: HDBACKUPS;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1493400870] SERVICE ALERT: HDBACKUPS;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1493400876] SERVICE ALERT: HDBACKUPS;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1493400878] SERVICE ALERT: HDBACKUPS;F:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1493400897] SERVICE ALERT: HDBACKUPS;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1493400902] SERVICE ALERT: HDBACKUPS;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1493400907] SERVICE ALERT: HDBACKUPS;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1493400909] SERVICE ALERT: HDBACKUPS;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1493400916] SERVICE ALERT: HDBACKUPS;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1493400917] SERVICE ALERT: HDBACKUPS;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1493400930] SERVICE ALERT: HDBACKUPS;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1493400932] SERVICE ALERT: HDBACKUPS;Service: User Profile;CRITICAL;SOFT;1;connect to address 164.76.4.109 and port 12489: No route to host
[1493400936] SERVICE ALERT: HDBACKUPS;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1493400938] SERVICE ALERT: HDBACKUPS;F:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1493400939] SERVICE ALERT: HDBACKUPS;Service: Windows Time;CRITICAL;SOFT;2;connect to address 164.76.4.109 and port 12489: No route to host
[1493400962] SERVICE ALERT: HDBACKUPS;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1493400962] SERVICE ALERT: HDBACKUPS;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1493400967] SERVICE ALERT: HDBACKUPS;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1493400969] SERVICE ALERT: HDBACKUPS;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1493400972] SERVICE ALERT: HDBACKUPS;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1493400976] SERVICE ALERT: HDBACKUPS;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1493400977] SERVICE ALERT: HDBACKUPS;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1493400990] SERVICE ALERT: HDBACKUPS;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1493400994] SERVICE ALERT: HDBACKUPS;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1493400998] SERVICE ALERT: HDBACKUPS;F:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1493401001] SERVICE ALERT: HDBACKUPS;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1493401003] SERVICE ALERT: HDBACKUPS;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1493401012] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1493401017] SERVICE ALERT: HDBACKUPS;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1493401022] SERVICE ALERT: HDBACKUPS;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1493401032] SERVICE ALERT: HDBACKUPS;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1493401054] SERVICE ALERT: HDBACKUPS;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1493401061] SERVICE ALERT: HDBACKUPS;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1493401063] SERVICE ALERT: HDBACKUPS;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1493401072] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 62% (15 min average)
[1493401082] SERVICE ALERT: HDBACKUPS;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1493401092] SERVICE ALERT: HDBACKUPS;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1493401114] SERVICE ALERT: HDBACKUPS;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1493401121] SERVICE ALERT: HDBACKUPS;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1493401131] SERVICE ALERT: HDBACKUPS;EMU Star SSL Cert Expirations;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1493401320] SERVICE ALERT: HDBACKUPS;EMU Star SSL Cert Expirations;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1493401509] SERVICE ALERT: HDBACKUPS;EMU Star SSL Cert Expirations;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1493402086] HOST ALERT: rackhampk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.40 ms
[1493403749] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;OK;HARD;3;c:\ - total: 99.90 Gb - used: 88.82 Gb (89%) - free 11.08 Gb (11%)
[1493403975] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;OK;HARD;3;
[1493404513] HOST ALERT: rackhampk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1493404563] HOST ALERT: rackhampk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1493404653] HOST ALERT: rackhampk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1493404695] SERVICE ALERT: alexanderpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1493404743] HOST ALERT: rackhampk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1493404755] SERVICE ALERT: alexanderpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1493404815] SERVICE ALERT: alexanderpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1493404833] HOST ALERT: rackhampk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1493404897] HOST ALERT: rackhampk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1493404987] HOST ALERT: rackhampk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1493405076] HOST ALERT: rackhampk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1493405113] HOST ALERT: rackhampk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1493405164] HOST ALERT: rackhampk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1493406006] SERVICE ALERT: ADWSUSSERVERS;VMWare Tools;CRITICAL;SOFT;1;connect to address 164.76.4.79 and port 12489: Connection refused
[1493406066] SERVICE ALERT: ADWSUSSERVERS;VMWare Tools;OK;SOFT;2;VMTools: Started
[1493409472] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1493409532] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1493409592] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 91% (15 min average)
[1493410134] SERVICE ALERT: hallefirst3;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1493410194] SERVICE ALERT: hallefirst3;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1493410254] SERVICE ALERT: hallefirst3;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1493411454] SERVICE ALERT: hallefirst3;Device Status;OK;HARD;3;Warming up
[1493414415] SERVICE ALERT: alexanderpk;Device Status;OK;HARD;3;TRAY 3 EMPTY PLAIN LETTER   Processing job from tray 2
[1493418591] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1493418651] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1493418711] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1493419311] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1493429400] HOST DOWNTIME ALERT: IMAGES;STARTED; Host has entered a period of scheduled downtime
[1493429837] SERVICE ALERT: TARGETX-2012;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1493429897] SERVICE ALERT: TARGETX-2012;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1493429957] SERVICE ALERT: TARGETX-2012;CPU Load;OK;SOFT;3;CPU Load 90% (15 min average)
[1493430417] HOST ALERT: IMAGES;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1493430482] SERVICE ALERT: IMAGES;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1493430500] HOST ALERT: IMAGES;UP;SOFT;2;PING OK - Packet loss = 79%, RTA = 0.25 ms
[1493430522] SERVICE ALERT: IMAGES;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1493431311] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1493431371] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 43% (15 min average)
[1493431971] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1493432031] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1493432091] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1493432691] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 39% (15 min average)
[1493433608] HOST ALERT: emu-fndn-fe7;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1493433625] SERVICE ALERT: emu-fndn-fe7;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1493433655] HOST ALERT: emu-fndn-fe7;DOWN;HARD;2;PING CRITICAL - Packet loss = 100%
[1493433685] SERVICE ALERT: emu-fndn-fe7;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1493433710] SERVICE ALERT: emu-fndn-fe7;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1493433710] SERVICE ALERT: emu-fndn-fe7;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1493433745] SERVICE ALERT: emu-fndn-fe7;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1493433769] SERVICE ALERT: emu-fndn-fe7;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1493433769] SERVICE ALERT: emu-fndn-fe7;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1493433797] SERVICE ALERT: emu-fndn-fe7;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1493433797] SERVICE ALERT: emu-fndn-fe7;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1493433829] SERVICE ALERT: emu-fndn-fe7;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1493433829] SERVICE ALERT: emu-fndn-fe7;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1493433857] SERVICE ALERT: emu-fndn-fe7;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1493433857] SERVICE ALERT: emu-fndn-fe7;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1493433882] SERVICE ALERT: emu-fndn-fe7;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1493433884] SERVICE ALERT: emu-fndn-fe7;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1493433917] SERVICE ALERT: emu-fndn-fe7;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1493433917] SERVICE ALERT: emu-fndn-fe7;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1493433942] SERVICE ALERT: emu-fndn-fe7;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1493433944] SERVICE ALERT: emu-fndn-fe7;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1493433968] SERVICE ALERT: emu-fndn-fe7;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1493433970] SERVICE ALERT: emu-fndn-fe7;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1493434002] SERVICE ALERT: emu-fndn-fe7;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1493434004] SERVICE ALERT: emu-fndn-fe7;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1493434028] SERVICE ALERT: emu-fndn-fe7;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1493434030] SERVICE ALERT: emu-fndn-fe7;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1493434054] SERVICE ALERT: emu-fndn-fe7;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1493434054] SERVICE ALERT: emu-fndn-fe7;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1493434088] SERVICE ALERT: emu-fndn-fe7;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1493434090] SERVICE ALERT: emu-fndn-fe7;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1493434113] SERVICE ALERT: emu-fndn-fe7;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1493434114] SERVICE ALERT: emu-fndn-fe7;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1493434142] SERVICE ALERT: emu-fndn-fe7;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1493434173] SERVICE ALERT: emu-fndn-fe7;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1493434173] SERVICE ALERT: emu-fndn-fe7;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1493434174] SERVICE ALERT: emu-fndn-fe7;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1493434202] SERVICE ALERT: emu-fndn-fe7;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1493434233] SERVICE ALERT: emu-fndn-fe7;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1493434262] SERVICE ALERT: emu-fndn-fe7;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1493434293] SERVICE ALERT: emu-fndn-fe7;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1493434306] SERVICE ALERT: emu-fndn-fe7;EMU Star SSL Cert Expirations;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1493434399] SERVICE ALERT: emu-fndn-fe7;EMU Star SSL Cert Expirations;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1493434399] SERVICE ALERT: emu-fndn-fe7;EMU Star SSL Cert Expirations;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1493434403] HOST ALERT: emu-fndn-fe7;UP;HARD;2;PING OK - Packet loss = 0%, RTA = 0.36 ms
[1493434409] SERVICE ALERT: emu-fndn-fe7;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1493434410] SERVICE ALERT: emu-fndn-fe7;Service: WMI;OK;HARD;3;Winmgmt: Started
[1493434497] SERVICE ALERT: emu-fndn-fe7;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1493434582] SERVICE ALERT: emu-fndn-fe7;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1493434583] SERVICE ALERT: emu-fndn-fe7;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 3 minute(s)
[1493434668] SERVICE ALERT: emu-fndn-fe7;C:\ Drive Space;OK;HARD;3;c:\ - total: 698.53 Gb - used: 43.65 Gb (6%) - free 654.89 Gb (94%)
[1493434669] SERVICE ALERT: emu-fndn-fe7;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1493434753] SERVICE ALERT: emu-fndn-fe7;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1493434754] SERVICE ALERT: emu-fndn-fe7;CPU Load;OK;HARD;3;CPU Load 0% (15 min average)
[1493434842] SERVICE ALERT: emu-fndn-fe7;Service: RPC;OK;HARD;3;RpcSs: Started
[1493434873] SERVICE ALERT: emu-fndn-fe7;Memory Usage;OK;HARD;3;Memory usage: total:18781.43 MB - used: 1729.00 MB (9%) - free: 17052.43 MB (91%)
[1493434925] SERVICE ALERT: emu-fndn-fe7;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1493435008] SERVICE ALERT: emu-fndn-fe7;EMU Star SSL Cert Expirations;OK;HARD;3;OK! All certificates are current.
[1493435097] SERVICE ALERT: emu-fndn-fe7;Service: Windows Time;OK;HARD;3;W32Time: Started
[1493436600] HOST DOWNTIME ALERT: IMAGES;STOPPED; Host has exited from a period of scheduled downtime