[1462939259] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462939318] SERVICE ALERT: sillpk;Device Status;OK;SOFT;3;Open tray indicated.
[1462939917] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462939977] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462940037] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462940637] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Open tray indicated.
[1462941174] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462941294] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1462941414] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 93% (15 min average)
[1462941837] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462941897] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462941958] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462942143] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 81% (15 min average)
[1462942557] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Open tray indicated.
[1462942743] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462942863] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462942983] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 93% (15 min average)
[1462943583] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462943703] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462943757] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462943817] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462943823] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462943877] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462944477] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Open tray indicated.
[1462945678] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462945737] SERVICE ALERT: sillpk;Device Status;OK;SOFT;2;Open tray indicated.
[1462946214] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462946334] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1462946454] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1462946527] SERVICE ALERT: BOEAPPTEST;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1462946647] SERVICE ALERT: BOEAPPTEST;CPU Load;OK;SOFT;2;CPU Load 29% (15 min average)
[1462946937] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462946997] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462947054] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462947054] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1462947057] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462947423] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462947654] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462947657] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Open tray indicated.
[1462948023] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462948858] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462948917] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462948977] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462949199] SERVICE ALERT: VUM;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1462949223] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462949319] SERVICE ALERT: VUM;CPU Load;OK;SOFT;2;CPU Load 6% (15 min average)
[1462949577] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Open tray indicated.
[1462949823] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462950777] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462950837] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462950898] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462951254] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462951497] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Open tray indicated.
[1462951623] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462951854] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1462952096] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462952156] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462952216] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462952223] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462952454] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462952574] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462952694] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1462952816] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Open tray indicated.
[1462953295] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462953416] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462953423] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 84% (15 min average)
[1462953476] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462953537] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462954023] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462954136] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Open tray indicated.
[1462954143] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462954263] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462954494] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 88% (15 min average)
[1462955094] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1462955214] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1462955334] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 88% (15 min average)
[1462955336] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462955396] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462955456] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462955934] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462956054] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1462956056] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Open tray indicated.
[1462956063] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 84% (15 min average)
[1462956174] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1462957257] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462957263] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1462957316] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462957376] SERVICE ALERT: sillpk;Device Status;OK;SOFT;3;Open tray indicated.
[1462957383] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 84% (15 min average)
[1462958574] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1462958576] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462958636] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462958696] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462959174] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1462959294] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1462959297] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Open tray indicated.
[1462959414] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1462959783] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462959903] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 92% (15 min average)
[1462960014] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1462960495] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462960555] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462960614] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462960615] SERVICE ALERT: sillpk;Device Status;OK;SOFT;3;Open tray indicated.
[1462960734] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1462961334] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462961454] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1462961574] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1462961703] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462961815] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462961823] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1462961875] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462961936] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462961943] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1462962174] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 87% (15 min average)
[1462962535] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Open tray indicated.
[1462962543] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462962774] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462962894] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1462963014] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 76% (15 min average)
[1462963135] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462963143] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462963195] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462963255] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462964214] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1462964334] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 85% (15 min average)
[1462964343] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1462964455] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Open tray indicated.
[1462964934] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462964943] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1462965054] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1462965055] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462965063] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462965116] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462965174] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 94% (15 min average)
[1462965175] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462965183] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462965774] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1462965775] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Open tray indicated.
[1462965894] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1462966014] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 94% (15 min average)
[1462966614] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1462966734] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1462966854] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1462966975] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462967035] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462967095] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462967454] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1462967583] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1462967695] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Open tray indicated.
[1462968054] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462968174] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 70% (15 min average)
[1462968183] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462968303] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462968423] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462968895] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462968955] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462969015] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462969023] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 81% (15 min average)
[1462969563] SERVICE ALERT: mckennypk;Device Status;OK;HARD;3;Genuine HP cartridge installed
[1462969803] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STARTED; Service appears to have started flapping (38.3% change >= 20.0% threshold)
[1462969803] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (42.9% change >= 20.0% threshold)
[1462969803] SERVICE FLAPPING ALERT: sillpk;Device Status;STARTED; Service appears to have started flapping (56.2% change >= 20.0% threshold)
[1462969803] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1462969803] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1462969803] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1462969803] HOST DOWNTIME ALERT: R25-ZETA2012;STARTED; Host has entered a period of scheduled downtime
[1462969803] HOST DOWNTIME ALERT: NETENG-PA;STARTED; Host has entered a period of scheduled downtime
[1462969803] SERVICE DOWNTIME ALERT: NETIDADPROXY;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1462969803] HOST DOWNTIME ALERT: TABLEAU;STARTED; Host has entered a period of scheduled downtime
[1462969803] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1462969803] HOST DOWNTIME ALERT: R25-ALPHA2012;STARTED; Host has entered a period of scheduled downtime
[1462969803] HOST DOWNTIME ALERT: R25-DELTA2012;STARTED; Host has entered a period of scheduled downtime
[1462969808] HOST ALERT: TABLEAU;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.41 ms
[1462969831] SERVICE ALERT: NETENG-PA;Service: RPC;OK;HARD;3;RpcSs: Started
[1462969835] HOST ALERT: NETENG-PA;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1462969840] SERVICE ALERT: TABLEAU;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1462969875] SERVICE ALERT: TABLEAU;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1462969881] SERVICE ALERT: TABLEAU;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1462969887] SERVICE ALERT: NETENG-PA;Memory Usage;OK;HARD;3;Memory usage: total:14333.42 MB - used: 1418.43 MB (10%) - free: 12914.99 MB (90%)
[1462969948] SERVICE ALERT: NETENG-PA;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1462969953] SERVICE ALERT: NETENG-PA;CPU Load;OK;HARD;3;CPU Load 0% (15 min average)
[1462970045] SERVICE ALERT: NETENG-PA;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1462970050] SERVICE ALERT: TABLEAU;CPU Load;OK;HARD;3;CPU Load 0% (15 min average)
[1462970064] SERVICE ALERT: NETENG-PA;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1462970083] SERVICE ALERT: TABLEAU;Service: Windows Time;OK;HARD;3;W32Time: Started
[1462970103] SERVICE ALERT: NETENG-PA;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1462970134] SERVICE ALERT: NETENG-PA;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1462970164] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462970190] SERVICE ALERT: NETENG-PA;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1462970216] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Open tray indicated.
[1462970223] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462970223] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462970223] SERVICE ALERT: TABLEAU;Uptime;OK;HARD;3;System Uptime - 0 day(s) 17 hour(s) 17 minute(s)
[1462970245] SERVICE ALERT: TABLEAU;C:\ Drive Space;OK;HARD;3;c:\ - total: 99.66 Gb - used: 68.75 Gb (69%) - free 30.91 Gb (31%)
[1462970248] SERVICE ALERT: NETENG-PA;Service: WMI;OK;HARD;3;Winmgmt: Started
[1462970249] SERVICE ALERT: TABLEAU;Service: WMI;OK;HARD;3;Winmgmt: Started
[1462970250] SERVICE ALERT: TABLEAU;VMWare Tools;OK;HARD;3;VMTools: Started
[1462970253] SERVICE ALERT: TABLEAU;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1462970276] SERVICE ALERT: TABLEAU;Service: RPC;OK;HARD;3;RpcSs: Started
[1462970276] SERVICE ALERT: TABLEAU;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1462970283] SERVICE ALERT: mckennypk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462970314] SERVICE ALERT: TABLEAU;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1462970343] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1462970343] SERVICE ALERT: NETENG-PA;Uptime;OK;HARD;3;System Uptime - 0 day(s) 18 hour(s) 55 minute(s)
[1462970355] SERVICE ALERT: TABLEAU;Memory Usage;OK;HARD;3;Memory usage: total:21759.74 MB - used: 7724.14 MB (35%) - free: 14035.59 MB (65%)
[1462970367] SERVICE ALERT: NETENG-PA;C:\ Drive Space;OK;HARD;3;c:\ - total: 59.90 Gb - used: 28.91 Gb (48%) - free 30.99 Gb (52%)
[1462970380] SERVICE ALERT: NETENG-PA;VMWare Tools;OK;HARD;3;VMTools: Started
[1462970395] SERVICE ALERT: NETENG-PA;Service: Windows Time;OK;HARD;3;W32Time: Started
[1462970463] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 91% (15 min average)
[1462971063] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462971183] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 84% (15 min average)
[1462971415] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462971475] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462971535] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462971774] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462971894] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1462972014] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 93% (15 min average)
[1462972135] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Open tray indicated.
[1462972383] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462972503] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462972623] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462973107] SERVICE ALERT: AD-RMAD;E:\ Drive Space;OK;HARD;3;e:\ - total: 599.87 Gb - used: 428.64 Gb (71%) - free 171.24 Gb (29%)
[1462973336] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462973395] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462973455] SERVICE ALERT: sillpk;Device Status;OK;SOFT;3;Open tray indicated.
[1462973814] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1462973884] SERVICE ALERT: mckennypk;Device Status;OK;HARD;3;Warming up
[1462973934] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1462974054] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 82% (15 min average)
[1462974483] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462974543] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462974603] SERVICE ALERT: mckennypk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462974603] SERVICE FLAPPING ALERT: mckennypk;Device Status;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1462974655] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462974715] SERVICE ALERT: sillpk;Device Status;OK;SOFT;2;Open tray indicated.
[1462975254] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462975374] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1462975494] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462975915] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462975975] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462976036] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462976094] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 80% (15 min average)
[1462976635] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Open tray indicated.
[1462976652] SERVICE ALERT: markjeffersonpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1462976682] HOST ALERT: markjeffersonpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462976692] SERVICE ALERT: markjeffersonpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1462976712] SERVICE ALERT: markjeffersonpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1462976719] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1462976722] HOST ALERT: markjeffersonpk;UP;SOFT;2;PING WARNING - Packet loss = 86%, RTA = 0.58 ms
[1462976746] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462976767] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;SOFT;3;Printer Paper Is OK! All Trays Still Full
[1462976773] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 72% - OK!
[1462976806] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462977294] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1462977414] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1462977423] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462977835] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462977895] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462977955] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462978006] SERVICE ALERT: markjeffersonpk;Device Status;OK;HARD;3;Ready
[1462978023] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462978203] SERVICE ALERT: mckennypk;Device Status;OK;HARD;3;50 72 65 73 73 20 17 20 74 6F 20 73 65 65 20 73  74 65 70  by step information
[1462978555] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Open tray indicated.
[1462979037] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (29.5% change >= 20.0% threshold)
[1462979038] SERVICE FLAPPING ALERT: mckennypk;Device Status;STARTED; Service appears to have started flapping (25.3% change >= 20.0% threshold)
[1462979038] SERVICE FLAPPING ALERT: sillpk;Device Status;STARTED; Service appears to have started flapping (50.4% change >= 20.0% threshold)
[1462979038] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1462979038] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1462979038] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1462979038] HOST DOWNTIME ALERT: R25-ZETA2012;STARTED; Host has entered a period of scheduled downtime
[1462979038] HOST DOWNTIME ALERT: NETENG-PA;STARTED; Host has entered a period of scheduled downtime
[1462979038] SERVICE DOWNTIME ALERT: NETIDADPROXY;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1462979038] HOST DOWNTIME ALERT: TABLEAU;STARTED; Host has entered a period of scheduled downtime
[1462979038] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1462979038] HOST DOWNTIME ALERT: R25-ALPHA2012;STARTED; Host has entered a period of scheduled downtime
[1462979038] HOST DOWNTIME ALERT: R25-DELTA2012;STARTED; Host has entered a period of scheduled downtime
[1462979214] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462979334] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1462979453] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1462979453] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STARTED; Service appears to have started flapping (24.9% change >= 20.0% threshold)
[1462980085] HOST DOWNTIME ALERT: R25-ALPHA2012;STOPPED; Host has exited from a period of scheduled downtime
[1462980094] HOST DOWNTIME ALERT: R25-DELTA2012;STOPPED; Host has exited from a period of scheduled downtime
[1462980116] HOST DOWNTIME ALERT: R25-ZETA2012;STOPPED; Host has exited from a period of scheduled downtime
[1462980421] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 37% (15 min average)
[1462980651] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 82% (15 min average)
[1462981022] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462981141] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462981261] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462981851] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1462981971] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 92% (15 min average)
[1462982460] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 85% (15 min average)
[1462983171] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462983291] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 85% (15 min average)
[1462983302] SERVICE ALERT: PNCDB-2012;D:\ Drive Space;OK;HARD;3;d:\ - total: 450.00 Gb - used: 350.00 Gb (78%) - free 100.00 Gb (22%)
[1462983527] SERVICE ALERT: mckennypk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1462983587] SERVICE ALERT: mckennypk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1462983647] SERVICE ALERT: mckennypk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1462983891] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462984011] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 89% (15 min average)
[1462984043] SERVICE ALERT: alexanderpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462984103] SERVICE ALERT: alexanderpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462984163] SERVICE ALERT: alexanderpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462984203] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462984246] SERVICE ALERT: mckennypk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1462984260] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462984263] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462984323] SERVICE ALERT: mckennypk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462984380] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462984500] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1462985616] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1462985676] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 100% - OK!
[1462986190] HOST ALERT: amseed;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462986202] HOST ALERT: banadmin;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462986205] HOST ALERT: banseed;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462986269] HOST ALERT: rmandb;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462986280] HOST ALERT: amseed;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1462986292] HOST ALERT: banadmin;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1462986292] HOST ALERT: wlsseed;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462986295] HOST ALERT: banseed;DOWN;SOFT;2;(Host check timed out after 30.00 seconds)
[1462986304] HOST ALERT: xeseed;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462986359] HOST ALERT: rmandb;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1462986361] HOST ALERT: wlsseed;UP;SOFT;2;PING OK - Packet loss = 50%, RTA = 0.28 ms
[1462986370] HOST ALERT: amseed;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1462986382] HOST ALERT: banadmin;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1462986385] HOST ALERT: banseed;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1462986394] HOST ALERT: xeseed;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1462986449] HOST ALERT: rmandb;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1462986460] HOST ALERT: amseed;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1462986472] HOST ALERT: banadmin;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1462986475] HOST ALERT: banseed;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1462986484] HOST ALERT: xeseed;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1462986524] HOST ALERT: amseed;UP;SOFT;5;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1462986539] HOST ALERT: banseed;UP;SOFT;5;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1462986539] HOST ALERT: rmandb;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1462986548] HOST ALERT: xeseed;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1462986562] HOST ALERT: banadmin;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1462986626] HOST ALERT: banadmin;UP;SOFT;6;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1462986629] HOST ALERT: rmandb;UP;SOFT;5;PING WARNING - Packet loss = 90%, RTA = 0.35 ms
[1462986890] HOST ALERT: reflin;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462986980] HOST ALERT: reflin;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1462987011] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1462987070] HOST ALERT: reflin;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1462987132] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462987160] HOST ALERT: reflin;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1462987250] HOST ALERT: reflin;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1462987251] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 90% (15 min average)
[1462987340] HOST ALERT: reflin;DOWN;SOFT;6;(Host check timed out after 30.00 seconds)
[1462987430] HOST ALERT: reflin;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1462987500] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 70% (15 min average)
[1462987520] HOST ALERT: reflin;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1462987610] HOST ALERT: reflin;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1462987700] HOST ALERT: reflin;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1462987713] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462987773] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462987833] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462988334] HOST ALERT: reflin;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1462988451] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1462988572] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1462988664] HOST ALERT: reflin;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462988691] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 78% (15 min average)
[1462988700] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462988754] HOST ALERT: reflin;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1462988818] HOST ALERT: reflin;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1462988818] HOST FLAPPING ALERT: reflin;STARTED; Host appears to have started flapping (22.5% change > 20.0% threshold)
[1462988820] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 92% (15 min average)
[1462989033] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Ready
[1462989352] SERVICE ALERT: EVENTS-2012;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1462989420] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462989472] SERVICE ALERT: EVENTS-2012;CPU Load;OK;SOFT;2;CPU Load 0% (15 min average)
[1462989540] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462989569] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1462989599] HOST ALERT: hallefirst2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462989629] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1462989639] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1462989659] HOST ALERT: hallefirst2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1462989660] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462989689] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1462989699] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1462989719] HOST ALERT: hallefirst2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1462989759] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1462989760] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1462989789] HOST ALERT: hallefirst2;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1462989820] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1462989850] HOST ALERT: hallefirst2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1462989880] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1462989891] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1462989910] HOST ALERT: hallefirst2;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1462989955] SERVICE FLAPPING ALERT: sillpk;Device Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1462990000] HOST ALERT: hallefirst2;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1462990011] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 89% (15 min average)
[1462990090] HOST ALERT: hallefirst2;DOWN;SOFT;8;(Host check timed out after 30.00 seconds)
[1462990180] HOST ALERT: hallefirst2;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1462990260] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462990270] HOST ALERT: hallefirst2;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1462990322] SERVICE FLAPPING ALERT: mckennypk;Device Status;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1462990356] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1462990416] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1462990475] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1462991460] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462992014] HOST ALERT: hallefirst2;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.46 ms
[1462992060] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462992084] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1462992153] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Ready
[1462992153] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1462992275] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 54% - OK!
[1462992336] HOST ALERT: itservicestest;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462992411] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1462992423] HOST ALERT: itservicestest;UP;SOFT;2;PING WARNING - Packet loss = 82%, RTA = 0.37 ms
[1462992660] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462993161] SERVICE ALERT: alexanderpk;Device Status;OK;HARD;3;Clearing paper path
[1462993260] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462993475] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1462994289] HOST FLAPPING ALERT: reflin;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1462994692] SERVICE ALERT: LANSWEEPER;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462994692] SERVICE ALERT: LANSWEEPER;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462994694] SERVICE ALERT: LANSWEEPER;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462994694] SERVICE ALERT: LANSWEEPER;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462994716] SERVICE ALERT: BBTS_DB;Size: TS_DATA02.DBF;WARNING;SOFT;1;The term '\\ad\Nagios\Scripts\WinTeamPowershell.ps1' is not recognized as the n
[1462994723] SERVICE ALERT: BBTS_DB;Size: TS_INDEX02.DBF;WARNING;SOFT;1;The term '\\ad\Nagios\Scripts\WinTeamPowershell.ps1' is not recognized as the n
[1462994792] SERVICE ALERT: LANSWEEPER;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1462994792] SERVICE ALERT: LANSWEEPER;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1462994794] SERVICE ALERT: LANSWEEPER;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 1 minute(s)
[1462994794] SERVICE ALERT: LANSWEEPER;C:\ Drive Space;OK;SOFT;2;c:\ - total: 79.66 Gb - used: 32.59 Gb (41%) - free 47.06 Gb (59%)
[1462994809] HOST DOWNTIME ALERT: NETENG-PA;STOPPED; Host has exited from a period of scheduled downtime
[1462994811] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462994813] SERVICE ALERT: BBTS_DB;Size: TS_DATA02.DBF;OK;SOFT;2;OK: Size is  13924.0078125  Mb
[1462994844] SERVICE ALERT: BBTS_DB;Size: TS_INDEX02.DBF;OK;SOFT;2;OK: Size is  100.0078125  Mb
[1462994931] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 91% (15 min average)
[1462995531] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1462995651] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1462995771] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 77% (15 min average)
[1462996260] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 78% (15 min average)
[1462996860] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462996980] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462997100] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462997700] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1462998171] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1462998292] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462998300] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462998411] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 89% (15 min average)
[1462998420] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462998502] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462998540] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1462998562] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462998622] SERVICE ALERT: hallefirst1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462998765] SERVICE ALERT: GOVGIS;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.102.118 and port 12489: No route to host
[1462998768] HOST ALERT: GOVGIS;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.118)
[1462998796] SERVICE ALERT: GOVGIS;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.102.118 and port 12489: No route to host
[1462998799] HOST ALERT: GOVGIS;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.118)
[1462998803] SERVICE ALERT: GOVGIS;Uptime;CRITICAL;SOFT;1;connect to address 164.76.102.118 and port 12489: No route to host
[1462998841] SERVICE ALERT: GOVGIS;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.102.118 and port 12489: No route to host
[1462998841] SERVICE ALERT: GOVGIS;CPU Load;CRITICAL;SOFT;1;connect to address 164.76.102.118 and port 12489: No route to host
[1462998841] SERVICE ALERT: GOVGIS;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.102.118 and port 12489: No route to host
[1462998844] HOST ALERT: GOVGIS;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.102.118)
[1462998876] SERVICE ALERT: GOVGIS;NSClient++ Version;CRITICAL;SOFT;2;connect to address 164.76.102.118 and port 12489: No route to host
[1462998879] HOST ALERT: GOVGIS;DOWN;SOFT;4;CRITICAL - Host Unreachable (164.76.102.118)
[1462998892] SERVICE ALERT: EMUSQLEXPRESS;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462998912] SERVICE ALERT: EMUSQLEXPRESS;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462998915] SERVICE ALERT: GOVGIS;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.102.118 and port 12489: No route to host
[1462998915] SERVICE ALERT: GOVGIS;Service: DNS Client;CRITICAL;SOFT;2;connect to address 164.76.102.118 and port 12489: No route to host
[1462998918] HOST ALERT: GOVGIS;DOWN;SOFT;5;CRITICAL - Host Unreachable (164.76.102.118)
[1462998922] HOST ALERT: EMUSQLEXPRESS;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462998922] SERVICE ALERT: EMUSCRIPT;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462998923] SERVICE ALERT: GOVGIS;Uptime;CRITICAL;SOFT;2;connect to address 164.76.102.118 and port 12489: No route to host
[1462998932] SERVICE ALERT: EMUSQLEXPRESS;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462998935] SERVICE ALERT: GOVGIS;Service: NetLogin;CRITICAL;SOFT;1;connect to address 164.76.102.118 and port 12489: No route to host
[1462998936] SERVICE ALERT: EMUSCRIPT;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462998938] HOST ALERT: GOVGIS;DOWN;SOFT;6;CRITICAL - Host Unreachable (164.76.102.118)
[1462998950] SERVICE ALERT: EMUSCRIPT;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462998952] SERVICE ALERT: EMUSQLEXPRESS;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462998952] HOST ALERT: EMUSCRIPT;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462998961] SERVICE ALERT: GOVGIS;Service: Windows Time;CRITICAL;SOFT;2;connect to address 164.76.102.118 and port 12489: No route to host
[1462998961] SERVICE ALERT: GOVGIS;CPU Load;CRITICAL;SOFT;2;connect to address 164.76.102.118 and port 12489: No route to host
[1462998961] SERVICE ALERT: GOVGIS;Service: Firewall;CRITICAL;SOFT;2;connect to address 164.76.102.118 and port 12489: No route to host
[1462998962] HOST ALERT: EMUSQLEXPRESS;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1462998964] HOST ALERT: GOVGIS;DOWN;SOFT;7;CRITICAL - Host Unreachable (164.76.102.118)
[1462998989] SERVICE ALERT: EMUSCRIPT;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462998993] SERVICE ALERT: EMUSQLEXPRESS;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462998996] SERVICE ALERT: GOVGIS;NSClient++ Version;CRITICAL;HARD;3;connect to address 164.76.102.118 and port 12489: No route to host
[1462998999] HOST ALERT: GOVGIS;DOWN;SOFT;8;CRITICAL - Host Unreachable (164.76.102.118)
[1462999002] SERVICE ALERT: GOVGIS;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.102.118 and port 12489: No route to host
[1462999002] SERVICE ALERT: EMUSQLEXPRESS;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462999010] SERVICE ALERT: EMUSQLEXPRESS;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462999011] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462999012] SERVICE ALERT: EMUSQLEXPRESS;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462999019] HOST ALERT: EMUSCRIPT;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1462999020] SERVICE ALERT: EMUSCRIPT;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462999023] HOST ALERT: EMUSQLEXPRESS;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1462999032] SERVICE ALERT: EMUSQLEXPRESS;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462999035] SERVICE ALERT: GOVGIS;Service: DNS Client;CRITICAL;HARD;3;connect to address 164.76.102.118 and port 12489: No route to host
[1462999035] SERVICE ALERT: GOVGIS;Service: Group Policy;CRITICAL;SOFT;2;connect to address 164.76.102.118 and port 12489: No route to host
[1462999036] HOST ALERT: DIRECTACCESS;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462999038] HOST ALERT: GOVGIS;DOWN;SOFT;9;CRITICAL - Host Unreachable (164.76.102.118)
[1462999042] SERVICE ALERT: EMUSCRIPT;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462999043] SERVICE ALERT: GOVGIS;Uptime;CRITICAL;HARD;3;connect to address 164.76.102.118 and port 12489: No route to host
[1462999044] SERVICE ALERT: EMUSCRIPT;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462999046] SERVICE ALERT: GOVGIS;Service: User Profile;CRITICAL;SOFT;1;connect to address 164.76.102.118 and port 12489: No route to host
[1462999050] SERVICE ALERT: EMUSQLEXPRESS;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462999050] HOST ALERT: EMUSCRIPT;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1462999052] SERVICE ALERT: EMUSQLEXPRESS;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462999055] SERVICE ALERT: GOVGIS;Service: NetLogin;CRITICAL;SOFT;2;connect to address 164.76.102.118 and port 12489: No route to host
[1462999056] SERVICE ALERT: EMUSCRIPT;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462999058] HOST ALERT: GOVGIS;DOWN;HARD;10;CRITICAL - Host Unreachable (164.76.102.118)
[1462999062] HOST ALERT: EMUSQLEXPRESS;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1462999064] SERVICE ALERT: EMUSCRIPT;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462999068] HOST ALERT: EMUSCRIPT;UP;SOFT;4;PING OK - Packet loss = 61%, RTA = 0.40 ms
[1462999070] SERVICE ALERT: EMUSCRIPT;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462999072] SERVICE ALERT: EMUSQLEXPRESS;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462999081] SERVICE ALERT: GOVGIS;Service: Firewall;CRITICAL;HARD;3;connect to address 164.76.102.118 and port 12489: No route to host
[1462999081] SERVICE ALERT: GOVGIS;Service: Windows Time;CRITICAL;HARD;3;connect to address 164.76.102.118 and port 12489: No route to host
[1462999081] SERVICE ALERT: GOVGIS;CPU Load;CRITICAL;HARD;3;connect to address 164.76.102.118 and port 12489: No route to host
[1462999084] SERVICE ALERT: GOVGIS;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.102.118 and port 12489: No route to host
[1462999089] SERVICE ALERT: EMUSCRIPT;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1462999100] HOST ALERT: DIRECTACCESS;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1462999102] SERVICE ALERT: EMUSQLEXPRESS;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462999102] HOST ALERT: EMUSQLEXPRESS;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1462999113] SERVICE ALERT: EMUSQLEXPRESS;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462999114] SERVICE ALERT: GOVGIS;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.102.118 and port 12489: No route to host
[1462999120] SERVICE ALERT: GOVGIS;Service: RPC;CRITICAL;SOFT;2;connect to address 164.76.102.118 and port 12489: No route to host
[1462999120] SERVICE ALERT: EMUSCRIPT;CPU Load;OK;SOFT;2;CPU Load 1% (15 min average)
[1462999122] SERVICE ALERT: EMUSQLEXPRESS;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462999130] SERVICE ALERT: EMUSQLEXPRESS;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462999131] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 92% (15 min average)
[1462999132] SERVICE ALERT: EMUSQLEXPRESS;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462999134] HOST DOWNTIME ALERT: GOVGIS;STARTED; Host has entered a period of scheduled downtime
[1462999140] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 79% (15 min average)
[1462999142] SERVICE ALERT: EMUSCRIPT;Uptime;OK;SOFT;3;System Uptime - 0 day(s) 0 hour(s) 1 minute(s)
[1462999143] HOST ALERT: EMUSQLEXPRESS;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1462999143] SERVICE ALERT: EMUSCRIPT;Memory Usage;OK;SOFT;2;Memory usage: total:10943.64 MB - used: 836.80 MB (8%) - free: 10106.84 MB (92%)
[1462999150] SERVICE ALERT: EMUSQLEXPRESS;CPU Load;OK;SOFT;2;CPU Load 28% (15 min average)
[1462999152] SERVICE ALERT: EMUSQLEXPRESS;Service: NetLogin;OK;SOFT;3;Netlogon: Started
[1462999152] SERVICE ALERT: EMUSQLEXPRESS;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462999155] SERVICE ALERT: GOVGIS;Service: Group Policy;CRITICAL;HARD;3;connect to address 164.76.102.118 and port 12489: No route to host
[1462999156] HOST ALERT: EMUSQLEXPRESS;UP;SOFT;7;PING OK - Packet loss = 28%, RTA = 0.37 ms
[1462999156] SERVICE ALERT: EMUSCRIPT;C:\ Drive Space;OK;SOFT;3;c:\ - total: 79.66 Gb - used: 20.56 Gb (26%) - free 59.10 Gb (74%)
[1462999163] SERVICE ALERT: EMUSCRIPT;Service: RPC;OK;SOFT;2;RpcSs: Started
[1462999166] SERVICE ALERT: GOVGIS;Service: User Profile;CRITICAL;SOFT;2;connect to address 164.76.102.118 and port 12489: No route to host
[1462999170] SERVICE ALERT: EMUSCRIPT;VMWare Tools;OK;SOFT;3;VMTools: Started
[1462999172] SERVICE ALERT: EMUSQLEXPRESS;Uptime;OK;SOFT;3;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1462999175] SERVICE ALERT: GOVGIS;Service: NetLogin;CRITICAL;HARD;3;connect to address 164.76.102.118 and port 12489: No route to host
[1462999202] SERVICE ALERT: EMUSQLEXPRESS;Memory Usage;OK;SOFT;2;Memory usage: total:10943.64 MB - used: 1875.52 MB (17%) - free: 9068.12 MB (83%)
[1462999205] SERVICE ALERT: GOVGIS;Service: WMI;CRITICAL;SOFT;2;connect to address 164.76.102.118 and port 12489: No route to host
[1462999213] SERVICE ALERT: EMUSQLEXPRESS;Service: WMI;OK;SOFT;3;Winmgmt: Started
[1462999222] SERVICE ALERT: EMUSQLEXPRESS;C:\ Drive Space;OK;SOFT;3;c:\ - total: 79.66 Gb - used: 25.04 Gb (31%) - free 54.61 Gb (69%)
[1462999222] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;Ready
[1462999227] SERVICE ALERT: GOVGIS;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.102.118 and port 12489: No route to host
[1462999227] SERVICE ALERT: GOVGIS;VMWare Tools;CRITICAL;SOFT;1;connect to address 164.76.102.118 and port 12489: No route to host
[1462999230] SERVICE ALERT: EMUSQLEXPRESS;VMWare Tools;OK;SOFT;3;VMTools: Started
[1462999234] SERVICE ALERT: GOVGIS;Memory Usage;CRITICAL;SOFT;2;connect to address 164.76.102.118 and port 12489: No route to host
[1462999242] SERVICE ALERT: GOVGIS;Service: RPC;CRITICAL;HARD;3;connect to address 164.76.102.118 and port 12489: No route to host
[1462999286] SERVICE ALERT: GOVGIS;Service: User Profile;CRITICAL;HARD;3;connect to address 164.76.102.118 and port 12489: No route to host
[1462999324] SERVICE ALERT: GOVGIS;Service: WMI;CRITICAL;HARD;3;connect to address 164.76.102.118 and port 12489: No route to host
[1462999331] HOST ALERT: loghost;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462999338] HOST ALERT: neteng;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462999340] HOST ALERT: oms;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462999347] SERVICE ALERT: GOVGIS;C:\ Drive Space;CRITICAL;SOFT;2;connect to address 164.76.102.118 and port 12489: No route to host
[1462999347] SERVICE ALERT: GOVGIS;VMWare Tools;CRITICAL;SOFT;2;connect to address 164.76.102.118 and port 12489: No route to host
[1462999353] SERVICE ALERT: GOVGIS;Memory Usage;CRITICAL;HARD;3;connect to address 164.76.102.118 and port 12489: No route to host
[1462999384] HOST ALERT: wwwstage1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462999421] HOST ALERT: loghost;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1462999428] HOST ALERT: neteng;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1462999430] HOST ALERT: oms;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1462999467] SERVICE ALERT: GOVGIS;C:\ Drive Space;CRITICAL;HARD;3;connect to address 164.76.102.118 and port 12489: No route to host
[1462999467] SERVICE ALERT: GOVGIS;VMWare Tools;CRITICAL;HARD;3;connect to address 164.76.102.118 and port 12489: No route to host
[1462999474] HOST ALERT: wwwstage1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1462999503] HOST ALERT: fanghorn;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462999511] HOST ALERT: loghost;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1462999518] HOST ALERT: neteng;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1462999520] HOST ALERT: oms;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1462999564] HOST ALERT: wwwstage1;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1462999567] HOST ALERT: fanghorn;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1462999583] HOST ALERT: loghost;UP;SOFT;4;PING OK - Packet loss = 61%, RTA = 0.35 ms
[1462999608] HOST ALERT: neteng;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1462999610] HOST ALERT: oms;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1462999628] HOST ALERT: wwwstage1;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1462999698] HOST ALERT: neteng;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1462999700] HOST ALERT: oms;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1462999712] SERVICE ALERT: EMUSQLEXPRESS;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1462999732] SERVICE ALERT: EMUSQLEXPRESS;Service: Windows Time;OK;HARD;3;W32Time: Started
[1462999788] HOST ALERT: neteng;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1462999790] HOST ALERT: oms;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1462999812] HOST DOWNTIME ALERT: TABLEAU;STOPPED; Host has exited from a period of scheduled downtime
[1462999852] HOST ALERT: neteng;UP;SOFT;7;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1462999859] HOST ALERT: oms;UP;SOFT;7;PING OK - Packet loss = 50%, RTA = 0.29 ms
[1463000940] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1463001060] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1463001180] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 66% (15 min average)
[1463001531] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463001651] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1463001753] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1463001771] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1463001780] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463001900] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 92% (15 min average)
[1463002042] SERVICE ALERT: ADRODC-LIV;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463002042] SERVICE ALERT: ADRODC-LIV;Service: Distributed File System Replication;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463002141] SERVICE ALERT: ADRODC-LIV;Uptime;OK;SOFT;2;System Uptime - 13 day(s) 9 hour(s) 6 minute(s)
[1463002142] SERVICE ALERT: ADRODC-LIV;Service: Distributed File System Replication;OK;SOFT;2;DFSR: Started
[1463002371] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 85% (15 min average)
[1463002500] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463002620] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463002740] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463005740] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 72% (15 min average)
[1463006340] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1463006460] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463006571] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1463006580] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1463006691] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 85% (15 min average)
[1463007780] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1463009580] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463009700] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463009820] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463010890] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1463011010] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1463011130] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 73% (15 min average)
[1463016421] SERVICE ALERT: TARGETX-2012;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1463016541] SERVICE ALERT: TARGETX-2012;CPU Load;OK;SOFT;2;CPU Load 6% (15 min average)
[1463019530] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1463019650] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 44% (15 min average)
[1463020020] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1463020200] HOST DOWNTIME ALERT: ADCA2;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: TI22;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: PNCAPP-2012;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: ADCA1;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: ADAUTHDC1-OLD;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: OPIE;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: ADAUTHDC1;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: ADAUTH-RMAD;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: OFFICEKMS;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: ADARC;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: AD-RMAD;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: TABLEAU;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: MBAMDB;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: MBAM;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: SP2010;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: LANSWEEPER;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: SP2010DB;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: IMAGES;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: SCCM-PRIM01;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: HDBACKUPS;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: FTK;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: SANDBOX-DB;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: Echo-01;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: EMUSQLEXPRESS;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: XDM;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: REFWIN;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: EMUSCRIPT;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: EMU-VALUE01;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: R25-ZETA2012;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: DPS-SECURE02;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: DIRECTACCESS;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: WINKMS;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: R25-WEBVIEWER;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: COE-WEBDAV;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: CEDARON;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: R25-WEBSERVICES;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: CBORD;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: ADWU1;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: UHSPHARM-2012;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: R25-INTERFACE;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: ADWSUSCLIENTS;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: ADDC55;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: NESSUS;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: R25-DELTA2012;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: ADDC5;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: ADDC44;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: TI22-PSYCLINIC;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: R25-ALPHA2012;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: NETENG-PA;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: ADDC4;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: ADDC33;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: PROCARE;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: ADDC3;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: ADDC22;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: TI22-COE;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: PNCWEB-2012;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: ADDC2;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: ADDC11;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: PNCDB-2012;STARTED; Host has entered a period of scheduled downtime
[1463020200] HOST DOWNTIME ALERT: ADDC1;STARTED; Host has entered a period of scheduled downtime
[1463020620] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463020850] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463020970] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463021090] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 83% (15 min average)
[1463021220] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463022130] SERVICE ALERT: R25-WEBSERVICES;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1463022250] SERVICE ALERT: R25-WEBSERVICES;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1463022271] SERVICE ALERT: XDM;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1463022370] SERVICE ALERT: PROCARE;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463022391] SERVICE ALERT: XDM;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1463022424] SERVICE ALERT: ADCA2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1463022490] SERVICE ALERT: PROCARE;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463022511] SERVICE ALERT: XDM;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463022541] SERVICE ALERT: NETENG-PA;Uptime;CRITICAL;SOFT;1;connect to address 164.76.4.142 and port 12489: Connection refused
[1463022544] SERVICE ALERT: ADCA2;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463022610] SERVICE ALERT: PROCARE;CPU Load;OK;SOFT;3;CPU Load 66% (15 min average)
[1463022641] SERVICE ALERT: TI22-COE;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1463022660] SERVICE ALERT: UHSPHARM-2012;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1463022661] SERVICE ALERT: NETENG-PA;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 2 minute(s)
[1463022664] SERVICE ALERT: ADCA2;CPU Load;OK;SOFT;3;CPU Load 23% (15 min average)
[1463022722] SERVICE ALERT: ADDC4;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.99.238 and port 12489: Connection refused
[1463022738] SERVICE ALERT: NESSUS;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.2.69 and port 12489: Connection refused
[1463022740] SERVICE ALERT: ADDC3;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463022743] SERVICE ALERT: ADDC3;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463022746] SERVICE ALERT: MBAM;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463022751] SERVICE ALERT: ADDC2;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1463022758] SERVICE ALERT: ADDC2;Service: DNS Service;CRITICAL;SOFT;1;DNS: Start pending
[1463022758] SERVICE ALERT: ADDC2;Service: Dell Backup Agent;CRITICAL;SOFT;1;ErdAgent: Start pending
[1463022758] SERVICE ALERT: ADDC1;Service: Dell Backup Agent;CRITICAL;SOFT;1;ErdAgent: Starting
[1463022761] SERVICE ALERT: TI22-COE;CPU Load;OK;SOFT;2;CPU Load 50% (15 min average)
[1463022780] SERVICE ALERT: UHSPHARM-2012;CPU Load;OK;SOFT;2;CPU Load 52% (15 min average)
[1463022791] SERVICE ALERT: REFWIN;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.158.184 and port 12489: Connection refused
[1463022815] SERVICE ALERT: PROCARE;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.158.143 and port 12489: Connection refused
[1463022818] SERVICE ALERT: PROCARE;VMWare Tools;CRITICAL;SOFT;1;connect to address 164.76.158.143 and port 12489: Connection refused
[1463022818] SERVICE ALERT: PROCARE;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.158.143 and port 12489: Connection refused
[1463022840] SERVICE ALERT: ADDC3;Memory Usage;OK;SOFT;2;Memory usage: total:20479.66 MB - used: 954.26 MB (5%) - free: 19525.40 MB (95%)
[1463022842] SERVICE ALERT: ADDC4;C:\ Drive Space;OK;SOFT;2;c:\ - total: 79.66 Gb - used: 17.54 Gb (22%) - free 62.12 Gb (78%)
[1463022843] SERVICE ALERT: ADDC3;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1463022845] HOST ALERT: PROCARE;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1463022846] SERVICE ALERT: MBAM;CPU Load;OK;SOFT;2;CPU Load 6% (15 min average)
[1463022858] SERVICE ALERT: NESSUS;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1463022871] SERVICE ALERT: ADDC2;Service: Windows Time;CRITICAL;SOFT;2;W32Time: Stopped
[1463022878] SERVICE ALERT: ADDC2;Service: DNS Service;OK;SOFT;2;DNS: Started
[1463022878] SERVICE ALERT: ADDC2;Service: Dell Backup Agent;OK;SOFT;2;ErdAgent: Started
[1463022878] SERVICE ALERT: ADDC1;Service: Dell Backup Agent;OK;SOFT;2;ErdAgent: Started
[1463022901] SERVICE ALERT: MBAMDB;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1463022903] SERVICE ALERT: XDM;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.2.79 and port 12489: Connection refused
[1463022905] HOST ALERT: PROCARE;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1463022910] SERVICE ALERT: REFWIN;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1463022926] SERVICE ALERT: PROCARE;C:\ Drive Space;OK;SOFT;2;c:\ - total: 49.90 Gb - used: 24.26 Gb (49%) - free 25.65 Gb (51%)
[1463022929] SERVICE ALERT: PROCARE;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1463022929] SERVICE ALERT: PROCARE;VMWare Tools;OK;SOFT;2;VMTools: Started
[1463022931] SERVICE ALERT: ADDC2;AD Replication;WARNING;SOFT;1;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1463022989] SERVICE ALERT: FTK;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463022991] SERVICE ALERT: ADDC2;Service: Windows Time;OK;SOFT;3;W32Time: Started
[1463022999] HOST ALERT: FTK;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1463023020] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463023021] SERVICE ALERT: MBAMDB;CPU Load;OK;SOFT;2;CPU Load 23% (15 min average)
[1463023022] SERVICE ALERT: ADDC33;Service: Intersite Messaging;CRITICAL;SOFT;1;connect to address 164.76.7.203 and port 12489: Connection refused
[1463023023] SERVICE ALERT: XDM;Memory Usage;OK;SOFT;2;Memory usage: total:7167.17 MB - used: 1048.45 MB (15%) - free: 6118.72 MB (85%)
[1463023031] SERVICE ALERT: ADDC55;Service: Distributed File System Replication;CRITICAL;SOFT;1;connect to address 164.76.7.205 and port 12489: Connection refused
[1463023033] SERVICE ALERT: R25-INTERFACE;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.4.82 and port 12489: Connection refused
[1463023034] SERVICE ALERT: ADDC55;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.7.205 and port 12489: Connection refused
[1463023034] SERVICE ALERT: ADDC1;AD Replication;WARNING;SOFT;1;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1463023035] SERVICE ALERT: R25-WEBVIEWER;Service: NetLogin;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1463023038] SERVICE ALERT: WINKMS;Service: Group Policy;CRITICAL;SOFT;1;gpsvc: Stopped
[1463023038] SERVICE ALERT: R25-WEBSERVICES;VMWare Tools;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1463023039] SERVICE ALERT: R25-WEBSERVICES;Service: Firewall;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1463023042] SERVICE ALERT: R25-INTERFACE;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.4.82 and port 12489: Connection refused
[1463023042] SERVICE ALERT: ADDC11;Service: AD Domain Service;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1463023042] SERVICE ALERT: ADDC11;Service: Group Policy;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1463023043] SERVICE ALERT: MBAMDB;Service: NetLogin;CRITICAL;SOFT;1;connect to address 164.76.7.97 and port 12489: Connection refused
[1463023048] SERVICE ALERT: ADDC22;Service: AD Web Service;CRITICAL;SOFT;1;connect to address 164.76.7.202 and port 12489: Connection refused
[1463023049] SERVICE ALERT: ADDC2;AD Replication;WARNING;SOFT;2;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1463023050] SERVICE ALERT: ADAUTHDC1;Service: Dell Backup Agent;CRITICAL;SOFT;1;connect to address 164.76.4.207 and port 12489: Connection refused
[1463023050] SERVICE ALERT: MBAMDB;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.7.97 and port 12489: Connection refused
[1463023073] HOST ALERT: MBAMDB;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1463023089] HOST ALERT: FTK;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1463023109] SERVICE ALERT: FTK;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463023111] SERVICE ALERT: XDM;CPU Load;OK;HARD;3;CPU Load 4% (15 min average)
[1463023111] SERVICE ALERT: R25-DELTA2012;Service: Windows Time;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1463023123] SERVICE ALERT: COE-WEBDAV;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463023126] SERVICE ALERT: FTK;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.7.103 and port 12489: Connection refused
[1463023131] SERVICE ALERT: R25-DELTA2012;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.4.177 and port 12489: Connection refused
[1463023133] SERVICE ALERT: R25-DELTA2012;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.4.177 and port 12489: Connection refused
[1463023134] SERVICE ALERT: ADDC3;AD Replication;WARNING;SOFT;1;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1463023137] HOST ALERT: MBAMDB;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1463023139] HOST ALERT: FTK;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1463023142] SERVICE ALERT: ADDC33;Service: Intersite Messaging;OK;SOFT;2;IsmServ: Started
[1463023143] SERVICE ALERT: ADDC11;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1463023151] SERVICE ALERT: ADDC55;Service: Distributed File System Replication;OK;SOFT;2;DFSR: Started
[1463023153] SERVICE ALERT: ADDC1;AD Replication;WARNING;SOFT;2;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1463023153] SERVICE ALERT: R25-INTERFACE;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1463023154] SERVICE ALERT: ADDC55;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1463023154] SERVICE ALERT: MBAMDB;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1463023155] SERVICE ALERT: R25-WEBVIEWER;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1463023158] SERVICE ALERT: EMUSCRIPT;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.4.107 and port 12489: Connection refused
[1463023158] SERVICE ALERT: R25-WEBSERVICES;VMWare Tools;OK;SOFT;2;VMTools: Started
[1463023158] SERVICE ALERT: WINKMS;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1463023159] SERVICE ALERT: R25-WEBSERVICES;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1463023162] SERVICE ALERT: ADDC11;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1463023162] SERVICE ALERT: ADDC11;Service: AD Domain Service;OK;SOFT;2;NTDS: Started
[1463023162] SERVICE ALERT: R25-INTERFACE;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1463023163] SERVICE ALERT: EMUSCRIPT;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.4.107 and port 12489: Connection refused
[1463023168] SERVICE ALERT: ADDC22;Service: AD Web Service;OK;SOFT;2;ADWS: Started
[1463023169] SERVICE ALERT: EMUSCRIPT;VMWare Tools;CRITICAL;SOFT;1;connect to address 164.76.4.107 and port 12489: Connection refused
[1463023169] SERVICE ALERT: ADDC2;AD Replication;WARNING;HARD;3;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1463023170] SERVICE ALERT: ADAUTHDC1;Service: Dell Backup Agent;CRITICAL;SOFT;2;ErdAgent: Starting
[1463023170] SERVICE ALERT: MBAMDB;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1463023170] SERVICE ALERT: ADWU1;Windows Update Services Event 10022;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1463023203] HOST ALERT: FTK;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1463023209] SERVICE ALERT: FTK;Service: Firewall;OK;SOFT;3;MpsSvc: Started
[1463023210] SERVICE ALERT: PROCARE;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1463023213] SERVICE ALERT: ADAUTHDC1;AD Replication;UNKNOWN;SOFT;1;CHECK_NRPE: Invalid packet type received from server.
[1463023223] SERVICE ALERT: COE-WEBDAV;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1463023231] SERVICE ALERT: R25-DELTA2012;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1463023246] SERVICE ALERT: FTK;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1463023248] SERVICE ALERT: R25-DELTA2012;C:\ Drive Space;OK;SOFT;2;c:\ - total: 39.66 Gb - used: 21.00 Gb (53%) - free 18.65 Gb (47%)
[1463023252] SERVICE ALERT: ADDC3;AD Replication;WARNING;SOFT;2;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1463023253] SERVICE ALERT: R25-DELTA2012;Service: RPC;OK;SOFT;2;RpcSs: Started
[1463023261] SERVICE ALERT: R25-ALPHA2012;CPU Load;CRITICAL;SOFT;1;connect to address 164.76.4.179 and port 12489: Connection refused
[1463023263] SERVICE ALERT: ADDC11;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1463023266] SERVICE ALERT: OPIE;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.158.60 and port 12489: Connection refused
[1463023272] SERVICE ALERT: ADDC1;AD Replication;WARNING;HARD;3;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1463023273] SERVICE ALERT: TI22;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1463023275] SERVICE ALERT: EMUSCRIPT;C:\ Drive Space;OK;SOFT;2;c:\ - total: 79.66 Gb - used: 20.75 Gb (26%) - free 58.90 Gb (74%)
[1463023283] SERVICE ALERT: EMUSCRIPT;Service: RPC;OK;SOFT;2;RpcSs: Started
[1463023289] SERVICE ALERT: EMUSCRIPT;VMWare Tools;OK;SOFT;2;VMTools: Started
[1463023290] SERVICE ALERT: ADAUTHDC1;Service: Dell Backup Agent;OK;SOFT;3;ErdAgent: Started
[1463023290] SERVICE ALERT: ADWU1;Windows Update Services Event 10022;OK;SOFT;2;OK: No entries found
[1463023296] SERVICE ALERT: PNCDB-2012;VMWare Tools;CRITICAL;SOFT;1;connect to address 164.76.2.176 and port 12489: Connection refused
[1463023313] SERVICE ALERT: SANDBOX-DB;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.99.176 and port 12489: Connection refused
[1463023321] SERVICE ALERT: ADAUTHDC1;AD Replication;UNKNOWN;SOFT;2;CHECK_NRPE: Invalid packet type received from server.
[1463023330] SERVICE ALERT: PROCARE;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1463023331] SERVICE ALERT: TI22-PSYCLINIC;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1463023372] SERVICE ALERT: ADDC3;AD Replication;WARNING;HARD;3;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1463023381] SERVICE ALERT: R25-ALPHA2012;CPU Load;OK;SOFT;2;CPU Load 66% (15 min average)
[1463023386] SERVICE ALERT: OPIE;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1463023390] SERVICE ALERT: EMUSQLEXPRESS;Service: User Profile;CRITICAL;SOFT;1;connect to address 164.76.7.169 and port 12489: Connection refused
[1463023393] SERVICE ALERT: TI22;CPU Load;OK;SOFT;2;CPU Load 49% (15 min average)
[1463023416] SERVICE ALERT: PNCDB-2012;VMWare Tools;OK;SOFT;2;VMTools: Started
[1463023433] SERVICE ALERT: SANDBOX-DB;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1463023441] SERVICE ALERT: ADAUTHDC1;AD Replication;UNKNOWN;HARD;3;CHECK_NRPE: Invalid packet type received from server.
[1463023446] SERVICE ALERT: MBAM;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1463023449] SERVICE ALERT: CBORD;Service: Group Policy;CRITICAL;SOFT;1;gpsvc: Stopped
[1463023450] SERVICE ALERT: PROCARE;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1463023451] SERVICE ALERT: TI22-PSYCLINIC;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1463023489] SERVICE ALERT: DPS-SECURE02;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463023490] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1463023508] SERVICE ALERT: DPS-SECURE02;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463023511] SERVICE ALERT: EMUSQLEXPRESS;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1463023519] HOST ALERT: DPS-SECURE02;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1463023521] SERVICE ALERT: DPS-SECURE02;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463023532] SERVICE ALERT: DPS-SECURE02;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463023532] SERVICE ALERT: DPS-SECURE02;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463023548] SERVICE ALERT: DPS-SECURE02;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463023551] HOST ALERT: DPS-SECURE02;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1463023563] SERVICE ALERT: DPS-SECURE02;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.2.46 and port 12489: Connection refused
[1463023566] SERVICE ALERT: MBAM;CPU Load;OK;SOFT;2;CPU Load 90% (15 min average)
[1463023569] SERVICE ALERT: CBORD;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1463023571] SERVICE ALERT: TI22-PSYCLINIC;CPU Load;OK;SOFT;3;CPU Load 51% (15 min average)
[1463023583] SERVICE ALERT: DPS-SECURE02;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.2.46 and port 12489: Connection refused
[1463023583] SERVICE ALERT: DPS-SECURE02;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.2.46 and port 12489: Connection refused
[1463023589] SERVICE ALERT: DPS-SECURE02;Service: RPC;CRITICAL;SOFT;2;connect to address 164.76.2.46 and port 12489: Connection refused
[1463023593] HOST ALERT: DPS-SECURE02;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1463023608] SERVICE ALERT: DPS-SECURE02;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1463023610] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1463023619] HOST ALERT: DPS-SECURE02;UP;SOFT;4;PING OK - Packet loss = 58%, RTA = 0.38 ms
[1463023620] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 79% (15 min average)
[1463023620] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1463023621] SERVICE ALERT: DPS-SECURE02;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1463023622] SERVICE ALERT: EMU-VALUE01;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463023632] SERVICE ALERT: DPS-SECURE02;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1463023632] SERVICE ALERT: DPS-SECURE02;CPU Load;OK;SOFT;2;CPU Load 2% (15 min average)
[1463023648] SERVICE ALERT: DPS-SECURE02;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1463023652] HOST ALERT: EMU-VALUE01;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1463023660] SERVICE ALERT: EMU-VALUE01;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463023662] SERVICE ALERT: EMU-VALUE01;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463023662] SERVICE ALERT: EMU-VALUE01;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463023669] SERVICE ALERT: EMU-VALUE01;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463023670] SERVICE ALERT: EMU-VALUE01;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463023674] SERVICE ALERT: DPS-SECURE02;Memory Usage;OK;SOFT;2;Memory usage: total:4790.25 MB - used: 862.07 MB (18%) - free: 3928.18 MB (82%)
[1463023678] SERVICE ALERT: EMU-VALUE01;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463023680] SERVICE ALERT: SP2010DB;Service: Firewall;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1463023686] SERVICE ALERT: EMU-VALUE01;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.4.85 and port 12489: Connection refused
[1463023690] SERVICE ALERT: EMU-VALUE01;CPU Load;CRITICAL;SOFT;1;connect to address 164.76.4.85 and port 12489: Connection refused
[1463023690] HOST ALERT: EMU-VALUE01;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1463023700] SERVICE ALERT: SP2010DB;Service: NetLogin;CRITICAL;SOFT;1;connect to address 164.76.4.123 and port 12489: Connection refused
[1463023700] SERVICE ALERT: EMU-VALUE01;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.4.85 and port 12489: Connection refused
[1463023702] SERVICE ALERT: EMU-VALUE01;Uptime;CRITICAL;SOFT;1;connect to address 164.76.4.85 and port 12489: Connection refused
[1463023703] SERVICE ALERT: DPS-SECURE02;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1463023703] SERVICE ALERT: DPS-SECURE02;C:\ Drive Space;OK;SOFT;2;c:\ - total: 278.53 Gb - used: 14.25 Gb (5%) - free 264.28 Gb (95%)
[1463023707] SERVICE ALERT: SP2010DB;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463023709] SERVICE ALERT: DPS-SECURE02;Service: RPC;OK;SOFT;3;RpcSs: Started
[1463023711] SERVICE ALERT: XDM;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1463023722] SERVICE ALERT: EMU-VALUE01;Memory Usage;CRITICAL;SOFT;2;connect to address 164.76.4.85 and port 12489: Connection refused
[1463023730] HOST ALERT: EMU-VALUE01;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1463023730] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1463023760] SERVICE ALERT: EMU-VALUE01;Service: WMI;CRITICAL;SOFT;2;connect to address 164.76.4.85 and port 12489: Connection refused
[1463023762] SERVICE ALERT: EMU-VALUE01;Service: Firewall;CRITICAL;SOFT;2;connect to address 164.76.4.85 and port 12489: Connection refused
[1463023762] SERVICE ALERT: EMU-VALUE01;NSClient++ Version;CRITICAL;SOFT;2;connect to address 164.76.4.85 and port 12489: Connection refused
[1463023769] SERVICE ALERT: EMU-VALUE01;C:\ Drive Space;CRITICAL;SOFT;2;connect to address 164.76.4.85 and port 12489: Connection refused
[1463023779] SERVICE ALERT: EMU-VALUE01;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1463023779] SERVICE ALERT: EMU-VALUE01;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1463023790] HOST ALERT: EMU-VALUE01;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1463023800] SERVICE ALERT: SP2010DB;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1463023801] SERVICE ALERT: EMU-VALUE01;CPU Load;OK;SOFT;2;CPU Load 19% (15 min average)
[1463023805] HOST ALERT: EMU-VALUE01;UP;SOFT;5;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1463023806] SERVICE ALERT: EMU-VALUE01;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1463023807] SERVICE ALERT: SP2010DB;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1463023820] SERVICE ALERT: SP2010DB;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1463023820] SERVICE ALERT: EMU-VALUE01;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1463023822] SERVICE ALERT: EMU-VALUE01;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 4 minute(s)
[1463023831] SERVICE ALERT: XDM;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1463023842] SERVICE ALERT: EMU-VALUE01;Memory Usage;OK;SOFT;3;Memory usage: total:9444.43 MB - used: 2126.07 MB (23%) - free: 7318.36 MB (77%)
[1463023880] SERVICE ALERT: EMU-VALUE01;Service: WMI;OK;SOFT;3;Winmgmt: Started
[1463023881] SERVICE ALERT: EMU-VALUE01;Service: Firewall;OK;SOFT;3;MpsSvc: Started
[1463023882] SERVICE ALERT: EMU-VALUE01;NSClient++ Version;OK;SOFT;3;NSClient++ 0.4.4.15 2015-11-25
[1463023889] SERVICE ALERT: EMU-VALUE01;C:\ Drive Space;OK;SOFT;3;c:\ - total: 465.42 Gb - used: 15.32 Gb (3%) - free 450.09 Gb (97%)
[1463023946] SERVICE ALERT: HDBACKUPS;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463023951] SERVICE ALERT: XDM;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1463023976] HOST ALERT: HDBACKUPS;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1463023992] HOST ALERT: HDBACKUPS;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1463024013] SERVICE ALERT: SANDBOX-DB;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463024046] SERVICE ALERT: HDBACKUPS;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1463024050] SERVICE ALERT: PROCARE;CPU Load;OK;HARD;3;CPU Load 0% (15 min average)
[1463024133] SERVICE ALERT: SANDBOX-DB;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463024219] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1463024253] SERVICE ALERT: SANDBOX-DB;CPU Load;OK;SOFT;3;CPU Load 6% (15 min average)
[1463024339] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463024459] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463024551] SERVICE ALERT: XDM;CPU Load;OK;HARD;3;CPU Load 0% (15 min average)
[1463024551] SERVICE FLAPPING ALERT: XDM;CPU Load;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1463024641] SERVICE ALERT: ADAUTHDC1;AD Replication;WARNING;HARD;3;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1463024930] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463025245] SERVICE ALERT: ADAUTHDC1;AD Replication;OK;HARD;3;OK: replication is up and running
[1463025530] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1463025574] SERVICE ALERT: ADDC2;AD Replication;OK;HARD;3;OK: replication is up and running