[1492660926] SERVICE FLAPPING ALERT: hallefirstcolor;Paper Tray Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1492660992] SERVICE ALERT: studentcenterpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1492661052] SERVICE ALERT: studentcenterpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1492661099] SERVICE ALERT: cava2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1492661112] SERVICE ALERT: studentcenterpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1492661159] SERVICE ALERT: cava2;CPU Load;OK;SOFT;2;CPU Load 6% (15 min average)
[1492661481] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1492661511] HOST ALERT: hallesecondsouth;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1492661541] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1492661549] HOST ALERT: hallesecondsouth;UP;SOFT;2;PING OK - Packet loss = 44%, RTA = 0.48 ms
[1492661596] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;OK;SOFT;3;Black Print Cartridge HP C8543X is at 34% - OK!
[1492661759] SERVICE ALERT: cava2;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1492661819] SERVICE ALERT: cava2;CPU Load;OK;SOFT;2;CPU Load 2% (15 min average)
[1492662095] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1492662313] SERVICE ALERT: studentcenterpk;Device Status;OK;HARD;3;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1492662402] SERVICE ALERT: halleground1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1492662600] HOST DOWNTIME ALERT: R25-INTERFACE;STOPPED; Host has exited from a period of scheduled downtime
[1492662600] HOST DOWNTIME ALERT: PNCAPP-2012;STOPPED; Host has exited from a period of scheduled downtime
[1492662600] HOST DOWNTIME ALERT: R25-DELTA2012;STOPPED; Host has exited from a period of scheduled downtime
[1492662600] HOST DOWNTIME ALERT: TI22-COE;STOPPED; Host has exited from a period of scheduled downtime
[1492662600] HOST DOWNTIME ALERT: PNCWEB-2012;STOPPED; Host has exited from a period of scheduled downtime
[1492662600] HOST DOWNTIME ALERT: R25-ALPHA2012;STOPPED; Host has exited from a period of scheduled downtime
[1492662600] HOST DOWNTIME ALERT: PROCARE;STOPPED; Host has exited from a period of scheduled downtime
[1492662600] HOST DOWNTIME ALERT: PNCDB-2012;STOPPED; Host has exited from a period of scheduled downtime
[1492662600] HOST DOWNTIME ALERT: XDM;STOPPED; Host has exited from a period of scheduled downtime
[1492662600] HOST DOWNTIME ALERT: FTK;STOPPED; Host has exited from a period of scheduled downtime
[1492662600] HOST DOWNTIME ALERT: REFWIN;STOPPED; Host has exited from a period of scheduled downtime
[1492662600] HOST DOWNTIME ALERT: OPIE;STOPPED; Host has exited from a period of scheduled downtime
[1492662600] HOST DOWNTIME ALERT: NETENG-PA;STOPPED; Host has exited from a period of scheduled downtime
[1492662600] HOST DOWNTIME ALERT: OFFICEKMS;STOPPED; Host has exited from a period of scheduled downtime
[1492662600] HOST DOWNTIME ALERT: WINKMS;STOPPED; Host has exited from a period of scheduled downtime
[1492662600] HOST DOWNTIME ALERT: R25-WEBSERVICES;STOPPED; Host has exited from a period of scheduled downtime
[1492662600] HOST DOWNTIME ALERT: R25-WEBVIEWER;STOPPED; Host has exited from a period of scheduled downtime
[1492662600] HOST DOWNTIME ALERT: NESSUS;STOPPED; Host has exited from a period of scheduled downtime
[1492662600] HOST DOWNTIME ALERT: HDBACKUPS;STOPPED; Host has exited from a period of scheduled downtime
[1492662600] HOST DOWNTIME ALERT: UHSPHARM-2012;STOPPED; Host has exited from a period of scheduled downtime
[1492662600] HOST DOWNTIME ALERT: TI22-PSYCLINIC;STOPPED; Host has exited from a period of scheduled downtime
[1492662600] HOST DOWNTIME ALERT: MBAMDB;STOPPED; Host has exited from a period of scheduled downtime
[1492662600] HOST DOWNTIME ALERT: MBAM;STOPPED; Host has exited from a period of scheduled downtime
[1492662600] HOST DOWNTIME ALERT: DPS-SECURE02;STOPPED; Host has exited from a period of scheduled downtime
[1492662600] HOST DOWNTIME ALERT: SANDBOX-DB;STOPPED; Host has exited from a period of scheduled downtime
[1492662600] HOST DOWNTIME ALERT: EMUSQLEXPRESS;STOPPED; Host has exited from a period of scheduled downtime
[1492662600] HOST DOWNTIME ALERT: TI22;STOPPED; Host has exited from a period of scheduled downtime
[1492662600] HOST DOWNTIME ALERT: EMUSCRIPT;STOPPED; Host has exited from a period of scheduled downtime
[1492662600] HOST DOWNTIME ALERT: TABLEAU;STOPPED; Host has exited from a period of scheduled downtime
[1492662600] HOST DOWNTIME ALERT: DIRECTACCESS;STOPPED; Host has exited from a period of scheduled downtime
[1492662600] HOST DOWNTIME ALERT: ADFS;STOPPED; Host has exited from a period of scheduled downtime
[1492662600] HOST DOWNTIME ALERT: COE-WEBDAV;STOPPED; Host has exited from a period of scheduled downtime
[1492662600] HOST DOWNTIME ALERT: CEDARON;STOPPED; Host has exited from a period of scheduled downtime
[1492662600] HOST DOWNTIME ALERT: CBORD;STOPPED; Host has exited from a period of scheduled downtime
[1492662600] HOST DOWNTIME ALERT: SCCM-PRIM01;STOPPED; Host has exited from a period of scheduled downtime
[1492662600] HOST DOWNTIME ALERT: 4WINDSMANAGER;STOPPED; Host has exited from a period of scheduled downtime
[1492662600] HOST DOWNTIME ALERT: R25-ZETA2012;STOPPED; Host has exited from a period of scheduled downtime
[1492662601] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1492663955] SERVICE ALERT: ADWU1;Service: NetLogin;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1492664015] SERVICE ALERT: ADWU1;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1492664265] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;1;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1492664385] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;2;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1492664495] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1492664505] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;HARD;3;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1492664555] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 74% (15 min average)
[1492664855] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1492664914] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1492665155] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1492665215] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 87% (15 min average)
[1492666114] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1492666174] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1492666414] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1492666474] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1492666534] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1492666775] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1492666834] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;2;Warning:  idm1  are more than 120 seconds behind!
[1492666893] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1492667734] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 80% (15 min average)
[1492668093] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1492668334] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1492668389] SERVICE ALERT: hallesecondsouth;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1492668394] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1492668419] HOST ALERT: hallesecondsouth;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1492668448] SERVICE ALERT: hallesecondsouth;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1492668453] HOST ALERT: hallesecondsouth;UP;SOFT;2;PING OK - Packet loss = 16%, RTA = 0.50 ms
[1492668454] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 77% (15 min average)
[1492668502] SERVICE ALERT: hallesecondsouth;Paper Tray Status;OK;SOFT;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1492669054] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1492669114] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1492669174] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1492669293] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1492670493] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1492670493] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1492671093] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1492671574] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (15 min average)
[1492671574] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1492671705] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;OK;HARD;3;OK: No entries found
[1492672293] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1492672774] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1492672834] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1492672893] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1492672894] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1492673493] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1492674094] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1492674094] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 91% (15 min average)
[1492674153] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1492674213] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1492674694] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1492674754] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1492674813] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1492674814] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1492675413] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1492675414] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1492676613] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1492677000] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service has exited from a period of scheduled downtime
[1492677003] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (41.4% change >= 20.0% threshold)
[1492677003] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (63.4% change >= 20.0% threshold)
[1492677003] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1492677003] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1492677003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1492677003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1492677003] HOST DOWNTIME ALERT: INTLAPP;STARTED; Host has entered a period of scheduled downtime
[1492677003] HOST DOWNTIME ALERT: INTLDB;STARTED; Host has entered a period of scheduled downtime
[1492677003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1492677003] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1492677003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1492677003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1492677003] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1492677003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1492677003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1492677003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1492677003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1492677003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1492677003] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1492677213] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1492678018] SERVICE ALERT: cava2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1492678078] SERVICE ALERT: cava2;CPU Load;OK;SOFT;2;CPU Load 5% (15 min average)
[1492678802] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (38.3% change >= 20.0% threshold)
[1492678802] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (64.7% change >= 20.0% threshold)
[1492678802] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1492678802] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1492678802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1492678802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1492678802] HOST DOWNTIME ALERT: INTLAPP;STARTED; Host has entered a period of scheduled downtime
[1492678802] HOST DOWNTIME ALERT: INTLDB;STARTED; Host has entered a period of scheduled downtime
[1492678802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1492678802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1492678802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1492678802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1492678802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1492678802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1492678802] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1492678802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1492678802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1492678802] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1492678802] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1492680813] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1492680873] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1492683214] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1492683274] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1492683334] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1492684452] SERVICE ALERT: eateriespk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1492684512] SERVICE ALERT: eateriespk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1492684572] SERVICE ALERT: eateriespk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1492688496] HOST ALERT: rackhampk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.47 ms
[1492688674] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1492688734] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 71% (15 min average)
[1492688799] SERVICE ALERT: owen101lab1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1492688803] HOST ALERT: owen101lab1;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.52 ms
[1492688832] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 42% - OK!
[1492688936] SERVICE ALERT: owen101lab1;Device Status;OK;HARD;3;Ready
[1492689696] HOST ALERT: rackhampk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1492689731] HOST ALERT: rackhampk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1492689821] HOST ALERT: rackhampk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1492689911] HOST ALERT: rackhampk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1492689992] HOST ALERT: rackhampk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1492690082] HOST ALERT: rackhampk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1492690172] HOST ALERT: rackhampk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1492690262] HOST ALERT: rackhampk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1492690324] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 1 Trays Empty
[1492690331] HOST ALERT: rackhampk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1492690375] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1492690378] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1492690390] SERVICE ALERT: hallefirstcolor;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1492690405] HOST ALERT: hallefirstcolor;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1492690421] HOST ALERT: rackhampk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1492690429] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;OK;SOFT;2;Image Transfer Kit HP CE249A is at 4% - OK!
[1492690433] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;OK;SOFT;2;
[1492690433] HOST ALERT: hallefirstcolor;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1492690444] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;SOFT;3;Printer Paper Is OK! All Trays Still Full
[1492690483] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Cartridge HP CE260X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1492690542] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Black Cartridge HP CE260X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1492690608] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1492690638] HOST ALERT: hallefirstcolor;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1492690668] SERVICE ALERT: hallefirstcolor;Toner Collection Unit HP CE265A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1492690698] HOST ALERT: hallefirstcolor;UP;SOFT;2;PING WARNING - Packet loss = 86%, RTA = 0.44 ms
[1492690698] HOST FLAPPING ALERT: hallefirstcolor;STARTED; Host appears to have started flapping (21.4% change > 20.0% threshold)
[1492690723] SERVICE ALERT: hallefirstcolor;Toner Collection Unit HP CE265A;OK;SOFT;2;
[1492691202] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;OK;HARD;3;Black Cartridge HP CE260X is at 57% - OK!
[1492691599] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1492691645] HOST FLAPPING ALERT: hallefirstcolor;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1492691659] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1492691719] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1492691734] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1492691794] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1492691819] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1492691828] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1492691849] HOST ALERT: prayharroldpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1492691854] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1492691879] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1492691883] HOST ALERT: prayharroldpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.59 ms
[1492691885] SERVICE ALERT: prayharroldpk;Device Status;OK;SOFT;2;Genuine HP cartridge installed
[1492691934] SERVICE ALERT: prayharroldpk;Paper Tray Status;OK;SOFT;3;Printer Paper Is OK! All Trays Still Full
[1492692060] SERVICE ALERT: hallegroundcolor;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1492692114] SERVICE ALERT: hallegroundcolor;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! All Trays Still Full
[1492692313] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 20% - OK!
[1492692758] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8011 found since last run. oneLevelSearchOps OK. searchOps is too high! 8019 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492692818] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492692933] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5778 found since last run. oneLevelSearchOps OK. searchOps is too high! 5996 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492692993] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492693054] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 50% (15 min average)
[1492693418] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7903 found since last run. oneLevelSearchOps OK. searchOps is too high! 7918 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492693478] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492693593] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6456 found since last run. oneLevelSearchOps OK. searchOps is too high! 6765 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492693653] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492694853] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6126 found since last run. oneLevelSearchOps OK. searchOps is too high! 6556 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492694913] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492695003] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (20.1% change >= 20.0% threshold)
[1492695003] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1492695003] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1492695003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1492695003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1492695003] HOST DOWNTIME ALERT: INTLAPP;STARTED; Host has entered a period of scheduled downtime
[1492695003] HOST DOWNTIME ALERT: INTLDB;STARTED; Host has entered a period of scheduled downtime
[1492695003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1492695003] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1492695003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1492695003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1492695003] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1492695003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1492695003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1492695003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1492695003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1492695003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1492695003] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1492695198] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1492695228] HOST ALERT: hallesecondsouth;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1492695256] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1492695260] HOST ALERT: hallesecondsouth;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1492695312] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;OK;SOFT;3;Black Print Cartridge HP C8543X is at 34% - OK!
[1492695362] SERVICE ALERT: mckennypk;Device Status;OK;HARD;3;RETRIEVE JOB  NO STORED JOBS
[1492695366] HOST ALERT: mckennypk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.63 ms
[1492695454] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1492695514] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1492695565] SERVICE ALERT: mckennypk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1492695574] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1492695652] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 52% - OK!
[1492696113] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6989 found since last run. oneLevelSearchOps OK. searchOps is too high! 7268 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492696173] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492696773] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5089 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492696774] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 84% (15 min average)
[1492696833] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492697077] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5163 found since last run. oneLevelSearchOps OK. searchOps is too high! 5172 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492697137] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492697433] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8379 found since last run. oneLevelSearchOps OK. searchOps is too high! 8750 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492697493] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492697974] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1492698034] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 89% (15 min average)
[1492698093] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5721 found since last run. oneLevelSearchOps OK. searchOps is too high! 5899 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492698153] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492698227] SERVICE ALERT: halleground2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1492698257] HOST ALERT: halleground2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1492698287] SERVICE ALERT: halleground2;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1492698294] HOST ALERT: halleground2;UP;SOFT;2;PING OK - Packet loss = 37%, RTA = 1.89 ms
[1492698337] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6296 found since last run. oneLevelSearchOps OK. searchOps is too high! 6301 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492698341] SERVICE ALERT: halleground2;Device Status;OK;SOFT;3;59.40 ERROR  To continue turn off then on 59.40 ERROR  46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1492698371] SERVICE ALERT: eateriespk;Device Status;OK;HARD;3;ORDER CARTRIDGE LESS THAN 1300 PAGES 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1492698397] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492698548] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1492698697] HOST ALERT: hallefirst2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1492698753] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7064 found since last run. oneLevelSearchOps OK. searchOps is too high! 7318 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492698761] HOST ALERT: hallefirst2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.34 ms
[1492698813] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492698997] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5830 found since last run. oneLevelSearchOps OK. searchOps is too high! 5838 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492699057] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492699147] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1492699208] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1492699268] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1492699413] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7458 found since last run. oneLevelSearchOps OK. searchOps is too high! 7706 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492699473] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492699518] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1492699572] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 34% - OK!
[1492699657] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9079 found since last run. oneLevelSearchOps OK. searchOps is too high! 9086 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492699717] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492700073] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10083 found since last run. oneLevelSearchOps OK. searchOps is too high! 10328 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492700133] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492700152] HOST ALERT: hallesecondsouth;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1492700178] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1492700208] HOST ALERT: hallesecondsouth;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1492700238] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1492700268] HOST ALERT: hallesecondsouth;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1492700298] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1492700307] SERVICE ALERT: hallesecondsouth;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1492700317] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5834 found since last run. oneLevelSearchOps OK. searchOps is too high! 5843 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492700328] HOST ALERT: hallesecondsouth;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1492700367] SERVICE ALERT: hallesecondsouth;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1492700374] SERVICE ALERT: hallesecondsouth;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1492700377] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492700397] HOST ALERT: hallesecondsouth;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1492700427] SERVICE ALERT: hallesecondsouth;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1492700434] SERVICE ALERT: hallesecondsouth;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1492700457] HOST ALERT: hallesecondsouth;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1492700494] SERVICE ALERT: hallesecondsouth;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1492700524] HOST ALERT: hallesecondsouth;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1492700614] HOST ALERT: hallesecondsouth;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1492700704] HOST ALERT: hallesecondsouth;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1492700733] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10023 found since last run. oneLevelSearchOps OK. searchOps is too high! 10123 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492700793] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492700794] HOST ALERT: hallesecondsouth;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1492701034] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1492701094] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 83% (15 min average)
[1492701304] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 99.90 Gb - used: 90.09 Gb (90%) - free 9.81 Gb (10%)
[1492701364] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 99.90 Gb - used: 90.09 Gb (90%) - free 9.81 Gb (10%)
[1492701393] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5603 found since last run. oneLevelSearchOps OK. searchOps is too high! 5766 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492701424] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 99.90 Gb - used: 90.10 Gb (90%) - free 9.80 Gb (10%)
[1492701428] HOST ALERT: hallesecondsouth;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.73 ms
[1492701453] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492701622] SERVICE ALERT: hallesecondsouth;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1492701688] SERVICE ALERT: hallesecondsouth;Device Status;OK;HARD;3;TRAY 3 EMPTY ANY TYPE LETTER   Processing job from tray 4
[1492701961] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1492702020] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1492702053] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7149 found since last run. oneLevelSearchOps OK. searchOps is too high! 7506 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492702080] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1492702092] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 34% - OK!
[1492702113] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492702713] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5076 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492702773] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492703374] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7018 found since last run. oneLevelSearchOps OK. searchOps is too high! 7401 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492703377] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5332 found since last run. oneLevelSearchOps OK. searchOps is too high! 5340 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492703433] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492703437] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492703613] SERVICE ALERT: ADWSUSSERVERS;Service: Group Policy;CRITICAL;SOFT;1;gpsvc: Stopped
[1492703632] SERVICE ALERT: ADWSUSSERVERS;E:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.4.79 and port 12489: Connection refused
[1492703634] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 364;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1492703639] SERVICE ALERT: ADWSUSSERVERS;Service: User Profile;CRITICAL;SOFT;1;connect to address 164.76.4.79 and port 12489: Connection refused
[1492703673] SERVICE ALERT: ADWSUSSERVERS;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1492703683] SERVICE ALERT: ADWSUSSERVERS;E:\ Drive Space;OK;SOFT;2;e:\ - total: 149.87 Gb - used: 106.75 Gb (71%) - free 43.12 Gb (29%)
[1492703699] SERVICE ALERT: ADWSUSSERVERS;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1492703745] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 364;OK;SOFT;2;Eventlog check ok
[1492704003] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1492704003] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1492704003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1492704003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1492704003] HOST DOWNTIME ALERT: INTLAPP;STARTED; Host has entered a period of scheduled downtime
[1492704003] HOST DOWNTIME ALERT: INTLDB;STARTED; Host has entered a period of scheduled downtime
[1492704003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1492704003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1492704003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1492704003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1492704003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1492704003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1492704003] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1492704003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1492704003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1492704003] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1492704003] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1492704033] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7526 found since last run. oneLevelSearchOps OK. searchOps is too high! 7773 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492704093] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492704203] SERVICE ALERT: CAVA2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1492704263] SERVICE ALERT: CAVA2;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1492704323] SERVICE ALERT: CAVA2;CPU Load;OK;SOFT;3;CPU Load 80% (15 min average)
[1492704693] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 11302 found since last run. oneLevelSearchOps OK. searchOps is too high! 11670 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492704753] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492705237] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6604 found since last run. oneLevelSearchOps OK. searchOps is too high! 6613 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492705294] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1492705296] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492705353] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8737 found since last run. oneLevelSearchOps OK. searchOps is too high! 9140 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492705353] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1492705413] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492705413] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1492705533] SERVICE ALERT: hallefirst3;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1492705593] SERVICE ALERT: hallefirst3;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1492705653] SERVICE ALERT: hallefirst3;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1492705680] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1492705752] SERVICE ALERT: sillpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1492705802] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1492705802] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1492705802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1492705802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1492705802] HOST DOWNTIME ALERT: INTLAPP;STARTED; Host has entered a period of scheduled downtime
[1492705802] HOST DOWNTIME ALERT: INTLDB;STARTED; Host has entered a period of scheduled downtime
[1492705802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1492705802] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1492705802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1492705802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1492705802] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1492705802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1492705802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1492705802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1492705802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1492705802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1492705802] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1492705812] SERVICE ALERT: sillpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1492705872] SERVICE ALERT: sillpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1492705896] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6683 found since last run. oneLevelSearchOps OK. searchOps is too high! 6697 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492705956] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492706013] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1492706013] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1492706013] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7718 found since last run. oneLevelSearchOps OK. searchOps is too high! 8028 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492706073] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492706206] HOST ALERT: sillpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1492706286] SERVICE ALERT: sillpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1492706296] HOST ALERT: sillpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1492706346] SERVICE ALERT: sillpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1492706376] HOST ALERT: sillpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1492706406] SERVICE ALERT: sillpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1492706436] HOST ALERT: sillpk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1492706508] HOST ALERT: sillpk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1492706556] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5044 found since last run. oneLevelSearchOps OK. searchOps is too high! 5051 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492706598] HOST ALERT: sillpk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1492706616] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492706673] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5281 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492706686] SERVICE ALERT: sillpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1492706688] HOST ALERT: sillpk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1492706733] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492706746] SERVICE ALERT: sillpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1492706770] SERVICE ALERT: cava2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1492706776] HOST ALERT: sillpk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1492706799] SERVICE ALERT: halleground1;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1492706806] SERVICE ALERT: sillpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1492706830] SERVICE ALERT: cava2;CPU Load;OK;SOFT;2;CPU Load 2% (15 min average)
[1492706836] HOST ALERT: sillpk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1492706859] SERVICE ALERT: halleground1;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1492706919] SERVICE ALERT: halleground1;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1492706926] HOST ALERT: sillpk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1492707216] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5618 found since last run. oneLevelSearchOps OK. searchOps is too high! 5626 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492707243] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 1 Trays Empty
[1492707276] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492707303] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 1 Trays Empty
[1492707364] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 1 Trays Empty
[1492707876] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7795 found since last run. oneLevelSearchOps OK. searchOps is too high! 7803 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492707933] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5354 found since last run. oneLevelSearchOps OK. searchOps is too high! 5617 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492707936] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492707969] HOST ALERT: irimdevl;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1492707993] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492708023] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;OK;HARD;3;c:\ - total: 99.90 Gb - used: 89.72 Gb (90%) - free 10.18 Gb (10%)
[1492708033] HOST ALERT: irimdevl;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1492708052] SERVICE ALERT: hallefirst3;Device Status;OK;HARD;3;Open right  door.
[1492708593] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8305 found since last run. oneLevelSearchOps OK. searchOps is too high! 8690 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492708653] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492709013] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1492709136] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7325 found since last run. oneLevelSearchOps OK. searchOps is too high! 7345 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492709162] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1492709196] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492709253] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 5137 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492709314] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492709913] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5521 found since last run. oneLevelSearchOps OK. searchOps is too high! 5695 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492709973] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492710213] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1492711173] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7672 found since last run. oneLevelSearchOps OK. searchOps is too high! 7871 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492711233] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492711412] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 86% (15 min average)
[1492711833] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10447 found since last run. oneLevelSearchOps OK. searchOps is too high! 10574 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492711893] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492712318] SERVICE ALERT: halleground1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1492712493] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8444 found since last run. oneLevelSearchOps OK. searchOps is too high! 8682 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492712553] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492712796] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6429 found since last run. oneLevelSearchOps OK. searchOps is too high! 6558 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492712856] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492713211] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1492713271] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1492713331] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1492713456] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6292 found since last run. oneLevelSearchOps OK. searchOps is too high! 6417 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492713516] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492713931] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1492714000] SERVICE ALERT: sillpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1492714004] HOST ALERT: sillpk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.47 ms
[1492714052] SERVICE ALERT: hallefirst3;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1492714112] SERVICE ALERT: hallefirst3;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1492714116] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7950 found since last run. oneLevelSearchOps OK. searchOps is too high! 8065 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492714172] SERVICE ALERT: hallefirst3;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1492714176] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492714199] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;ORDER CARTRIDGE LESS THAN 2100 PAGES 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Processing job from tray 2
[1492714272] SERVICE ALERT: sillpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 6% - OK!
[1492714353] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5301 found since last run. oneLevelSearchOps OK. searchOps is too high! 5323 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492714413] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492714868] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1492715013] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5327 found since last run. oneLevelSearchOps OK. searchOps is too high! 5346 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492715073] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492715131] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 85% (15 min average)
[1492715372] SERVICE ALERT: hallefirst3;Device Status;OK;HARD;3;Clearing paper path
[1492715372] SERVICE FLAPPING ALERT: hallefirst3;Device Status;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1492715875] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1492715918] SERVICE ALERT: halleground1;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1492715934] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1492715976] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5323 found since last run. oneLevelSearchOps OK. searchOps is too high! 5442 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492715977] SERVICE ALERT: halleground1;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1492715994] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1492716036] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492716037] SERVICE ALERT: halleground1;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1492716331] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1492716391] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1492716451] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1492716476] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1492716530] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 48% - OK!
[1492716594] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Ready
[1492716873] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6867 found since last run. oneLevelSearchOps OK. searchOps is too high! 7034 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492716933] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492717051] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1492717194] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1492717254] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1492717314] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1492717772] SERVICE ALERT: hallefirst3;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1492717832] SERVICE ALERT: hallefirst3;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1492717835] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8080 found since last run. oneLevelSearchOps OK. searchOps is too high! 8171 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492717893] SERVICE ALERT: hallefirst3;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1492717896] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492718468] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1492718492] SERVICE ALERT: hallefirst3;Device Status;OK;HARD;3;Open  tray 4. Remove any  jammed
[1492718495] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5593 found since last run. oneLevelSearchOps OK. searchOps is too high! 5617 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492718527] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1492718554] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492718588] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1492718722] HOST ALERT: rackhampk;UP;HARD;10;PING OK - Packet loss = 66%, RTA = 0.32 ms
[1492718848] HOST ALERT: hallefirst2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1492718912] HOST ALERT: hallefirst2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1492719114] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Ready
[1492719114] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1492719729] HOST ALERT: rackhampk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1492719819] HOST ALERT: rackhampk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1492719909] HOST ALERT: rackhampk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1492719989] HOST ALERT: rackhampk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1492720051] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 52% (15 min average)
[1492720079] HOST ALERT: rackhampk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1492720169] HOST ALERT: rackhampk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1492720259] HOST ALERT: rackhampk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1492720304] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;1;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1492720329] HOST ALERT: rackhampk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1492720419] HOST ALERT: rackhampk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1492720424] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;2;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1492720509] HOST ALERT: rackhampk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1492720544] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;HARD;3;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1492720651] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1492720711] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 92% (15 min average)
[1492720954] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6100 found since last run. oneLevelSearchOps OK. searchOps is too high! 6209 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492721014] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492721230] SERVICE ALERT: cava2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1492721290] SERVICE ALERT: cava2;CPU Load;OK;SOFT;2;CPU Load 1% (15 min average)
[1492722036] SERVICE ALERT: halleground1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1492722333] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5709 found since last run. oneLevelSearchOps OK. searchOps is too high! 5911 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492722339] SERVICE ALERT: porterpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1492722393] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492722399] SERVICE ALERT: porterpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1492722459] SERVICE ALERT: porterpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1492722511] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1492722571] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1492722631] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1492722637] SERVICE ALERT: halleground1;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1492722697] SERVICE ALERT: halleground1;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! 1 of 3 Trays Still Full
[1492723245] HOST ALERT: halleground2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1492723297] SERVICE ALERT: halleground1;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1492723309] HOST ALERT: halleground2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1492723357] SERVICE ALERT: halleground1;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1492723417] SERVICE ALERT: halleground1;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1492723417] SERVICE FLAPPING ALERT: halleground1;Paper Tray Status;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1492723593] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5010 found since last run. oneLevelSearchOps OK. searchOps is too high! 5084 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492723653] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492724141] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1492724200] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1492724253] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5336 found since last run. oneLevelSearchOps OK. searchOps is too high! 5425 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492724259] SERVICE ALERT: halleground2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1492724260] SERVICE ALERT: porterpk;Device Status;OK;HARD;3;Warming up
[1492724313] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492724913] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5573 found since last run. oneLevelSearchOps OK. searchOps is too high! 5625 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492724973] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492726059] SERVICE ALERT: halleground2;Device Status;OK;HARD;3;Clearing paper path
[1492726173] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5713 found since last run. oneLevelSearchOps OK. searchOps is too high! 5918 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492726233] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492726416] SERVICE ALERT: halleground1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1492726659] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1492726719] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1492726780] SERVICE ALERT: halleground2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1492727144] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;OK;HARD;3;OK: No entries found
[1492728030] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 86% (15 min average)
[1492728630] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1492728690] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1492728750] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1492729233] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5054 found since last run. oneLevelSearchOps OK. searchOps is too high! 5182 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492729293] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492729313] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1492729350] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 81% (15 min average)
[1492729892] SERVICE FLAPPING ALERT: hallefirst3;Device Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1492729893] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6244 found since last run. oneLevelSearchOps OK. searchOps is too high! 6318 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492729953] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492730550] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1492730610] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1492730670] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1492731153] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5553 found since last run. oneLevelSearchOps OK. searchOps is too high! 5743 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1492731213] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1492732033] HOST ALERT: prayharroldpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1492732084] HOST ALERT: prayharroldpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1492733070] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1492733670] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1492734270] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1492734330] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1492734388] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 90% (15 min average)
[1492735416] SERVICE FLAPPING ALERT: halleground1;Paper Tray Status;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1492735588] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1492735648] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1492735708] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1492736308] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 64% (15 min average)
[1492736908] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1492736968] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1492737028] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1492738200] HOST DOWNTIME ALERT: IMAGES;STARTED; Host has entered a period of scheduled downtime
[1492738228] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1492738828] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1492738888] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1492738892] SERVICE ALERT: hallefirst3;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1492738948] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1492738952] SERVICE ALERT: hallefirst3;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1492739013] SERVICE ALERT: hallefirst3;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1492739162] SERVICE ALERT: IMAGES;MDT Monitor;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1492739192] HOST ALERT: IMAGES;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1492739282] HOST ALERT: IMAGES;DOWN;HARD;2;PING CRITICAL - Packet loss = 100%
[1492739312] SERVICE ALERT: IMAGES;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1492739314] HOST ALERT: IMAGES;UP;HARD;2;PING OK - Packet loss = 0%, RTA = 0.23 ms
[1492739352] SERVICE ALERT: IMAGES;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1492739444] SERVICE ALERT: IMAGES;Service: PXE Boot;CRITICAL;SOFT;1;WDSServer: Starting
[1492739504] SERVICE ALERT: IMAGES;Service: PXE Boot;OK;SOFT;2;WDSServer: Started
[1492739548] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 63% (15 min average)
[1492739742] SERVICE ALERT: IMAGES;MDT Monitor;UNKNOWN;HARD;3;Failed to open service MDT_Monitor: 1060: The specified service does not exist as an installed service.
[1492740035] SERVICE ALERT: SANDBOX-DB;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1492740095] SERVICE ALERT: SANDBOX-DB;CPU Load;OK;SOFT;2;CPU Load 89% (15 min average)
[1492742548] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1492742608] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1492742668] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1492745400] HOST DOWNTIME ALERT: IMAGES;STOPPED; Host has exited from a period of scheduled downtime