[1461297615] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 54% (15 min average)
[1461297856] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461299056] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461299656] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461300856] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 80% (15 min average)
[1461301456] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461301576] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461301696] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 76% (15 min average)
[1461302296] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461302416] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461302536] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461303136] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461304336] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 77% (15 min average)
[1461304936] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461305056] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1461305176] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461305776] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461308776] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 73% (15 min average)
[1461309376] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461309496] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461309616] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461310816] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461311416] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461312016] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461312616] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 88% (15 min average)
[1461315016] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461315134] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461315253] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1461318132] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1461318192] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1461318253] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;SOFT;3;Black Print Cartridge HP C8543X is at 12% - OK!
[1461318253] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 91% (15 min average)
[1461320053] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1461320173] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1461320293] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461320893] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461321493] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1461324083] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1461324142] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1461324203] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1461325093] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1461326562] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 798 found since last run. oneLevelSearchOps OK. searchOps is too high! 799 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 837 found since last run.
[1461326662] SERVICE ALERT: eateriespk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461326683] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1461326722] SERVICE ALERT: eateriespk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461326782] SERVICE ALERT: eateriespk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461326893] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1461327013] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461327133] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461327262] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461327322] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461327383] SERVICE ALERT: hallefirst1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461327982] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;Open front door.
[1461328178] HOST ALERT: crossroadspk;UP;HARD;10;PING WARNING - Packet loss = 93%, RTA = 0.47 ms
[1461328192] SERVICE ALERT: crossroadspk;Device Status;OK;HARD;3;Waiting for tray 2 to lift   Ready
[1461328523] SERVICE ALERT: crossroadspk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1461328582] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461328642] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461328702] SERVICE ALERT: hallefirst1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461328742] SERVICE ALERT: crossroadspk;Black Cartridge HP CC364X;OK;HARD;3;Black Cartridge HP CC364X is at 90% - OK!
[1461329302] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;Ready
[1461329302] SERVICE FLAPPING ALERT: hallefirst1;Device Status;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1461329533] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 87% (15 min average)
[1461330733] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1461330853] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461330973] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1461331573] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 82% (15 min average)
[1461332173] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1461332293] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1461332413] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461332603] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1461333013] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (15 min average)
[1461333282] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 698 found since last run. oneLevelSearchOps OK. searchOps is too high! 699 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 716 found since last run.
[1461333402] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1461333613] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1461333733] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461333852] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1461333853] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461333912] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 12% - OK!
[1461334002] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 1137 found since last run. oneLevelSearchOps OK. searchOps is too high! 1138 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 886 found since last run.
[1461334122] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 395 found since last run. oneLevelSearchOps OK. searchOps is too high! 396 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1461334242] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1461335588] SERVICE ALERT: halleground2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1461335648] SERVICE ALERT: halleground2;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1461335653] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1461335708] SERVICE ALERT: halleground2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1461335861] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1461335902] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1461335921] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1461335962] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1461335981] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1461336022] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1461336183] SERVICE ALERT: hallefirst2;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1461336243] SERVICE ALERT: hallefirst2;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1461336302] SERVICE ALERT: hallefirst2;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1461336853] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1461336903] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1461336973] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461337093] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1461337372] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461337432] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461337492] SERVICE ALERT: mckennypk;Device Status;OK;SOFT;3;Warming up
[1461337702] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461337762] SERVICE ALERT: hallefirst1;Device Status;OK;SOFT;2;RETRIEVE JOB  NO STORED JOBS
[1461338098] SERVICE ALERT: mckennypk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1461338118] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1461338158] SERVICE ALERT: mckennypk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1461338178] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1461338218] SERVICE ALERT: mckennypk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1461338238] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1461338322] SERVICE ALERT: mckennypk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1461338362] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461338382] SERVICE ALERT: mckennypk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1461338422] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461338442] SERVICE ALERT: mckennypk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1461338482] SERVICE ALERT: hallefirst1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461339082] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;50 72 65 73 73 20 17 20 74 6F 20 73 65 65 20 73  74 65 70  by step information
[1461339682] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461339742] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461339803] SERVICE ALERT: hallefirst1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461339938] HOST ALERT: halleground2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1461340002] HOST ALERT: halleground2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.55 ms
[1461340093] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1461340174] SERVICE ALERT: halleground2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1461340216] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 9% - OK!
[1461340402] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;TRAY 4 EMPTY PLAIN LETTER   Ready
[1461340503] SERVICE ALERT: halleground2;Device Status;OK;HARD;3;Ready
[1461340612] SERVICE ALERT: mckennypk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461340632] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 90% - OK!
[1461340693] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461340813] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461340837] SERVICE ALERT: mckennypk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1461340933] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461342733] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461342920] HOST DOWNTIME ALERT: idm1.emich.edu;STARTED; Host has entered a period of scheduled downtime
[1461342920] HOST DOWNTIME ALERT: idm2.emich.edu;STARTED; Host has entered a period of scheduled downtime
[1461342920] HOST DOWNTIME ALERT: idm3.emich.edu;STARTED; Host has entered a period of scheduled downtime
[1461342920] HOST DOWNTIME ALERT: idm4.emich.edu;STARTED; Host has entered a period of scheduled downtime
[1461342920] HOST DOWNTIME ALERT: idmrbs.emich.edu;STARTED; Host has entered a period of scheduled downtime
[1461342920] HOST DOWNTIME ALERT: idmrbstest.emich.edu;STARTED; Host has entered a period of scheduled downtime
[1461343333] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461343572] SERVICE ALERT: prayharroldpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461343632] SERVICE ALERT: prayharroldpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461343693] SERVICE ALERT: prayharroldpk;Device Status;OK;SOFT;3;Open right  door.
[1461344127] SERVICE ALERT: idm3.emich.edu;LDAP Stats;UNKNOWN;SOFT;1;Difference is negative. LDAP recently rebooted.
[1461344247] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1461344346] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1461345262] SERVICE ALERT: idm1.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 4
[1461345293] SERVICE ALERT: idm2.emich.edu;NSDS Errors;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1461345312] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 7
[1461345312] SERVICE ALERT: idm2.emich.edu;eDirectory Service;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1461345319] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 7
[1461345383] SERVICE ALERT: idm1.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 7
[1461345432] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 7
[1461345433] SERVICE ALERT: idm2.emich.edu;eDirectory Service;OK;SOFT;2;OK: ndsd status is good.
[1461345439] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 7
[1461345452] SERVICE ALERT: idm2.emich.edu;NSDS Errors;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 60 seconds.
[1461345502] SERVICE ALERT: idm1.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 4
[1461345552] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 2
[1461345559] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 2
[1461345572] SERVICE ALERT: idm2.emich.edu;NSDS Errors;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 60 seconds.
[1461345733] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461345922] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1461345983] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1461346042] SERVICE ALERT: crossroadspk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1461346102] SERVICE ALERT: idm1.emich.edu;NSDS Errors;OK;HARD;3;OK: 0 Errors Reported
[1461346112] SERVICE ALERT: idm2.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1461346152] SERVICE ALERT: idm4.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1461346159] SERVICE ALERT: idm3.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1461346611] SERVICE ALERT: mckennypk;Device Status;OK;HARD;3;Genuine HP cartridge installed
[1461346712] SERVICE ALERT: idm2.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 4
[1461346752] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 5
[1461346759] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 6
[1461346832] SERVICE ALERT: idm2.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 7
[1461346849] SERVICE ALERT: owen101lab3;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461346852] SERVICE ALERT: idm1.emich.edu;LDAP Stats;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1461346872] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 7
[1461346879] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 7
[1461346882] HOST ALERT: idm1.emich.edu;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1461346895] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;CRITICAL;SOFT;1;Unknown: Error reported?
[1461346899] HOST ALERT: idm1.emich.edu;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1461346908] SERVICE ALERT: owen101lab3;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461346932] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 90% (15 min average)
[1461346952] SERVICE ALERT: idm2.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 16
[1461346961] SERVICE ALERT: idm1.emich.edu;LDAP Stats;CRITICAL;SOFT;2;CRITICAL: Failed to connect to LDAP!
[1461346968] SERVICE ALERT: owen101lab3;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461346992] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 11
[1461346999] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 16
[1461347012] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;CRITICAL;SOFT;2;CRITICAL: cn=ADAUTH Driver,cn=driverset1,o=system is stopped!
[1461347080] SERVICE ALERT: idm1.emich.edu;Loopback Driver;CRITICAL;SOFT;1;CRITICAL: cn=Generic Loopback,cn=driverset1,o=system is stopped!
[1461347080] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1461347133] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;CRITICAL;HARD;3;CRITICAL: cn=ADAUTH Driver,cn=driverset1,o=system is stopped!
[1461347200] SERVICE ALERT: idm1.emich.edu;Loopback Driver;OK;SOFT;2;0
[1461347211] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461347272] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461347331] SERVICE ALERT: mckennypk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461347331] SERVICE FLAPPING ALERT: mckennypk;Device Status;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1461347532] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461347552] SERVICE ALERT: idm2.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1461347552] SERVICE FLAPPING ALERT: idm2.emich.edu;NSDS Errors;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1461347592] SERVICE ALERT: idm4.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1461347592] SERVICE FLAPPING ALERT: idm4.emich.edu;NSDS Errors;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1461347599] SERVICE ALERT: idm3.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1461347599] SERVICE FLAPPING ALERT: idm3.emich.edu;NSDS Errors;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1461347652] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1461347733] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;OK;HARD;3;0
[1461347772] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461347842] SERVICE ALERT: crossroadspk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1461347930] SERVICE ALERT: mckennypk;Device Status;OK;HARD;3;Ready
[1461348372] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1461349503] SERVICE ALERT: hallefirst2;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1461349562] SERVICE ALERT: hallefirst2;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1461349572] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461349623] SERVICE ALERT: hallefirst2;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1461350172] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461350222] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1461351802] SERVICE FLAPPING ALERT: hallefirst1;Device Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1461353172] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461353556] SERVICE ALERT: hallegroundcolor;Cyan Cartridge HP CE261A;OK;HARD;3;
[1461353772] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461353966] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1461354026] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1461354087] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1461354372] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1461354478] SERVICE ALERT: prayharroldpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1461354538] SERVICE ALERT: prayharroldpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1461354597] SERVICE ALERT: prayharroldpk;Paper Tray Status;OK;SOFT;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1461354872] HOST ALERT: prayharroldpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1461354936] HOST ALERT: prayharroldpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.56 ms
[1461355092] SERVICE ALERT: prayharroldpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461355152] SERVICE ALERT: prayharroldpk;Device Status;OK;SOFT;2;Open front door.
[1461356172] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1461356292] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1461356412] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461357202] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461357262] SERVICE ALERT: hallefirst1;Device Status;OK;SOFT;2;Move left accessory, if attached.
[1461357313] HOST DOWNTIME ALERT: idm2.emich.edu;STOPPED; Host has exited from a period of scheduled downtime
[1461357313] HOST DOWNTIME ALERT: idm1.emich.edu;STOPPED; Host has exited from a period of scheduled downtime
[1461357313] HOST DOWNTIME ALERT: idm3.emich.edu;STOPPED; Host has exited from a period of scheduled downtime
[1461357313] HOST DOWNTIME ALERT: idmrbstest.emich.edu;STOPPED; Host has exited from a period of scheduled downtime
[1461357313] HOST DOWNTIME ALERT: idmrbs.emich.edu;STOPPED; Host has exited from a period of scheduled downtime
[1461357313] HOST DOWNTIME ALERT: idm4.emich.edu;STOPPED; Host has exited from a period of scheduled downtime
[1461357422] SERVICE ALERT: hallefirst2;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1461357483] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! All Trays Still Full
[1461357612] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1461358212] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461358332] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1461358452] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461358952] SERVICE FLAPPING ALERT: idm2.emich.edu;NSDS Errors;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1461358992] SERVICE FLAPPING ALERT: idm4.emich.edu;NSDS Errors;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1461358999] SERVICE FLAPPING ALERT: idm3.emich.edu;NSDS Errors;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1461359330] SERVICE FLAPPING ALERT: mckennypk;Device Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1461359412] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1461359532] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 88% (15 min average)
[1461359883] SERVICE ALERT: hallefirst2;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1461359942] SERVICE ALERT: hallefirst2;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1461360003] SERVICE ALERT: hallefirst2;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1461362052] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461363252] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1461364452] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1461365052] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461365172] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461365292] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461366492] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 75% (15 min average)
[1461367092] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461367212] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461367332] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461368532] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 64% (15 min average)
[1461369003] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1461369132] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461369252] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461369372] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1461371062] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461371123] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461371182] SERVICE ALERT: hallefirst1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461372972] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461373572] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461373884] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;CRITICAL;HARD;3;Black Print Cartridge HP C8543X is at 1% - CRITICAL!
[1461374172] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461375372] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461378952] SERVICE ALERT: INTLTESTAPP;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461378972] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461379052] SERVICE ALERT: INTLTESTAPP;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1461379572] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461380772] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 75% (15 min average)
[1461381372] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1461381492] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1461381612] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461382212] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)