[1457503911] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 84.90 Gb - used: 78.54 Gb (93%) - free 6.36 Gb (7%)
[1457506362] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1457506884] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457506962] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1457507004] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 95% (5 min average)
[1457507124] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 74% (5 min average)
[1457507724] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 94% (5 min average)
[1457507844] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 77% (5 min average)
[1457510244] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1457510364] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 97% (5 min average)
[1457510484] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 95% (5 min average)
[1457511684] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 71% (5 min average)
[1457512884] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457513004] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 78% (5 min average)
[1457515359] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 47% (5 min average)
[1457515959] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 91% (5 min average)
[1457516079] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 82% (5 min average)
[1457518012] SERVICE ALERT: DB01-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457518132] SERVICE ALERT: DB01-2012;CPU Load;OK;SOFT;2;CPU Load 57% (5 min average)
[1457518732] SERVICE ALERT: DB01-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1457518852] SERVICE ALERT: DB01-2012;CPU Load;OK;SOFT;2;CPU Load 74% (5 min average)
[1457520279] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1457520399] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 80% (5 min average)
[1457520999] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 94% (5 min average)
[1457521119] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 63% (5 min average)
[1457521719] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457521839] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1457521959] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 78% (5 min average)
[1457522559] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457522679] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1457522799] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1457523499] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Genuine HP cartridge installed
[1457523672] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1457523987] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1457524015] HOST ALERT: hallefirst2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1457524047] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1457524077] HOST ALERT: hallefirst2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1457524105] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1457524107] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1457524135] HOST ALERT: hallefirst2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1457524165] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1457524195] HOST ALERT: hallefirst2;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1457524225] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1457524255] HOST ALERT: hallefirst2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1457524279] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1457524309] HOST ALERT: hallefirst2;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1457524338] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1457524368] HOST ALERT: hallefirst2;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1457524397] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1457524427] HOST ALERT: hallefirst2;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1457524517] HOST ALERT: hallefirst2;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1457524607] HOST ALERT: hallefirst2;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1457524852] SERVICE ALERT: DB01-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457524972] SERVICE ALERT: DB01-2012;CPU Load;OK;SOFT;2;CPU Load 5% (5 min average)
[1457526505] HOST ALERT: hallefirst2;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.49 ms
[1457526655] HOST ALERT: hallefirst2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1457526745] HOST ALERT: hallefirst2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1457526827] HOST ALERT: hallefirst2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1457526917] HOST ALERT: hallefirst2;DOWN;SOFT;4;(Host check timed out after 30.00 seconds)
[1457527007] HOST ALERT: hallefirst2;DOWN;SOFT;5;(Host check timed out after 30.00 seconds)
[1457527097] HOST ALERT: hallefirst2;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1457527137] HOST ALERT: hallefirst2;DOWN;SOFT;7;(Host check timed out after 30.00 seconds)
[1457527227] HOST ALERT: hallefirst2;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1457527317] HOST ALERT: hallefirst2;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1457527407] HOST ALERT: hallefirst2;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1457527490] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1457527550] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1457527610] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1457530599] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 38% (5 min average)
[1457532111] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 84.90 Gb - used: 75.37 Gb (89%) - free 9.53 Gb (11%)
[1457532879] HOST ALERT: AD-RMAD;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1457532893] SERVICE ALERT: AD-RMAD;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457532923] HOST ALERT: AD-RMAD;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1457532973] SERVICE ALERT: AD-RMAD;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457532987] SERVICE ALERT: AD-RMAD;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457532999] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457533013] SERVICE ALERT: AD-RMAD;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457533043] HOST ALERT: AD-RMAD;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1457533054] SERVICE ALERT: AD-RMAD;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457533057] SERVICE ALERT: AD-RMAD;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457533060] SERVICE ALERT: AD-RMAD;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457533084] HOST ALERT: AD-RMAD;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1457533093] SERVICE ALERT: AD-RMAD;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457533101] SERVICE ALERT: AD-RMAD;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457533107] SERVICE ALERT: AD-RMAD;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457533119] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1457533123] HOST ALERT: AD-RMAD;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1457533174] SERVICE ALERT: AD-RMAD;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457533177] SERVICE ALERT: AD-RMAD;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457533178] SERVICE ALERT: cava2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457533180] SERVICE ALERT: AD-RMAD;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457533192] SERVICE ALERT: AD-RMAD;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457533204] HOST ALERT: AD-RMAD;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1457533213] SERVICE ALERT: AD-RMAD;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457533221] SERVICE ALERT: AD-RMAD;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457533227] SERVICE ALERT: AD-RMAD;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457533239] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1457533243] HOST ALERT: AD-RMAD;DOWN;SOFT;3;(Host check timed out after 30.00 seconds)
[1457533250] SERVICE ALERT: AD-RMAD;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457533250] SERVICE ALERT: AD-RMAD;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457533251] SERVICE ALERT: AD-RMAD;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457533256] SERVICE ALERT: AD-RMAD;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457533280] HOST ALERT: AD-RMAD;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1457533294] SERVICE ALERT: AD-RMAD;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457533297] SERVICE ALERT: AD-RMAD;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457533298] SERVICE ALERT: cava2;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1457533300] SERVICE ALERT: AD-RMAD;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457533300] SERVICE ALERT: AD-RMAD;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457533302] SERVICE ALERT: AD-RMAD;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457533315] SERVICE ALERT: AD-RMAD;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457533324] HOST ALERT: AD-RMAD;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1457533329] SERVICE ALERT: AD-RMAD;E:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457533341] SERVICE ALERT: AD-RMAD;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457533359] HOST ALERT: AD-RMAD;DOWN;SOFT;6;(Host check timed out after 30.01 seconds)
[1457533370] SERVICE ALERT: AD-RMAD;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457533370] SERVICE ALERT: AD-RMAD;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457533371] SERVICE ALERT: AD-RMAD;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457533376] SERVICE ALERT: AD-RMAD;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457533400] HOST ALERT: AD-RMAD;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1457533420] SERVICE ALERT: AD-RMAD;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457533422] SERVICE ALERT: AD-RMAD;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457533435] SERVICE ALERT: AD-RMAD;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457533449] SERVICE ALERT: AD-RMAD;E:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457533450] HOST ALERT: AD-RMAD;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1457533490] SERVICE ALERT: AD-RMAD;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457533490] SERVICE ALERT: AD-RMAD;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457533491] SERVICE ALERT: AD-RMAD;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457533496] SERVICE ALERT: AD-RMAD;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457533520] HOST ALERT: AD-RMAD;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1457533540] SERVICE ALERT: AD-RMAD;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457533542] SERVICE ALERT: AD-RMAD;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457533555] SERVICE ALERT: AD-RMAD;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457533569] SERVICE ALERT: AD-RMAD;E:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457533570] HOST ALERT: AD-RMAD;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1457533839] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 84% (5 min average)
[1457533839] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1457533893] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1457534210] SERVICE ALERT: studentcenterpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1457534261] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1457534376] SERVICE ALERT: ADDC5;GPO Replication;WARNING;SOFT;1;Warning! 1 policies are not replicating properly!
[1457534439] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457534559] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1457534682] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 94% (5 min average)
[1457534682] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (27.6% change >= 20.0% threshold)
[1457534727] SERVICE ALERT: ADDC5;GPO Replication;WARNING;SOFT;2;Export-CSV : The process cannot access the file 'C:\Program
[1457534730] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (27.6% change >= 20.0% threshold)
[1457534778] SERVICE ALERT: INTLAPP;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 63.90 Gb - used: 54.67 Gb (86%) - free 9.23 Gb (14%)
[1457534898] SERVICE ALERT: INTLAPP;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 63.90 Gb - used: 54.67 Gb (86%) - free 9.23 Gb (14%)
[1457534983] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (27.6% change >= 20.0% threshold)
[1457535018] SERVICE ALERT: INTLAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 63.90 Gb - used: 54.67 Gb (86%) - free 9.23 Gb (14%)
[1457535149] SERVICE ALERT: ADDC5;GPO Replication;WARNING;HARD;3;Warning! 1 policies are not replicating properly!
[1457535252] SERVICE ALERT: SP2010;C:\ Drive Space;CRITICAL;SOFT;1;c:\ - total: 59.90 Gb - used: 54.48 Gb (91%) - free 5.42 Gb (9%)
[1457535279] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 88% (5 min average)
[1457535372] SERVICE ALERT: SP2010;C:\ Drive Space;CRITICAL;SOFT;2;c:\ - total: 59.90 Gb - used: 54.48 Gb (91%) - free 5.42 Gb (9%)
[1457535492] SERVICE ALERT: SP2010;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 59.90 Gb - used: 54.48 Gb (91%) - free 5.42 Gb (9%)
[1457535788] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (29.5% change >= 20.0% threshold)
[1457535856] SERVICE ALERT: DIRECTACCESS;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457535883] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457535902] SERVICE ALERT: DIRECTACCESS;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457535949] SERVICE ALERT: DIRECTACCESS;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457535976] SERVICE ALERT: DIRECTACCESS;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457535996] SERVICE ALERT: DIRECTACCESS;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457535998] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1457536022] SERVICE ALERT: DIRECTACCESS;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457536026] SERVICE ALERT: DIRECTACCESS;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457536042] SERVICE ALERT: DIRECTACCESS;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457536043] SERVICE ALERT: DIRECTACCESS;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457536089] SERVICE ALERT: DIRECTACCESS;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457536096] SERVICE ALERT: DIRECTACCESS;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457536116] SERVICE ALERT: DIRECTACCESS;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457536118] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1457536118] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (35.1% change >= 20.0% threshold)
[1457536136] SERVICE ALERT: DIRECTACCESS;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457536142] SERVICE ALERT: DIRECTACCESS;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457536146] SERVICE ALERT: DIRECTACCESS;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457536162] SERVICE ALERT: DIRECTACCESS;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457536163] SERVICE ALERT: DIRECTACCESS;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457536167] SERVICE ALERT: DIRECTACCESS;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457536183] SERVICE ALERT: DIRECTACCESS;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457536209] SERVICE ALERT: DIRECTACCESS;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457536230] SERVICE ALERT: DIRECTACCESS;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457536236] SERVICE ALERT: DIRECTACCESS;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457536256] SERVICE ALERT: DIRECTACCESS;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457536266] SERVICE ALERT: DIRECTACCESS;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457536276] SERVICE ALERT: DIRECTACCESS;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457536282] SERVICE ALERT: DIRECTACCESS;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457536283] SERVICE ALERT: DIRECTACCESS;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457536287] SERVICE ALERT: DIRECTACCESS;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457536303] SERVICE ALERT: DIRECTACCESS;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457536323] SERVICE ALERT: DIRECTACCESS;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457536329] SERVICE ALERT: DIRECTACCESS;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457536350] SERVICE ALERT: DIRECTACCESS;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457536370] SERVICE ALERT: DIRECTACCESS;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457536376] SERVICE ALERT: DIRECTACCESS;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457536396] SERVICE ALERT: DIRECTACCESS;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457536423] SERVICE ALERT: DIRECTACCESS;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457536443] SERVICE ALERT: DIRECTACCESS;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457536470] SERVICE ALERT: DIRECTACCESS;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457536490] SERVICE ALERT: DIRECTACCESS;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457536509] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1457536516] SERVICE ALERT: DIRECTACCESS;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457536563] SERVICE ALERT: DIRECTACCESS;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457536610] SERVICE ALERT: DIRECTACCESS;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457536718] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (5 min average)
[1457537403] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (39.7% change >= 20.0% threshold)
[1457537713] HOST ALERT: AD-RMAD;UP;HARD;10;PING OK - Packet loss = 79%, RTA = 0.45 ms
[1457537735] SERVICE ALERT: AD-RMAD;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1457537749] SERVICE ALERT: AD-RMAD;E:\ Drive Space;OK;HARD;3;e:\ - total: 399.87 Gb - used: 274.72 Gb (69%) - free 125.16 Gb (31%)
[1457537776] SERVICE ALERT: AD-RMAD;Memory Usage;OK;HARD;3;Memory usage: total:10239.64 MB - used: 2151.53 MB (21%) - free: 8088.11 MB (79%)
[1457537776] SERVICE ALERT: AD-RMAD;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1457537776] SERVICE ALERT: AD-RMAD;Service: WMI;OK;HARD;3;Winmgmt: Started
[1457537776] SERVICE ALERT: AD-RMAD;CPU Load;OK;HARD;3;CPU Load 2% (5 min average)
[1457537776] SERVICE ALERT: AD-RMAD;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1457537776] SERVICE ALERT: AD-RMAD;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1457537776] SERVICE ALERT: AD-RMAD;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1457537776] SERVICE ALERT: AD-RMAD;Service: Windows Time;OK;HARD;3;W32Time: Started
[1457537776] SERVICE ALERT: AD-RMAD;VMWare Tools;OK;HARD;3;VMTools: Started
[1457537776] SERVICE ALERT: AD-RMAD;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1457537776] SERVICE ALERT: AD-RMAD;Uptime;OK;HARD;3;System Uptime - 13 day(s) 2 hour(s) 2 minute(s)
[1457537776] SERVICE ALERT: AD-RMAD;C:\ Drive Space;OK;HARD;3;c:\ - total: 79.66 Gb - used: 33.26 Gb (42%) - free 46.39 Gb (58%)
[1457537776] SERVICE ALERT: AD-RMAD;Service: RPC;OK;HARD;3;RpcSs: Started
[1457537966] SERVICE ALERT: SP2010;C:\ Drive Space;OK;HARD;3;c:\ - total: 79.90 Gb - used: 54.49 Gb (68%) - free 25.41 Gb (32%)
[1457538269] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (38.8% change >= 20.0% threshold)
[1457538400] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (38.8% change >= 20.0% threshold)
[1457538453] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (38.8% change >= 20.0% threshold)
[1457539235] SERVICE ALERT: mckennypk;Device Status;CRITICAL;SOFT;1;(Return code of 13 is out of bounds)
[1457539235] HOST ALERT: mckennypk;DOWN;SOFT;1;(Return code of 13 is out of bounds)
[1457539237] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (36.9% change >= 20.0% threshold)
[1457539295] SERVICE ALERT: mckennypk;Device Status;OK;SOFT;2;Ready 164.76.152.26
[1457539299] HOST ALERT: mckennypk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.04 ms
[1457539778] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (36.0% change >= 20.0% threshold)
[1457540039] SERVICE ALERT: AD-RMAD;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1457540159] SERVICE ALERT: AD-RMAD;Service: Windows Time;CRITICAL;SOFT;2;W32Time: Stopped
[1457540269] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (36.0% change >= 20.0% threshold)
[1457540279] SERVICE ALERT: AD-RMAD;Service: Windows Time;CRITICAL;HARD;3;W32Time: Stopped
[1457540318] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1457540438] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 88% (5 min average)
[1457540903] SERVICE ALERT: owen101lab3;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1457540963] SERVICE ALERT: owen101lab3;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1457541023] SERVICE ALERT: owen101lab3;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1457541038] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457541124] SERVICE ALERT: AD-RMAD;Service: Windows Time;OK;HARD;3;W32Time: Started
[1457541124] SERVICE FLAPPING ALERT: AD-RMAD;Service: Windows Time;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1457541158] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 96% (5 min average)
[1457541193] SERVICE FLAPPING ALERT: AD-RMAD;Service: Windows Time;DISABLED; Flap detection has been disabled
[1457541241] SERVICE ALERT: AD-RMAD;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1457541261] SERVICE ALERT: AD-RMAD;Service: Windows Time;CRITICAL;SOFT;2;W32Time: Stopped
[1457541278] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 85% (5 min average)
[1457541381] SERVICE ALERT: AD-RMAD;Service: Windows Time;CRITICAL;HARD;3;W32Time: Stopped
[1457541531] SERVICE ALERT: ADDC2;Service: Distributed File System Replication;WARNING;SOFT;1;DFSR: Error
[1457541570] SERVICE ALERT: ADDC2;Service: Intersite Messaging;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457541605] SERVICE ALERT: ADDC4;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1457541624] SERVICE ALERT: ADDC2;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1457541651] SERVICE ALERT: ADDC2;Service: Distributed File System Replication;OK;SOFT;2;DFSR: Started
[1457541659] SERVICE ALERT: AD-RMAD;Service: Windows Time;OK;HARD;3;W32Time: Started
[1457541670] SERVICE ALERT: ADDC2;Service: Intersite Messaging;OK;SOFT;2;IsmServ: Started
[1457541691] SERVICE FLAPPING ALERT: AD-RMAD;Service: Windows Time;STARTED; Service appears to have started flapping (32.9% change >= 20.0% threshold)
[1457541704] SERVICE ALERT: ADRODC-DET;Service: MDT Monitor;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457541722] SERVICE ALERT: ADRODC-DET;Service: Distributed File System Replication;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457541725] SERVICE ALERT: ADDC4;Service: Windows Time;CRITICAL;SOFT;2;W32Time: Stopped
[1457541733] SERVICE ALERT: ADRODC-DET;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457541734] HOST ALERT: ADRODC-DET;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1457541744] SERVICE ALERT: ADDC2;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1457541751] SERVICE ALERT: ADRODC-DET;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457541755] SERVICE ALERT: ADRODC-DET;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457541775] SERVICE ALERT: hallefirstcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1457541777] SERVICE ALERT: ADRODC-DET;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457541779] SERVICE ALERT: ADRODC-DET;Service: DNS Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457541781] HOST ALERT: ADRODC-DET;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1457541789] SERVICE ALERT: ADRODC-DET;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457541805] HOST ALERT: hallefirstcolor;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1457541819] HOST ALERT: ADRODC-DET;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1457541824] SERVICE ALERT: ADRODC-DET;Service: MDT Monitor;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457541824] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Yellow Cartridge HP CE262A'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1457541828] HOST ALERT: hallefirstcolor;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1457541829] SERVICE ALERT: hallefirstcolor;Image Fuser Kit HP 110V-CE246A;OK;SOFT;2;Image Fuser Kit HP 110V-CE246A, 220V-CE247A is at 100% - OK!
[1457541842] SERVICE ALERT: ADRODC-DET;Service: Distributed File System Replication;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457541842] SERVICE ALERT: ADRODC-DET;Service: DFS Namespace;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457541845] SERVICE ALERT: ADDC4;Service: Windows Time;OK;SOFT;3;W32Time: Started
[1457541853] SERVICE ALERT: ADRODC-DET;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457541854] HOST ALERT: ADRODC-DET;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1457541871] SERVICE ALERT: ADRODC-DET;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457541875] SERVICE ALERT: ADRODC-DET;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457541884] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;OK;SOFT;2;Yellow Cartridge HP CE262A is at 86% - OK!
[1457541897] SERVICE ALERT: ADRODC-DET;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457541899] SERVICE ALERT: ADRODC-DET;Service: DNS Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457541901] HOST ALERT: ADRODC-DET;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1457541909] SERVICE ALERT: ADRODC-DET;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457541939] HOST ALERT: ADRODC-DET;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1457541944] SERVICE ALERT: ADRODC-DET;Service: MDT Monitor;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457541962] SERVICE ALERT: ADRODC-DET;Service: Distributed File System Replication;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457541962] SERVICE ALERT: ADRODC-DET;Service: DFS Namespace;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457541969] SERVICE ALERT: ADRODC-DET;Service: Intersite Messaging;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457541973] SERVICE ALERT: ADRODC-DET;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457541974] HOST ALERT: ADRODC-DET;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1457541977] SERVICE ALERT: ADRODC-DET;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457541981] SERVICE ALERT: ADRODC-DET;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457541991] SERVICE ALERT: ADRODC-DET;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457541995] SERVICE ALERT: ADRODC-DET;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457542007] HOST ALERT: ADRODC-DET;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1457542017] SERVICE ALERT: ADRODC-DET;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457542019] SERVICE ALERT: ADRODC-DET;Service: DNS Service;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457542022] SERVICE ALERT: ADRODC-DET;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457542029] SERVICE ALERT: ADRODC-DET;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457542032] SERVICE ALERT: ADRODC-DET;Service: Key Distribution Center;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457542039] SERVICE ALERT: ADRODC-DET;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457542043] SERVICE ALERT: ADRODC-DET;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457542047] HOST ALERT: ADRODC-DET;DOWN;SOFT;9;(Host check timed out after 30.00 seconds)
[1457542049] SERVICE ALERT: ADRODC-DET;Service: PXE Boot;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457542067] SERVICE ALERT: ADRODC-DET;Service: AD Domain Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457542071] SERVICE ALERT: ADRODC-DET;AD Replication;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 30 seconds.
[1457542079] HOST ALERT: ADRODC-DET;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1457542079] SERVICE ALERT: ADRODC-DET;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457542082] SERVICE ALERT: ADRODC-DET;Service: DFS Namespace;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457542084] SERVICE ALERT: ADRODC-DET;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457542085] SERVICE ALERT: ADRODC-DET;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457542089] SERVICE ALERT: ADRODC-DET;Service: Intersite Messaging;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457542097] SERVICE ALERT: ADRODC-DET;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457542101] SERVICE ALERT: ADRODC-DET;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457542117] SERVICE ALERT: ADRODC-DET;Service: AD Web Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457542142] SERVICE ALERT: ADRODC-DET;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457542152] SERVICE ALERT: ADRODC-DET;Service: Key Distribution Center;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457542159] SERVICE ALERT: ADRODC-DET;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457542163] SERVICE ALERT: ADRODC-DET;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457542169] SERVICE ALERT: ADRODC-DET;Service: PXE Boot;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457542187] SERVICE ALERT: ADRODC-DET;Service: AD Domain Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457542191] SERVICE ALERT: ADRODC-DET;AD Replication;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 30 seconds.
[1457542199] SERVICE ALERT: ADRODC-DET;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457542204] SERVICE ALERT: ADRODC-DET;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457542205] SERVICE ALERT: ADRODC-DET;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457542209] SERVICE ALERT: ADRODC-DET;Service: Intersite Messaging;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457542217] SERVICE ALERT: ADRODC-DET;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457542221] SERVICE ALERT: ADRODC-DET;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457542223] SERVICE ALERT: owen101lab3;Device Status;OK;HARD;3;41.3 UNEXPECTED SIZE IN TRAY 2 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   LOAD TRAY 2 PLAIN LETTER To use another 74 72 61 79 20 70 72 65 73 73 20 1E
[1457542237] SERVICE ALERT: ADRODC-DET;Service: AD Web Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457542253] SERVICE ALERT: ADRODC-DET;Service: Remote Management;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457542262] SERVICE ALERT: ADRODC-DET;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457542271] SERVICE ALERT: ADDC5;GPO Replication;OK;HARD;3;WARNING: [PROCESS] Something wrong happened
[1457542271] SERVICE FLAPPING ALERT: ADDC5;GPO Replication;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1457542272] SERVICE ALERT: ADRODC-DET;Service: Key Distribution Center;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457542279] SERVICE ALERT: ADRODC-DET;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457542283] SERVICE ALERT: ADRODC-DET;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457542289] SERVICE ALERT: ADRODC-DET;Service: PXE Boot;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457542307] SERVICE ALERT: ADRODC-DET;Service: AD Domain Service;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457542311] SERVICE ALERT: ADRODC-DET;AD Replication;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 30 seconds.
[1457542319] SERVICE ALERT: ADRODC-DET;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457542324] SERVICE ALERT: ADRODC-DET;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457542325] SERVICE ALERT: ADRODC-DET;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457542357] SERVICE ALERT: ADRODC-DET;Service: AD Web Service;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457542373] SERVICE ALERT: ADRODC-DET;Service: Remote Management;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457542478] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457542493] SERVICE ALERT: ADRODC-DET;Service: Remote Management;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1457542598] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1457542718] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1457542815] SERVICE FLAPPING ALERT: AD-RMAD;Service: Windows Time;STARTED; Service appears to have started flapping (32.1% change >= 20.0% threshold)
[1457542815] SERVICE FLAPPING ALERT: ADDC5;GPO Replication;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1457542815] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (38.6% change >= 20.0% threshold)
[1457542881] SERVICE FLAPPING ALERT: AD-RMAD;Service: Windows Time;STARTED; Service appears to have started flapping (31.3% change >= 20.0% threshold)
[1457542881] SERVICE FLAPPING ALERT: ADDC5;GPO Replication;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1457542881] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (38.6% change >= 20.0% threshold)
[1457542898] SERVICE ALERT: cava2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457543018] SERVICE ALERT: cava2;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1457543955] SERVICE FLAPPING ALERT: ADDC5;GPO Replication;DISABLED; Flap detection has been disabled
[1457544756] SERVICE ALERT: BOEAPP;C:\ Drive Space;OK;HARD;3;c:\ - total: 219.90 Gb - used: 160.71 Gb (73%) - free 59.19 Gb (27%)
[1457547730] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;SOFT;1;WARNING: Replication with ADDC5 failed!
[1457547850] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;SOFT;2;WARNING: Replication with ADDC5 failed!
[1457547950] SERVICE ALERT: ADRODC-LIV;AD Replication;OK;SOFT;3;OK: replication is up and running
[1457547970] SERVICE ALERT: ADRODC-DET;Service: PXE Boot;OK;HARD;3;WDSServer: Started
[1457547971] SERVICE ALERT: ADRODC-DET;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1457547974] HOST ALERT: ADRODC-DET;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 7.89 ms
[1457547975] SERVICE ALERT: ADRODC-DET;Service: WMI;OK;HARD;3;Winmgmt: Started
[1457547997] SERVICE ALERT: ADRODC-DET;Uptime;OK;HARD;3;System Uptime - 4 day(s) 17 hour(s) 46 minute(s)
[1457547999] SERVICE ALERT: ADRODC-DET;Service: DNS Service;OK;HARD;3;DNS: Started
[1457548009] SERVICE ALERT: ADRODC-DET;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1457548062] SERVICE ALERT: ADRODC-DET;Service: DFS Namespace;OK;HARD;3;Dfs: Started
[1457548189] SERVICE ALERT: ADRODC-DET;Service: Intersite Messaging;OK;HARD;3;IsmServ: Started
[1457548242] SERVICE ALERT: ADRODC-DET;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1457548252] SERVICE ALERT: ADRODC-DET;Service: Key Distribution Center;OK;HARD;3;KDC: Started
[1457548259] SERVICE ALERT: ADRODC-DET;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1457548287] SERVICE ALERT: ADRODC-DET;Service: AD Domain Service;OK;HARD;3;NTDS: Started
[1457548299] SERVICE ALERT: ADRODC-DET;Service: Windows Time;OK;HARD;3;W32Time: Started
[1457548304] SERVICE ALERT: ADRODC-DET;Service: RPC;OK;HARD;3;RpcSs: Started
[1457548305] SERVICE ALERT: ADRODC-DET;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1457548308] SERVICE ALERT: ADRODC-DET;Memory Usage;OK;HARD;3;Memory usage: total:37586.86 MB - used: 3473.36 MB (9%) - free: 34113.50 MB (91%)
[1457548337] SERVICE ALERT: ADRODC-DET;Service: AD Web Service;OK;HARD;3;ADWS: Started
[1457548423] SERVICE ALERT: ADRODC-DET;AD Replication;WARNING;HARD;3;WARNING: Replication with ADDC5 failed!
[1457548468] SERVICE ALERT: ADRODC-DET;C:\ Drive Space;OK;HARD;3;c:\ - total: 1115.16 Gb - used: 267.74 Gb (24%) - free 847.42 Gb (76%)
[1457548473] SERVICE ALERT: ADRODC-DET;Service: Remote Management;OK;HARD;3;WinRM: Started
[1457548514] SERVICE ALERT: ADRODC-DET;CPU Load;OK;HARD;3;CPU Load 0% (5 min average)
[1457548524] SERVICE ALERT: ADRODC-DET;Service: MDT Monitor;OK;HARD;3;MDT_Monitor: Started
[1457548542] SERVICE ALERT: ADRODC-DET;Service: Distributed File System Replication;OK;HARD;3;DFSR: Started
[1457548553] SERVICE ALERT: ADRODC-DET;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1457549134] SERVICE ALERT: ADRODC-DET;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457549144] SERVICE ALERT: ADRODC-DET;Service: MDT Monitor;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457549162] SERVICE ALERT: ADRODC-DET;Service: Distributed File System Replication;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457549164] HOST ALERT: ADRODC-DET;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1457549173] SERVICE ALERT: ADRODC-DET;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457549181] SERVICE ALERT: ADRODC-DET;Service: PXE Boot;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457549191] SERVICE ALERT: ADRODC-DET;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457549195] SERVICE ALERT: ADRODC-DET;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457549203] HOST ALERT: ADRODC-DET;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1457549217] SERVICE ALERT: ADRODC-DET;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457549219] SERVICE ALERT: ADRODC-DET;Service: DNS Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457549229] SERVICE ALERT: ADRODC-DET;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457549243] HOST ALERT: ADRODC-DET;UP;SOFT;3;PING WARNING - Packet loss = 81%, RTA = 8.26 ms
[1457549243] SERVICE ALERT: ADRODC-DET;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1457549244] SERVICE ALERT: ADRODC-DET;Service: MDT Monitor;OK;SOFT;2;MDT_Monitor: Started
[1457549262] SERVICE ALERT: ADRODC-DET;Service: Distributed File System Replication;OK;SOFT;2;DFSR: Started
[1457549273] SERVICE ALERT: ADRODC-DET;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1457549281] SERVICE ALERT: ADRODC-DET;Service: PXE Boot;OK;SOFT;2;WDSServer: Started
[1457549291] SERVICE ALERT: ADRODC-DET;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1457549295] SERVICE ALERT: ADRODC-DET;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1457549317] SERVICE ALERT: ADRODC-DET;Uptime;OK;SOFT;2;System Uptime - 4 day(s) 18 hour(s) 8 minute(s)
[1457549318] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (4.8% change < 5.0% threshold)
[1457549319] SERVICE ALERT: ADRODC-DET;Service: DNS Service;OK;SOFT;2;DNS: Started
[1457549329] SERVICE ALERT: ADRODC-DET;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1457550684] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1457550745] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1457550805] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1457551308] SERVICE ALERT: ADDC5;GPO Replication;WARNING;SOFT;1;Warning! 1 policies are not replicating properly!
[1457551405] SERVICE ALERT: eateriespk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1457551604] SERVICE ALERT: ADDC5;GPO Replication;WARNING;SOFT;2;Warning! 1 policies are not replicating properly!
[1457551809] SERVICE ALERT: ADRODC-DET;Service: Intersite Messaging;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457551839] HOST ALERT: ADRODC-DET;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1457551862] SERVICE ALERT: ADRODC-DET;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457551872] SERVICE ALERT: ADRODC-DET;Service: Key Distribution Center;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457551879] SERVICE ALERT: ADRODC-DET;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457551892] HOST ALERT: ADRODC-DET;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1457551901] SERVICE ALERT: ADDC5;GPO Replication;OK;SOFT;3;WARNING: [PROCESS] Something wrong happened
[1457551907] SERVICE ALERT: ADRODC-DET;Service: AD Domain Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457551919] SERVICE ALERT: ADRODC-DET;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457551924] SERVICE ALERT: ADRODC-DET;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457551925] SERVICE ALERT: ADRODC-DET;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457551928] SERVICE ALERT: ADRODC-DET;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457551929] SERVICE ALERT: ADRODC-DET;Service: Intersite Messaging;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457551937] HOST ALERT: ADRODC-DET;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1457551950] HOST ALERT: ADRODC-DET;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 8.34 ms
[1457551962] SERVICE ALERT: ADRODC-DET;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1457551972] SERVICE ALERT: ADRODC-DET;Service: Key Distribution Center;OK;SOFT;2;KDC: Started
[1457551979] SERVICE ALERT: ADRODC-DET;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1457552007] SERVICE ALERT: ADRODC-DET;Service: AD Domain Service;OK;SOFT;2;NTDS: Started
[1457552019] SERVICE ALERT: ADRODC-DET;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1457552024] SERVICE ALERT: ADRODC-DET;Service: RPC;OK;SOFT;2;RpcSs: Started
[1457552025] SERVICE ALERT: ADRODC-DET;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1457552028] SERVICE ALERT: ADRODC-DET;Memory Usage;OK;SOFT;2;Memory usage: total:37586.86 MB - used: 3440.14 MB (9%) - free: 34146.72 MB (91%)
[1457552029] SERVICE ALERT: ADRODC-DET;Service: Intersite Messaging;OK;SOFT;3;IsmServ: Started
[1457552429] SERVICE ALERT: owen101lab3;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1457552459] HOST ALERT: owen101lab3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1457552489] SERVICE ALERT: owen101lab3;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1457552519] HOST ALERT: owen101lab3;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1457552549] SERVICE ALERT: owen101lab3;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1457552579] HOST ALERT: owen101lab3;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1457552649] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1457552669] HOST ALERT: owen101lab3;DOWN;SOFT;4;(Host check timed out after 30.00 seconds)
[1457552709] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1457552739] HOST ALERT: owen101lab3;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1457552769] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1457552799] HOST ALERT: owen101lab3;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1457552828] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1457552854] SERVICE ALERT: ADRODC-DET;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457552858] HOST ALERT: owen101lab3;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1457552864] SERVICE ALERT: ADRODC-DET;Service: MDT Monitor;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457552882] SERVICE ALERT: ADRODC-DET;Service: DFS Namespace;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457552882] SERVICE ALERT: ADRODC-DET;Service: Distributed File System Replication;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457552884] HOST ALERT: ADRODC-DET;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1457552888] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1457552893] SERVICE ALERT: ADRODC-DET;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457552901] SERVICE ALERT: ADRODC-DET;Service: PXE Boot;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457552911] SERVICE ALERT: ADRODC-DET;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457552915] SERVICE ALERT: ADRODC-DET;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457552918] HOST ALERT: owen101lab3;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1457552923] HOST ALERT: ADRODC-DET;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1457552937] SERVICE ALERT: ADRODC-DET;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457552939] SERVICE ALERT: ADRODC-DET;Service: DNS Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457552948] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1457552949] SERVICE ALERT: ADRODC-DET;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457552957] SERVICE ALERT: ADRODC-DET;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1457552961] HOST ALERT: ADRODC-DET;UP;SOFT;3;PING WARNING - Packet loss = 80%, RTA = 8.28 ms
[1457552964] SERVICE ALERT: ADRODC-DET;Service: MDT Monitor;OK;SOFT;2;MDT_Monitor: Started
[1457552978] HOST ALERT: owen101lab3;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1457552982] SERVICE ALERT: ADRODC-DET;Service: DFS Namespace;OK;SOFT;2;Dfs: Started
[1457552982] SERVICE ALERT: ADRODC-DET;Service: Distributed File System Replication;OK;SOFT;2;DFSR: Started
[1457552993] SERVICE ALERT: ADRODC-DET;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1457553001] SERVICE ALERT: ADRODC-DET;Service: PXE Boot;OK;SOFT;2;WDSServer: Started
[1457553011] SERVICE ALERT: ADRODC-DET;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1457553015] SERVICE ALERT: ADRODC-DET;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1457553037] SERVICE ALERT: ADRODC-DET;Uptime;OK;SOFT;2;System Uptime - 4 day(s) 19 hour(s) 10 minute(s)
[1457553039] SERVICE ALERT: ADRODC-DET;Service: DNS Service;OK;SOFT;2;DNS: Started
[1457553049] SERVICE ALERT: ADRODC-DET;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1457553059] SERVICE FLAPPING ALERT: AD-RMAD;Service: Windows Time;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1457553068] HOST ALERT: owen101lab3;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1457553616] SERVICE ALERT: EVENTS-2012;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1457553734] SERVICE ALERT: EVENTS-2012;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1457558027] SERVICE ALERT: ADRODC-DET;AD Replication;OK;HARD;3;OK: replication is up and running
[1457559518] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 94% (5 min average)
[1457559963] SERVICE ALERT: owen101lab1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1457560023] SERVICE ALERT: owen101lab1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1457560084] SERVICE ALERT: owen101lab1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1457560118] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 74% (5 min average)
[1457560683] SERVICE ALERT: owen101lab1;Device Status;OK;HARD;3;Ready
[1457560718] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1457560723] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1457560761] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1457560783] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1457560791] HOST ALERT: hallefirst1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1457560821] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1457560825] HOST ALERT: hallefirst1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1457560838] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1457560843] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1457560873] HOST ALERT: hallefirst1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1457560881] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1457560911] HOST ALERT: hallefirst1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1457560958] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1457560975] HOST ALERT: hallefirst1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.49 ms
[1457560975] HOST FLAPPING ALERT: hallefirst1;STARTED; Host appears to have started flapping (23.0% change > 20.0% threshold)
[1457561057] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1457561117] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1457561177] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1457561650] SERVICE ALERT: hallefirst3;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1457561699] SERVICE ALERT: SP2010;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457561710] SERVICE ALERT: hallefirst3;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1457561714] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1457561723] SERVICE ALERT: hallefirst3;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1457561770] SERVICE ALERT: hallefirst3;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1457561774] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1457561783] SERVICE ALERT: hallefirst3;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1457561819] SERVICE ALERT: SP2010;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1457561834] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1457561843] SERVICE ALERT: hallefirst3;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1457562377] HOST FLAPPING ALERT: hallefirst1;STOPPED; Host appears to have stopped flapping (4.0% change < 5.0% threshold)
[1457562407] HOST ALERT: hallefirst1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1457562428] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 36% - OK!
[1457562438] SERVICE ALERT: hallefirst3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1457562489] HOST ALERT: hallefirst1;UP;SOFT;2;PING OK - Packet loss = 78%, RTA = 0.65 ms
[1457562637] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 26% - OK!
[1457562674] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;Ready
[1457562964] SERVICE ALERT: hallefirst3;Device Status;OK;HARD;3;TRAY 2 EMPTY PLAIN LETTER   Ready
[1457562972] SERVICE ALERT: hallefirst1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1457564501] SERVICE ALERT: studentcenterpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1457564558] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 94% (5 min average)
[1457564558] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1457564562] SERVICE ALERT: studentcenterpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1457564621] SERVICE ALERT: studentcenterpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1457564859] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1457564919] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 22% - OK!
[1457565158] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1457565221] SERVICE ALERT: studentcenterpk;Device Status;OK;HARD;3;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1457566421] SERVICE ALERT: studentcenterpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1457566481] SERVICE ALERT: studentcenterpk;Device Status;OK;SOFT;2;Clearing paper path
[1457567081] SERVICE ALERT: studentcenterpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1457567142] SERVICE ALERT: studentcenterpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1457567201] SERVICE ALERT: studentcenterpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1457568944] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;SOFT;1;WARNING: Replication with ADDC5 failed!
[1457569064] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;SOFT;2;WARNING: Replication with ADDC5 failed!
[1457569184] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;HARD;3;WARNING: Replication with ADDC5 failed!
[1457570919] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1457570979] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1457571039] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;OK;SOFT;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1457571401] SERVICE ALERT: studentcenterpk;Device Status;OK;HARD;3;Ready
[1457571401] SERVICE FLAPPING ALERT: studentcenterpk;Device Status;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1457571487] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 1 Trays Empty
[1457571548] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 1 Trays Empty
[1457571607] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 1 Trays Empty
[1457572207] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1457575526] SERVICE ALERT: C4SAS-SE;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1457575533] HOST ALERT: hallethirdsouth;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1457575623] HOST ALERT: hallethirdsouth;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1457575646] SERVICE ALERT: C4SAS-SE;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1457575713] HOST ALERT: hallethirdsouth;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1457575748] HOST ALERT: hallethirdsouth;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1457575829] HOST ALERT: hallethirdsouth;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1457575919] HOST ALERT: hallethirdsouth;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1457576009] HOST ALERT: hallethirdsouth;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1457576099] HOST ALERT: hallethirdsouth;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1457576189] HOST ALERT: hallethirdsouth;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1457576278] HOST ALERT: hallethirdsouth;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1457576556] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1457579201] SERVICE FLAPPING ALERT: studentcenterpk;Device Status;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1457579989] SERVICE ALERT: ADRODC-LIV;AD Replication;OK;HARD;3;OK: replication is up and running
[1457580301] SERVICE ALERT: CBORD;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457580401] SERVICE ALERT: CBORD;Uptime;OK;SOFT;2;System Uptime - 11 day(s) 23 hour(s) 22 minute(s)
[1457580600] HOST DOWNTIME ALERT: MBAM;STARTED; Host has entered a period of scheduled downtime
[1457580600] HOST DOWNTIME ALERT: TI22-PSYCLINIC;STARTED; Host has entered a period of scheduled downtime
[1457580600] HOST DOWNTIME ALERT: REFWIN;STARTED; Host has entered a period of scheduled downtime
[1457580600] HOST DOWNTIME ALERT: LANSWEEPER;STARTED; Host has entered a period of scheduled downtime
[1457580600] HOST DOWNTIME ALERT: WINKMS;STARTED; Host has entered a period of scheduled downtime
[1457580600] HOST DOWNTIME ALERT: IMAGES;STARTED; Host has entered a period of scheduled downtime
[1457580600] HOST DOWNTIME ALERT: PROCARE;STARTED; Host has entered a period of scheduled downtime
[1457580600] HOST DOWNTIME ALERT: HDBACKUPS;STARTED; Host has entered a period of scheduled downtime
[1457580600] HOST DOWNTIME ALERT: FTK;STARTED; Host has entered a period of scheduled downtime
[1457580600] HOST DOWNTIME ALERT: TI22;STARTED; Host has entered a period of scheduled downtime
[1457580600] HOST DOWNTIME ALERT: PNCWEB-2012;STARTED; Host has entered a period of scheduled downtime
[1457580600] HOST DOWNTIME ALERT: Echo-01;STARTED; Host has entered a period of scheduled downtime
[1457580600] HOST DOWNTIME ALERT: EMUSQLEXPRESS;STARTED; Host has entered a period of scheduled downtime
[1457580600] HOST DOWNTIME ALERT: PNCDB-2012;STARTED; Host has entered a period of scheduled downtime
[1457580600] HOST DOWNTIME ALERT: EMUSCRIPT;STARTED; Host has entered a period of scheduled downtime
[1457580600] HOST DOWNTIME ALERT: TABLEAU;STARTED; Host has entered a period of scheduled downtime
[1457580600] HOST DOWNTIME ALERT: OPIE;STARTED; Host has entered a period of scheduled downtime
[1457580600] HOST DOWNTIME ALERT: DPS-SECURE02;STARTED; Host has entered a period of scheduled downtime
[1457580600] HOST DOWNTIME ALERT: MBAMDB;STARTED; Host has entered a period of scheduled downtime
[1457580600] HOST DOWNTIME ALERT: CEDARON;STARTED; Host has entered a period of scheduled downtime
[1457580600] HOST DOWNTIME ALERT: OFFICEKMS;STARTED; Host has entered a period of scheduled downtime
[1457580600] HOST DOWNTIME ALERT: AD-RMAD;STARTED; Host has entered a period of scheduled downtime
[1457580600] HOST DOWNTIME ALERT: EMU-VALUE01;STARTED; Host has entered a period of scheduled downtime
[1457580600] HOST DOWNTIME ALERT: UHSPHARM-2012;STARTED; Host has entered a period of scheduled downtime
[1457580600] HOST DOWNTIME ALERT: SANDBOX-DB;STARTED; Host has entered a period of scheduled downtime
[1457580601] HOST DOWNTIME ALERT: NESSUS;STARTED; Host has entered a period of scheduled downtime
[1457580601] HOST DOWNTIME ALERT: XDM;STARTED; Host has entered a period of scheduled downtime
[1457580753] SERVICE ALERT: ADWSUSSERVERS;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457580772] HOST ALERT: ADWSUSSERVERS;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.133)
[1457580776] SERVICE ALERT: ADWSUSSERVERS;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457580779] HOST ALERT: ADWSUSSERVERS;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.133)
[1457580799] HOST ALERT: ADWSUSSERVERS;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.36 ms
[1457580853] SERVICE ALERT: ADWSUSSERVERS;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1457580876] SERVICE ALERT: ADWSUSSERVERS;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1457582578] SERVICE ALERT: MBAM;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457582601] SERVICE ALERT: XDM;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457582603] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1457582608] SERVICE ALERT: PROCARE;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457582641] SERVICE ALERT: EMUSQLEXPRESS;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1457582655] SERVICE ALERT: Echo-01;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1457582698] SERVICE ALERT: MBAM;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1457582721] SERVICE ALERT: XDM;CPU Load;WARNING;SOFT;2;CPU Load 94% (5 min average)
[1457582723] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 95% (5 min average)
[1457582728] SERVICE ALERT: PROCARE;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1457582761] SERVICE ALERT: EMUSQLEXPRESS;CPU Load;OK;SOFT;2;CPU Load 46% (5 min average)
[1457582775] SERVICE ALERT: Echo-01;CPU Load;OK;SOFT;2;CPU Load 37% (5 min average)
[1457582818] SERVICE ALERT: MBAM;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1457582841] SERVICE ALERT: XDM;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1457582843] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1457582848] SERVICE ALERT: PROCARE;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1457582963] SERVICE ALERT: NETENG-PA;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457583021] SERVICE ALERT: NETENG-PA;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.102.208 and port 12489: No route to host
[1457583037] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1457583052] SERVICE ALERT: NESSUS;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.2.69 and port 12489: Connection refused
[1457583083] SERVICE ALERT: NETENG-PA;CPU Load;OK;SOFT;2;CPU Load 4% (5 min average)
[1457583097] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1457583126] SERVICE ALERT: PROCARE;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457583129] SERVICE ALERT: NETENG-PA;Memory Usage;OK;SOFT;2;Memory usage: total:14333.44 MB - used: 818.23 MB (6%) - free: 13515.21 MB (94%)
[1457583135] SERVICE ALERT: MBAMDB;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1457583135] SERVICE ALERT: MBAM;Service: Group Policy;CRITICAL;SOFT;1;gpsvc: Stopped
[1457583156] HOST ALERT: PROCARE;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1457583158] SERVICE ALERT: halleground2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1457583172] SERVICE ALERT: NESSUS;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1457583173] HOST ALERT: PROCARE;UP;SOFT;2;PING OK - Packet loss = 72%, RTA = 0.44 ms
[1457583219] SERVICE ALERT: MBAMDB;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457583219] SERVICE ALERT: MBAMDB;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.7.97 and port 12489: Connection refused
[1457583226] SERVICE ALERT: PROCARE;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1457583255] SERVICE ALERT: MBAM;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1457583275] SERVICE ALERT: MBAMDB;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457583298] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1457583319] SERVICE ALERT: MBAMDB;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1457583339] SERVICE ALERT: MBAMDB;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1457583358] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 30% - OK!
[1457583375] SERVICE ALERT: MBAMDB;CPU Load;OK;SOFT;3;CPU Load 50% (5 min average)
[1457583392] SERVICE ALERT: EMUSQLEXPRESS;VMWare Tools;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1457583407] SERVICE ALERT: EMUSQLEXPRESS;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.102.163 and port 12489: Connection refused
[1457583433] SERVICE ALERT: WINKMS;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457583448] SERVICE ALERT: PROCARE;CPU Load;WARNING;HARD;3;CPU Load 94% (5 min average)
[1457583448] SERVICE ALERT: HDBACKUPS;Uptime;CRITICAL;SOFT;1;connect to address 164.76.102.104 and port 12489: No route to host
[1457583451] HOST ALERT: HDBACKUPS;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.104)
[1457583475] SERVICE ALERT: TI22;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.99.227 and port 12489: Connection refused
[1457583478] SERVICE ALERT: HDBACKUPS;F:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.102.104 and port 12489: No route to host
[1457583481] HOST ALERT: HDBACKUPS;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.104)
[1457583496] SERVICE ALERT: HDBACKUPS;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.102.104 and port 12489: No route to host
[1457583499] SERVICE ALERT: HDBACKUPS;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.102.104 and port 12489: No route to host
[1457583499] HOST ALERT: HDBACKUPS;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.102.104)
[1457583512] SERVICE ALERT: EMUSQLEXPRESS;VMWare Tools;OK;SOFT;2;VMTools: Started
[1457583514] SERVICE ALERT: HDBACKUPS;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.102.104 and port 12489: No route to host
[1457583514] SERVICE ALERT: Echo-01;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.2.241 and port 12489: Connection refused
[1457583517] HOST ALERT: HDBACKUPS;DOWN;SOFT;4;CRITICAL - Host Unreachable (164.76.102.104)
[1457583519] SERVICE ALERT: PNCWEB-2012;Service: User Profile;CRITICAL;SOFT;1;connect to address 164.76.2.178 and port 12489: Connection refused
[1457583521] SERVICE ALERT: UHSPHARM-2012;VMWare Tools;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1457583527] SERVICE ALERT: EMUSQLEXPRESS;Service: RPC;OK;SOFT;2;RpcSs: Started
[1457583533] SERVICE ALERT: WINKMS;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1457583538] SERVICE ALERT: WINKMS;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457583541] SERVICE ALERT: HDBACKUPS;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.102.104 and port 12489: No route to host
[1457583544] HOST ALERT: HDBACKUPS;DOWN;SOFT;5;CRITICAL - Host Unreachable (164.76.102.104)
[1457583556] SERVICE ALERT: HDBACKUPS;Uptime;CRITICAL;SOFT;2;connect to address 164.76.102.104 and port 12489: Connection refused
[1457583561] SERVICE ALERT: CBORD;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.99.168 and port 12489: Connection refused
[1457583562] SERVICE ALERT: PNCAPP-2012;C:\ Drive Space;UNKNOWN;SOFT;1;NSClient - ERROR: Invalid return from command: check_drivesize
[1457583580] HOST ALERT: HDBACKUPS;UP;SOFT;6;PING WARNING - Packet loss = 80%, RTA = 0.36 ms
[1457583586] SERVICE ALERT: TI22;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1457583595] SERVICE ALERT: HDBACKUPS;F:\ Drive Space;OK;SOFT;2;f:\ - total: 3315.22 Gb - used: 1493.24 Gb (45%) - free 1821.99 Gb (55%)
[1457583609] SERVICE ALERT: PNCAPP-2012;CPU Load;CRITICAL;SOFT;1;connect to address 164.76.2.177 and port 12489: Connection refused
[1457583613] SERVICE ALERT: HDBACKUPS;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1457583616] SERVICE ALERT: HDBACKUPS;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1457583631] SERVICE ALERT: HDBACKUPS;Memory Usage;OK;SOFT;2;Memory usage: total:37532.83 MB - used: 1286.65 MB (3%) - free: 36246.19 MB (97%)
[1457583634] SERVICE ALERT: Echo-01;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1457583639] SERVICE ALERT: PNCWEB-2012;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1457583641] SERVICE ALERT: UHSPHARM-2012;VMWare Tools;OK;SOFT;2;VMTools: Started
[1457583658] SERVICE ALERT: HDBACKUPS;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1457583658] SERVICE ALERT: WINKMS;CPU Load;OK;SOFT;2;CPU Load 50% (5 min average)
[1457583676] SERVICE ALERT: HDBACKUPS;Uptime;OK;SOFT;3;System Uptime - 0 day(s) 0 hour(s) 2 minute(s)
[1457583681] SERVICE ALERT: CBORD;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1457583682] SERVICE ALERT: PNCAPP-2012;C:\ Drive Space;OK;SOFT;2;c:\ - total: 59.66 Gb - used: 20.17 Gb (34%) - free 39.49 Gb (66%)
[1457583701] SERVICE ALERT: CEDARON;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1457583719] SERVICE ALERT: DPS-SECURE02;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457583729] SERVICE ALERT: PNCAPP-2012;CPU Load;OK;SOFT;2;CPU Load 51% (5 min average)
[1457583749] HOST ALERT: DPS-SECURE02;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1457583757] SERVICE ALERT: halleground2;Device Status;OK;HARD;3;Ready
[1457583772] SERVICE ALERT: DPS-SECURE02;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457583772] SERVICE ALERT: DPS-SECURE02;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457583772] SERVICE ALERT: DPS-SECURE02;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457583796] SERVICE ALERT: EMU-VALUE01;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457583801] SERVICE ALERT: EMU-VALUE01;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457583802] HOST ALERT: DPS-SECURE02;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1457583821] SERVICE ALERT: CEDARON;CPU Load;OK;SOFT;2;CPU Load 61% (5 min average)
[1457583823] SERVICE ALERT: DPS-SECURE02;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457583826] HOST ALERT: EMU-VALUE01;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1457583839] SERVICE ALERT: DPS-SECURE02;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457583849] SERVICE ALERT: EMU-VALUE01;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457583853] HOST ALERT: DPS-SECURE02;UP;SOFT;3;PING WARNING - Packet loss = 96%, RTA = 0.37 ms
[1457583872] SERVICE ALERT: DPS-SECURE02;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 1 minute(s)
[1457583872] SERVICE ALERT: DPS-SECURE02;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1457583872] SERVICE ALERT: DPS-SECURE02;C:\ Drive Space;OK;SOFT;2;c:\ - total: 278.53 Gb - used: 14.10 Gb (5%) - free 264.43 Gb (95%)
[1457583879] HOST ALERT: EMU-VALUE01;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1457583916] SERVICE ALERT: EMU-VALUE01;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457583921] SERVICE ALERT: EMU-VALUE01;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1457583923] SERVICE ALERT: DPS-SECURE02;CPU Load;OK;SOFT;2;CPU Load 8% (5 min average)
[1457583939] SERVICE ALERT: DPS-SECURE02;Service: Windows Time;OK;SOFT;3;W32Time: Started
[1457583946] HOST ALERT: EMU-VALUE01;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1457583949] SERVICE ALERT: EMU-VALUE01;Service: User Profile;CRITICAL;SOFT;2;connect to address 164.76.4.85 and port 12489: Connection refused
[1457583966] SERVICE ALERT: EMU-VALUE01;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.4.85 and port 12489: Connection refused
[1457583979] HOST ALERT: EMU-VALUE01;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1457584012] SERVICE ALERT: EMU-VALUE01;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.4.85 and port 12489: Connection refused
[1457584012] SERVICE ALERT: XDM;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.2.79 and port 12489: Connection refused
[1457584016] SERVICE ALERT: EMU-VALUE01;C:\ Drive Space;CRITICAL;HARD;3;connect to address 164.76.4.85 and port 12489: Connection refused
[1457584018] SERVICE ALERT: MBAM;CPU Load;OK;HARD;3;CPU Load 5% (5 min average)
[1457584021] SERVICE ALERT: EMU-VALUE01;Service: WMI;CRITICAL;HARD;3;connect to address 164.76.4.85 and port 12489: Connection refused
[1457584023] SERVICE ALERT: EMU-VALUE01;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.4.85 and port 12489: Connection refused
[1457584042] HOST ALERT: EMU-VALUE01;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1457584048] SERVICE ALERT: PROCARE;CPU Load;OK;HARD;3;CPU Load 5% (5 min average)
[1457584059] HOST ALERT: EMU-VALUE01;UP;SOFT;6;PING OK - Packet loss = 58%, RTA = 0.46 ms
[1457584069] SERVICE ALERT: EMU-VALUE01;Service: User Profile;OK;SOFT;3;ProfSvc: Started
[1457584086] SERVICE ALERT: EMU-VALUE01;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1457584132] SERVICE ALERT: EMU-VALUE01;Service: RPC;OK;SOFT;2;RpcSs: Started
[1457584132] SERVICE ALERT: XDM;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1457584143] SERVICE ALERT: EMU-VALUE01;Memory Usage;OK;SOFT;2;Memory usage: total:9444.43 MB - used: 2516.52 MB (27%) - free: 6927.91 MB (73%)
[1457584616] SERVICE ALERT: EMU-VALUE01;C:\ Drive Space;OK;HARD;3;c:\ - total: 465.42 Gb - used: 15.50 Gb (3%) - free 449.92 Gb (97%)
[1457584621] SERVICE ALERT: EMU-VALUE01;Service: WMI;OK;HARD;3;Winmgmt: Started
[1457584641] SERVICE ALERT: XDM;CPU Load;OK;HARD;3;CPU Load 77% (5 min average)
[1457585243] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 77% (5 min average)
[1457585719] SERVICE ALERT: FTK;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.102.240 and port 12489: No route to host
[1457585719] SERVICE ALERT: FTK;Uptime;CRITICAL;SOFT;1;connect to address 164.76.102.240 and port 12489: No route to host
[1457585722] HOST ALERT: FTK;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.240)
[1457585752] SERVICE ALERT: FTK;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.102.240 and port 12489: No route to host
[1457585755] HOST ALERT: FTK;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.240)
[1457585762] SERVICE ALERT: FTK;CPU Load;CRITICAL;SOFT;1;connect to address 164.76.102.240 and port 12489: No route to host
[1457585815] SERVICE ALERT: FTK;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.102.240 and port 12489: No route to host
[1457585818] HOST ALERT: FTK;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.102.240)
[1457585830] SERVICE ALERT: FTK;Uptime;CRITICAL;SOFT;2;connect to address 164.76.102.240 and port 12489: No route to host
[1457585833] HOST ALERT: FTK;DOWN;SOFT;4;CRITICAL - Host Unreachable (164.76.102.240)
[1457585839] SERVICE ALERT: FTK;C:\ Drive Space;CRITICAL;SOFT;2;connect to address 164.76.102.240 and port 12489: No route to host
[1457585865] SERVICE ALERT: FTK;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.102.240 and port 12489: No route to host
[1457585868] HOST ALERT: FTK;DOWN;SOFT;5;CRITICAL - Host Unreachable (164.76.102.240)
[1457585871] SERVICE ALERT: FTK;Service: Windows Time;CRITICAL;SOFT;2;connect to address 164.76.102.240 and port 12489: No route to host
[1457585871] SERVICE ALERT: FTK;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.102.240 and port 12489: No route to host
[1457585882] SERVICE ALERT: FTK;CPU Load;CRITICAL;SOFT;2;connect to address 164.76.102.240 and port 12489: No route to host
[1457585885] HOST ALERT: FTK;DOWN;SOFT;6;CRITICAL - Host Unreachable (164.76.102.240)
[1457585932] SERVICE ALERT: FTK;Memory Usage;OK;SOFT;2;Memory usage: total:65524.64 MB - used: 19030.60 MB (29%) - free: 46494.05 MB (71%)
[1457585940] HOST ALERT: FTK;UP;SOFT;7;PING OK - Packet loss = 73%, RTA = 0.38 ms
[1457585947] SERVICE ALERT: FTK;Uptime;OK;SOFT;3;System Uptime - 0 day(s) 0 hour(s) 1 minute(s)
[1457585948] SERVICE ALERT: FTK;C:\ Drive Space;OK;SOFT;3;c:\ - total: 408.28 Gb - used: 139.55 Gb (34%) - free 268.73 Gb (66%)
[1457585982] SERVICE ALERT: FTK;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1457585989] SERVICE ALERT: FTK;Service: Windows Time;OK;SOFT;3;W32Time: Started
[1457585989] SERVICE ALERT: FTK;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1457585999] SERVICE ALERT: FTK;CPU Load;OK;SOFT;3;CPU Load 4% (5 min average)