[1491105603] SERVICE ALERT: eateriespk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1491105663] SERVICE ALERT: eateriespk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1491105723] SERVICE ALERT: eateriespk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1491105846] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1491105906] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1491105966] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1491107166] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 39% (15 min average)
[1491107401] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1491108366] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1491108426] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1491108486] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1491109686] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 88% (15 min average)
[1491109949] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1491110009] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1491110070] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;3;OK. All replicas within 120 seconds.
[1491110669] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1491110731] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1491110790] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1491111989] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1491112086] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1491112146] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1491112206] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1491112590] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1491112649] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1491112708] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1491113179] SERVICE ALERT: ADDC5;AD Group Changes;UNKNOWN;SOFT;1;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1491113399] SERVICE ALERT: ADDC5;AD Group Changes;OK;SOFT;2;WARNING:  0  groups have changed!
[1491113908] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1491113908] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1491114508] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1491114568] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;2;Warning:  idm1  are more than 120 seconds behind!
[1491114606] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1491114628] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1491115828] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1491116428] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1491117006] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1491117066] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 92% (15 min average)
[1491117628] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1491117666] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1491117726] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1491117786] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1491118386] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 41% (15 min average)
[1491118829] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1491118986] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1491119046] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1491119106] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1491119706] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1491120306] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1491120366] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1491120426] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 90% (15 min average)
[1491120627] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1491121227] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1491121287] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1491121800] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service has exited from a period of scheduled downtime
[1491121803] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (35.5% change >= 20.0% threshold)
[1491121803] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (49.7% change >= 20.0% threshold)
[1491121803] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1491121803] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1491121803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1491121803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1491121803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1491121803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1491121803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1491121803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1491121803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1491121803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1491121803] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1491121803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1491121803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1491121803] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1491121803] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1491123603] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (29.3% change >= 20.0% threshold)
[1491123603] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (42.2% change >= 20.0% threshold)
[1491123603] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1491123603] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1491123603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1491123603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1491123603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1491123603] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1491123603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1491123603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1491123603] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1491123603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1491123603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1491123603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1491123603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1491123603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1491123603] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1491123687] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1491123747] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1491123807] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1491124407] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1491126207] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1491126267] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1491126327] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;3;OK. All replicas within 120 seconds.
[1491126927] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1491126988] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1491127626] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1491127686] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1491127746] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1491128346] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1491129546] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1491130146] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1491131946] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 85% (15 min average)
[1491135987] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1491136146] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1491136206] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1491139803] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1491139803] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1491139803] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1491139803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1491139803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1491139803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1491139803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1491139803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1491139803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1491139803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1491139803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1491139803] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1491139803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1491139803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1491139803] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1491139803] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1491142205] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1491148904] SERVICE ALERT: CAVA2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1491148964] SERVICE ALERT: CAVA2;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1491149024] SERVICE ALERT: CAVA2;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1491150224] SERVICE ALERT: CAVA2;CPU Load;OK;HARD;3;CPU Load 59% (15 min average)
[1491150603] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1491150603] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1491150603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1491150603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1491150603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1491150603] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1491150603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1491150603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1491150603] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1491150603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1491150603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1491150603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1491150603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1491150603] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1491150603] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1491150605] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1491150665] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1491150725] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1491151325] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 60% (15 min average)
[1491156123] SERVICE ALERT: eateriespk;Device Status;OK;HARD;3;There is no help available for this item. 54 6F 20 65 78 69 74 20 70 72 65 73 73 20 1F 2E
[1491156723] SERVICE ALERT: eateriespk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1491156725] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1491156783] SERVICE ALERT: eateriespk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1491156785] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 91% (15 min average)
[1491156843] SERVICE ALERT: eateriespk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1491157443] SERVICE ALERT: eateriespk;Device Status;OK;HARD;3;Clearing paper path
[1491159604] SERVICE ALERT: ADWSUSSERVERS;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.4.79 and port 12489: Connection refused
[1491159664] SERVICE ALERT: ADWSUSSERVERS;Service: RPC;OK;SOFT;2;RpcSs: Started
[1491162049] HOST ALERT: studentcenterpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1491162139] HOST ALERT: studentcenterpk;UP;SOFT;2;PING WARNING - Packet loss = 86%, RTA = 0.51 ms
[1491162599] HOST ALERT: studentcenterpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1491162649] HOST ALERT: studentcenterpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1491162739] HOST ALERT: studentcenterpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1491162829] HOST ALERT: studentcenterpk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1491162918] HOST ALERT: studentcenterpk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1491163008] HOST ALERT: studentcenterpk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1491163059] HOST ALERT: studentcenterpk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1491163149] HOST ALERT: studentcenterpk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1491163199] HOST ALERT: studentcenterpk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1491163249] HOST ALERT: studentcenterpk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1491167756] HOST ALERT: studentcenterpk;UP;HARD;10;PING OK - Packet loss = 37%, RTA = 0.52 ms
[1491167859] HOST ALERT: studentcenterpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1491167923] HOST ALERT: studentcenterpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.49 ms
[1491167999] HOST ALERT: studentcenterpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1491167999] HOST FLAPPING ALERT: studentcenterpk;STARTED; Host appears to have started flapping (23.7% change > 20.0% threshold)
[1491168042] HOST ALERT: studentcenterpk;UP;SOFT;2;PING OK - Packet loss = 79%, RTA = 0.57 ms
[1491168185] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1491168245] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1491168305] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 75% (15 min average)
[1491168372] HOST ALERT: studentcenterpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1491168459] HOST ALERT: studentcenterpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1491168549] HOST ALERT: studentcenterpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1491168599] HOST ALERT: studentcenterpk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1491168649] HOST ALERT: studentcenterpk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1491168738] HOST ALERT: studentcenterpk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1491168828] HOST ALERT: studentcenterpk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1491168905] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1491168918] HOST ALERT: studentcenterpk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1491168965] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1491169008] HOST ALERT: studentcenterpk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1491169024] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1491169059] HOST ALERT: studentcenterpk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1491169629] HOST FLAPPING ALERT: studentcenterpk;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1491170224] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 86% (15 min average)
[1491170513] SERVICE ALERT: cava2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1491170573] SERVICE ALERT: cava2;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1491170633] SERVICE ALERT: cava2;CPU Load;OK;SOFT;3;CPU Load 12% (15 min average)
[1491170824] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1491170884] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1491170944] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 85% (15 min average)
[1491171233] SERVICE ALERT: cava2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1491171293] SERVICE ALERT: cava2;CPU Load;OK;SOFT;2;CPU Load 9% (15 min average)
[1491171544] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1491171604] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 84% (15 min average)
[1491173093] SERVICE ALERT: cava2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1491173153] SERVICE ALERT: cava2;CPU Load;OK;SOFT;2;CPU Load 10% (15 min average)
[1491174004] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1491174064] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1491174124] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 86% (15 min average)
[1491174183] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5179 found since last run. oneLevelSearchOps OK. searchOps is too high! 5335 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491174243] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491174843] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5476 found since last run. oneLevelSearchOps OK. searchOps is too high! 5544 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491174903] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491175924] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1491175984] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1491176044] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1491177244] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1491177244] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1491177844] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1491177903] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6344 found since last run. oneLevelSearchOps OK. searchOps is too high! 6552 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491177904] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1491177963] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491178563] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5151 found since last run. oneLevelSearchOps OK. searchOps is too high! 5290 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491178623] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491179704] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1491179764] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1491179824] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1491179902] SERVICE ALERT: eateriespk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.59.141. Is the printer off?  Is the network down?
[1491179932] HOST ALERT: eateriespk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1491179962] SERVICE ALERT: eateriespk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.59.141. Is the printer off?  Is the network down?
[1491179992] HOST ALERT: eateriespk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1491180022] SERVICE ALERT: eateriespk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.59.141. Is the printer off?  Is the network down?
[1491180052] HOST ALERT: eateriespk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1491180142] HOST ALERT: eateriespk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1491180232] HOST ALERT: eateriespk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1491180249] SERVICE ALERT: eateriespk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.59.141. Is the printer off?  Is the network down?
[1491180279] HOST ALERT: eateriespk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1491180309] SERVICE ALERT: eateriespk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.59.141. Is the printer off?  Is the network down?
[1491180339] HOST ALERT: eateriespk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1491180369] SERVICE ALERT: eateriespk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.59.141. Is the printer off?  Is the network down?
[1491180399] HOST ALERT: eateriespk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1491180399] SERVICE ALERT: eateriespk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.59.141. Is the printer off?  Is the network down?
[1491180429] HOST ALERT: eateriespk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1491180459] SERVICE ALERT: eateriespk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.59.141. Is the printer off?  Is the network down?
[1491180489] HOST ALERT: eateriespk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1491180519] SERVICE ALERT: eateriespk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.59.141. Is the printer off?  Is the network down?
[1491182417] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1491182421] HOST ALERT: eateriespk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.45 ms
[1491182740] SERVICE ALERT: TARGETX-2012;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1491182763] SERVICE ALERT: eateriespk;Device Status;OK;HARD;3;TRAY 3 EMPTY PLAIN LETTER   Ready
[1491182800] SERVICE ALERT: TARGETX-2012;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1491182824] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 74% (15 min average)
[1491182860] SERVICE ALERT: TARGETX-2012;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1491182914] SERVICE ALERT: eateriespk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 33% - OK!
[1491183000] HOST DOWNTIME ALERT: IMAGES;STARTED; Host has entered a period of scheduled downtime
[1491183460] SERVICE ALERT: TARGETX-2012;CPU Load;OK;HARD;3;CPU Load 87% (15 min average)
[1491183599] HOST ALERT: studentcenterpk;UP;HARD;10;PING WARNING - Packet loss = 96%, RTA = 0.49 ms
[1491183965] HOST ALERT: studentcenterpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1491184033] SERVICE ALERT: IMAGES;MDT Monitor;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1491184040] HOST ALERT: IMAGES;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1491184055] HOST ALERT: studentcenterpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1491184104] HOST ALERT: IMAGES;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1491184145] HOST ALERT: studentcenterpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1491184199] HOST ALERT: studentcenterpk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1491184245] HOST ALERT: studentcenterpk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1491184335] HOST ALERT: studentcenterpk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1491184425] HOST ALERT: studentcenterpk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1491184515] HOST ALERT: studentcenterpk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1491184605] HOST ALERT: studentcenterpk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1491184613] SERVICE ALERT: IMAGES;MDT Monitor;UNKNOWN;HARD;3;Failed to open service MDT_Monitor: 1060: The specified service does not exist as an installed service.
[1491184623] SERVICE ALERT: studentcenterpk;Device Status;OK;HARD;3;Warming up
[1491184624] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1491184627] HOST ALERT: studentcenterpk;UP;SOFT;10;PING OK - Packet loss = 0%, RTA = 1.11 ms
[1491184684] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 84% (15 min average)
[1491185229] SERVICE ALERT: studentcenterpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.150.252. Is the printer off?  Is the network down?
[1491185284] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1491185289] SERVICE ALERT: studentcenterpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.150.252. Is the printer off?  Is the network down?
[1491185344] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1491185349] SERVICE ALERT: studentcenterpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.150.252. Is the printer off?  Is the network down?
[1491185404] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 94% (15 min average)
[1491185416] SERVICE ALERT: eateriespk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1491185445] HOST ALERT: studentcenterpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1491185535] HOST ALERT: studentcenterpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1491185599] HOST ALERT: studentcenterpk;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.47 ms
[1491186017] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1491186077] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1491186137] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1491186137] SERVICE FLAPPING ALERT: eateriespk;Paper Tray Status;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1491186604] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1491186664] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1491186724] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1491187023] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5339 found since last run. oneLevelSearchOps OK. searchOps is too high! 5555 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1491187083] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1491187324] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 68% (15 min average)
[1491189124] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1491189137] SERVICE ALERT: eateriespk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1491189184] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1491189243] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1491189843] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (15 min average)
[1491190200] HOST DOWNTIME ALERT: IMAGES;STOPPED; Host has exited from a period of scheduled downtime
[1491191043] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1491191103] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1491191163] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 92% (15 min average)
[1491191763] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1491191823] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 70% (15 min average)