[1455862492] HOST ALERT: halleground2;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.57 ms
[1455862517] SERVICE ALERT: halleground2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1455862577] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 68% - OK!
[1455862752] SERVICE ALERT: halleground2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1455883907] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1455883967] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1455884026] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1455884952] SERVICE ALERT: halleground2;Device Status;OK;HARD;3;Processing job from tray 4
[1455886426] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1455888098] SERVICE ALERT: mckennypk;Device Status;OK;HARD;3;Ready 164.76.152.26
[1455888100] HOST ALERT: mckennypk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1455888244] SERVICE ALERT: mckennypk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1455888360] SERVICE ALERT: mckennypk;Black Cartridge HP CC364X;OK;HARD;3;Black Cartridge HP CC364X is at 16% - OK!
[1455898086] SERVICE ALERT: ADDC5;GPO Replication;WARNING;SOFT;1;Warning! 1 policies are not replicating properly!
[1455898389] SERVICE ALERT: ADDC5;GPO Replication;WARNING;SOFT;2;Warning! 1 policies are not replicating properly!
[1455898686] SERVICE ALERT: ADDC5;GPO Replication;WARNING;HARD;3;Warning! 1 policies are not replicating properly!
[1455902668] SERVICE ALERT: OFFICEKMS;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1455902788] SERVICE ALERT: OFFICEKMS;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1455903166] SERVICE ALERT: EVENTSTEST-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1455903286] SERVICE ALERT: EVENTSTEST-2012;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1455905888] SERVICE ALERT: ADDC5;GPO Replication;OK;HARD;3;OK! All GPOs are properly replicated.
[1455905888] SERVICE FLAPPING ALERT: ADDC5;GPO Replication;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1455908164] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1455908223] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1455908283] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1455912649] SERVICE FLAPPING ALERT: ADDC5;GPO Replication;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1455913778] SERVICE ALERT: ADCA1;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1455913803] SERVICE FLAPPING ALERT: ADDC5;GPO Replication;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1455913898] SERVICE ALERT: ADCA1;CPU Load;OK;SOFT;2;CPU Load 5% (5 min average)
[1455923883] SERVICE ALERT: studentcenterpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1455924844] SERVICE ALERT: DB01-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1455924964] SERVICE ALERT: DB01-2012;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1455934443] SERVICE ALERT: SANDBOX-DB;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1455934484] SERVICE ALERT: EVENTSTEST-2012;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1455934562] SERVICE ALERT: SANDBOX-DB;CPU Load;OK;SOFT;2;CPU Load 80% (5 min average)
[1455934604] SERVICE ALERT: EVENTSTEST-2012;CPU Load;OK;SOFT;2;CPU Load 77% (5 min average)
[1455940248] SERVICE ALERT: ADWSUSSERVERS;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1455940342] SERVICE ALERT: ADWSUSSERVERS;Service: Firewall;OK;SOFT;2;MpsSvc: Started