[1481950919] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481950979] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481951039] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1481951348] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 69% (15 min average)
[1481951639] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1481951948] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1481952068] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 75% (15 min average)
[1481952239] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481952299] SERVICE ALERT: rackhampk;Device Status;OK;SOFT;2;Open tray indicated.
[1481952668] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1481952788] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1481952899] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481952908] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 32% (15 min average)
[1481952960] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481953019] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1481953619] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1481954015] SERVICE ALERT: ADWU1;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.4.133 and port 12489: Connection refused
[1481954126] SERVICE ALERT: ADWU1;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1481954219] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481954279] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481954334] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;1;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1481954339] SERVICE ALERT: rackhampk;Device Status;OK;SOFT;3;Open tray indicated.
[1481954454] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;2;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1481954574] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;HARD;3;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1481954939] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481955000] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481955059] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1481955659] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1481955814] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1481955908] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1481955934] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1481956028] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1481956053] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1481956148] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1481956200] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1481956259] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481956319] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481956379] SERVICE ALERT: rackhampk;Device Status;OK;SOFT;3;Open tray indicated.
[1481956979] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481957039] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481957100] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1481957699] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1481957852] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1481957948] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 58% (15 min average)
[1481958299] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481958359] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481958419] SERVICE ALERT: rackhampk;Device Status;OK;SOFT;3;Open tray indicated.
[1481958452] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1481958863] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 146 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481958983] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481959019] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481959052] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1481959079] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481959139] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1481959148] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1481959172] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1481959268] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1481959292] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1481959292] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1481959388] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 85% (15 min average)
[1481959740] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1481959788] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 112 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481959908] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481959988] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1481960108] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 81% (15 min average)
[1481960183] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 150 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481960303] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481960339] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481960399] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481960459] SERVICE ALERT: rackhampk;Device Status;OK;SOFT;3;Open tray indicated.
[1481960492] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1481960508] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 115 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481960628] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481960708] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1481960828] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1481960948] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 94% (15 min average)
[1481961059] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481961119] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481961179] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1481961228] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 126 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481961348] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481961503] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 134 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481961549] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1481961623] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481961668] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1481961692] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1481961774] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;OK;HARD;3;OK: No entries found
[1481961780] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1481961788] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 94% (15 min average)
[1481961948] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 120 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481962068] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1481962292] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1481962379] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481962388] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1481962413] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;2;Warning:  idm1  are more than 120 seconds behind!
[1481962439] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481962499] SERVICE ALERT: rackhampk;Device Status;OK;SOFT;3;Open tray indicated.
[1481962508] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1481962532] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1481962628] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1481963099] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481963159] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481963219] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1481963228] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (15 min average)
[1481963732] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1481963819] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1481963828] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1481963948] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1481964068] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1481964420] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481964479] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481964539] SERVICE ALERT: rackhampk;Device Status;OK;SOFT;3;Open tray indicated.
[1481964932] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1481965052] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1481965139] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481965199] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481965259] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1481965268] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1481965859] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1481965868] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1481966460] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481966468] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 64% (15 min average)
[1481966519] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481966579] SERVICE ALERT: rackhampk;Device Status;OK;SOFT;3;Open tray indicated.
[1481967003] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (44.1% change >= 20.0% threshold)
[1481967003] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (38.7% change >= 20.0% threshold)
[1481967003] SERVICE FLAPPING ALERT: rackhampk;Device Status;STARTED; Service appears to have started flapping (64.5% change >= 20.0% threshold)
[1481967003] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1481967003] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1481967003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1481967003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1481967003] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1481967003] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1481967003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1481967003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1481967003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1481967003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1481967003] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1481967003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1481967003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1481967003] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1481967003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1481967003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1481967003] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1481967068] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1481967179] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481967188] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1481967239] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481967300] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1481967308] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 90% (15 min average)
[1481967899] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1481968508] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1481968628] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1481968748] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1481968803] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (47.2% change >= 20.0% threshold)
[1481968803] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (31.9% change >= 20.0% threshold)
[1481968803] SERVICE FLAPPING ALERT: rackhampk;Device Status;STARTED; Service appears to have started flapping (64.5% change >= 20.0% threshold)
[1481968803] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1481968803] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1481968803] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1481968803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1481968803] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1481968803] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1481968803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1481968803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1481968803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1481968803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1481968803] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1481968803] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1481968803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1481968803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1481968803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1481968803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1481968803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1481969098] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481969158] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481969218] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1481969348] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 78% (15 min average)
[1481969818] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1481970419] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481970478] SERVICE ALERT: rackhampk;Device Status;OK;SOFT;2;Open tray indicated. TRAY 3
[1481970548] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1481970668] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1481970788] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 81% (15 min average)
[1481971078] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481971138] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481971198] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1481971798] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1481971988] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1481972108] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 90% (15 min average)
[1481972398] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481972459] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481972518] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1481972708] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1481972828] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1481972948] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1481973548] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1481973717] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1481974200] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service has exited from a period of scheduled downtime
[1481974317] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481974377] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481974437] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1481974652] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1481974747] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1481975037] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1481975347] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1481976237] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481976296] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481976356] SERVICE ALERT: rackhampk;Device Status;OK;SOFT;3;Open tray indicated.
[1481976956] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481977016] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481977076] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1481977676] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1481978277] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481978336] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481978396] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1481979596] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1481980196] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481980256] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481980317] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1481980743] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 67% (15 min average)
[1481980916] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1481981343] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1481981463] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1481981583] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 67% (15 min average)
[1481982116] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481982176] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481982183] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1481982236] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1481982303] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1481982423] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1481982836] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1481983023] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1481983437] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481983496] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481983556] SERVICE ALERT: rackhampk;Device Status;OK;SOFT;3;Open tray indicated.
[1481984156] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481984216] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481984223] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1481984276] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1481984343] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1481984876] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1481984943] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1481985003] SERVICE FLAPPING ALERT: rackhampk;Device Status;STARTED; Service appears to have started flapping (70.5% change >= 20.0% threshold)
[1481985003] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1481985003] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1481985003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1481985003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1481985003] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1481985003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1481985003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1481985003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1481985003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1481985003] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1481985003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1481985003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1481985003] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1481985003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1481985003] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1481985003] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1481985063] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1481985183] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1481985183] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (39.3% change >= 20.0% threshold)
[1481985476] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481985536] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481985597] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1481985783] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1481986383] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1481986503] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1481986623] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1481986796] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1481987223] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 79% (15 min average)
[1481987396] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481987456] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481987516] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1481987823] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1481987943] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1481988063] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1481988116] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1481988716] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481988777] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481988836] SERVICE ALERT: rackhampk;Device Status;OK;SOFT;3;Open tray indicated.
[1481989436] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481989496] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481989556] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1481990156] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1481990463] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1481990756] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481990817] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481990876] SERVICE ALERT: rackhampk;Device Status;OK;SOFT;3;Open tray indicated.
[1481991063] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1481991183] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1481991303] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1481991476] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481991536] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481991596] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1481991903] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 91% (15 min average)
[1481992196] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1481992503] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1481992623] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1481992743] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 74% (15 min average)
[1481992796] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481992856] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481992917] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1481993943] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1481994063] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1481994116] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1481994175] SERVICE ALERT: CAVA2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1481994183] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1481994295] SERVICE ALERT: CAVA2;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1481994415] SERVICE ALERT: CAVA2;CPU Load;OK;SOFT;3;CPU Load 2% (15 min average)
[1481994716] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481994776] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481994783] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1481994836] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1481995383] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1481995436] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1481995803] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (65.1% change >= 20.0% threshold)
[1481995803] SERVICE FLAPPING ALERT: rackhampk;Device Status;STARTED; Service appears to have started flapping (75.6% change >= 20.0% threshold)
[1481995803] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1481995803] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1481995803] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1481995803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1481995803] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1481995803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1481995803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1481995803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1481995803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1481995803] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1481995803] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1481995803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1481995803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1481995803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1481995803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1481995803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1481996636] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481996696] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481996756] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1481997357] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1481998556] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481998616] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481998676] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1481999276] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1481999876] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1481999937] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1481999996] SERVICE ALERT: rackhampk;Device Status;OK;SOFT;3;Open tray indicated.
[1482000183] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1482000595] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482000655] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482000715] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482000783] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1482000903] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1482001023] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1482001315] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1482001915] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482001975] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482002036] SERVICE ALERT: rackhampk;Device Status;OK;SOFT;3;Open tray indicated.
[1482002223] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 71% (15 min average)
[1482002635] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482002695] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482002755] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482002823] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1482002943] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1482003063] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1482003355] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1482003955] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482004015] SERVICE ALERT: rackhampk;Device Status;OK;SOFT;2;Open tray indicated.
[1482004263] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (15 min average)
[1482004616] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482004675] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482004734] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482005463] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1482005583] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1482005703] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 84% (15 min average)
[1482005934] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1482006303] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1482006423] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1482006534] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482006543] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1482006594] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482006654] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482007255] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1482007854] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482007914] SERVICE ALERT: rackhampk;Device Status;OK;SOFT;2;Open tray indicated.
[1482008514] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482008574] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482008634] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482008943] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1482009064] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482009124] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482009184] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482009234] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1482009468] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 100 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482009588] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482009784] SERVICE ALERT: hallesouthcommons;Device Status;OK;HARD;3;Open top cover.
[1482009834] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482009895] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482009954] SERVICE ALERT: rackhampk;Device Status;OK;SOFT;3;Open tray indicated.
[1482010554] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482010614] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482010674] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482010984] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482011044] SERVICE ALERT: hallesouthcommons;Device Status;OK;SOFT;2;Open top cover.
[1482011274] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1482011459] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1482011519] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1482011578] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1482011874] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482011935] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482011943] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1482011994] SERVICE ALERT: rackhampk;Device Status;OK;SOFT;3;Open tray indicated.
[1482012063] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1482012183] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1482012244] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482012304] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482012364] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482012594] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482012654] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482012714] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482012964] SERVICE ALERT: hallesouthcommons;Device Status;OK;HARD;3;Open top cover.
[1482012964] SERVICE FLAPPING ALERT: hallesouthcommons;Device Status;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1482013314] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1482013788] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 105 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482013908] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482013914] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482013974] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482014035] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482014164] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482014223] SERVICE ALERT: hallesouthcommons;Device Status;OK;SOFT;2;Open top cover.
[1482015023] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 107 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482015143] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1482015234] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1482015423] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482015483] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482015543] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482015775] SERVICE ALERT: cava2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1482015834] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482015894] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482015895] SERVICE ALERT: cava2;CPU Load;OK;SOFT;2;CPU Load 76% (15 min average)
[1482015954] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482016143] SERVICE ALERT: hallesouthcommons;Device Status;OK;HARD;3;Open top cover.
[1482016554] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1482017155] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482017214] SERVICE ALERT: rackhampk;Device Status;OK;SOFT;2;Open tray indicated.
[1482017343] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482017403] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482017463] SERVICE ALERT: hallesouthcommons;Device Status;OK;SOFT;3;Open top cover.
[1482017814] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482017874] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482017934] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482018183] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 57% (15 min average)
[1482018534] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1482018663] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482018723] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482018783] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1482018784] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482018903] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 90% (15 min average)
[1482019134] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482019194] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482019255] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482019383] SERVICE ALERT: hallesouthcommons;Device Status;OK;HARD;3;Open top cover.
[1482019503] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1482019623] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1482019743] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1482019983] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482020043] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482020103] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482020454] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1482020703] SERVICE ALERT: hallesouthcommons;Device Status;OK;HARD;3;Open top cover.
[1482021054] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482021114] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482021174] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482021774] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1482021903] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482021964] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482022023] SERVICE ALERT: hallesouthcommons;Device Status;OK;SOFT;3;Open top cover.
[1482022743] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1482022974] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482023033] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482023093] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482023223] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482023283] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482023343] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1482023343] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482023463] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1482023583] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1482023693] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1482023943] SERVICE ALERT: hallesouthcommons;Device Status;OK;HARD;3;Open top cover.
[1482024893] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482024953] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482025014] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482025144] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482025203] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482025263] SERVICE ALERT: hallesouthcommons;Device Status;OK;SOFT;3;Open top cover.
[1482025613] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1482026213] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482026273] SERVICE ALERT: rackhampk;Device Status;OK;SOFT;2;Open tray indicated.
[1482026463] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482026523] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482026583] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482026873] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482026933] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482026993] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482027183] SERVICE ALERT: hallesouthcommons;Device Status;OK;HARD;3;Open top cover.
[1482027569] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;SOFT;1;e:\ - total: 799.87 Gb - used: 650.35 Gb (81%) - free 149.52 Gb (19%)
[1482027593] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1482027689] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;SOFT;2;e:\ - total: 799.87 Gb - used: 650.35 Gb (81%) - free 149.52 Gb (19%)
[1482027809] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;HARD;3;e:\ - total: 799.87 Gb - used: 650.35 Gb (81%) - free 149.52 Gb (19%)
[1482028194] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482028253] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482028313] SERVICE ALERT: rackhampk;Device Status;OK;SOFT;3;Open tray indicated.
[1482028383] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482028443] SERVICE ALERT: hallesouthcommons;Device Status;OK;SOFT;2;Open top cover.
[1482028494] SERVICE ALERT: cava2;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1482028614] SERVICE ALERT: cava2;CPU Load;OK;SOFT;2;CPU Load 48% (15 min average)
[1482028913] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482028973] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482028983] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1482029009] SERVICE ALERT: AD-RMAD;E:\ Drive Space;OK;HARD;3;e:\ - total: 799.87 Gb - used: 625.54 Gb (78%) - free 174.34 Gb (22%)
[1482029033] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482029583] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1482029633] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1482029643] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482029703] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482029763] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482030233] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482030294] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482030353] SERVICE ALERT: rackhampk;Device Status;OK;SOFT;3;Open tray indicated.
[1482030363] SERVICE ALERT: hallesouthcommons;Device Status;OK;HARD;3;Open top cover.
[1482030953] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482031013] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482031073] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482031389] SERVICE ALERT: epprogsql2;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 249.66 Gb - used: 225.68 Gb (90%) - free 23.98 Gb (10%)
[1482031509] SERVICE ALERT: epprogsql2;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 249.66 Gb - used: 225.68 Gb (90%) - free 23.98 Gb (10%)
[1482031563] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482031623] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482031629] SERVICE ALERT: epprogsql2;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 249.66 Gb - used: 225.68 Gb (90%) - free 23.98 Gb (10%)
[1482031673] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1482031683] SERVICE ALERT: hallesouthcommons;Device Status;OK;SOFT;3;Open top cover.
[1482031983] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 84% (15 min average)
[1482032273] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482032333] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482032394] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482032583] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1482032703] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1482032823] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 92% (15 min average)
[1482032883] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482032943] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482033003] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482033423] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1482033543] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1482033593] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1482033604] SERVICE ALERT: hallesouthcommons;Device Status;OK;HARD;3;Open top cover.
[1482033663] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 86% (15 min average)
[1482034193] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482034203] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482034253] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482034263] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482034263] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1482034313] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482034323] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482034383] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1482034503] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1482034913] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1482034923] SERVICE ALERT: hallesouthcommons;Device Status;OK;HARD;3;Open top cover.
[1482035103] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 72% (15 min average)
[1482035513] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482035574] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482035633] SERVICE ALERT: rackhampk;Device Status;OK;SOFT;3;Open tray indicated.
[1482035703] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1482035823] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 64% (15 min average)
[1482035829] SERVICE ALERT: epprogsql2;C:\ Drive Space;OK;HARD;3;c:\ - total: 249.66 Gb - used: 219.80 Gb (88%) - free 29.86 Gb (12%)
[1482036123] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482036184] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482036233] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1482036243] SERVICE ALERT: hallesouthcommons;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482036293] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1482036353] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1482036423] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1482036543] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1482036663] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1482036843] SERVICE ALERT: hallesouthcommons;Device Status;OK;HARD;3;Open top cover.
[1482036953] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.