[1463458845] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463459444] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1463460044] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463460644] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463464350] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1463464470] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1463464590] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 92% (15 min average)
[1463464788] SERVICE ALERT: MBAM;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463464908] SERVICE ALERT: MBAM;CPU Load;OK;SOFT;2;CPU Load 1% (15 min average)
[1463465444] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463466044] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463466644] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463466990] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1463467110] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1463467230] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1463467245] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463467830] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463467830] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1463467844] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463468430] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1463469030] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1463469150] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1463469270] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 90% (15 min average)
[1463469644] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463470244] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463470470] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1463470590] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 80% (15 min average)
[1463470844] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1463473244] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 84% (15 min average)
[1463475044] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463475164] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1463475284] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 79% (15 min average)
[1463479484] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1463479604] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463479724] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463480190] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1463480324] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463480924] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 84% (15 min average)
[1463481187] SERVICE ALERT: INTLAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1463481307] SERVICE ALERT: INTLAPP;CPU Load;OK;SOFT;2;CPU Load 0% (15 min average)
[1463481524] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463481644] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1463481764] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 80% (15 min average)
[1463482964] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1463483084] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463483204] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463483804] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463484404] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 80% (15 min average)
[1463484500] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1463486190] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1463486207] SERVICE ALERT: epprogsql2;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463486237] HOST ALERT: epprogsql2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1463486259] SERVICE ALERT: epprogsql2;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463486260] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463486260] SERVICE ALERT: epprogsql2;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463486289] HOST ALERT: epprogsql2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1463486310] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 89% (15 min average)
[1463486327] SERVICE ALERT: epprogsql2;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463486348] SERVICE ALERT: epprogsql2;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463486351] HOST DOWNTIME ALERT: epprogsql2;STARTED; Host has entered a period of scheduled downtime
[1463486357] HOST ALERT: epprogsql2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1463486379] SERVICE ALERT: epprogsql2;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463486380] SERVICE ALERT: epprogsql2;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463486380] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463486409] HOST ALERT: epprogsql2;DOWN;SOFT;4;(Host check timed out after 30.00 seconds)
[1463486434] SERVICE ALERT: epprogsql2;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463486447] SERVICE ALERT: epprogsql2;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1463486464] HOST ALERT: epprogsql2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1463486465] SERVICE ALERT: epprogsql2;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463486468] SERVICE ALERT: epprogsql2;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463486495] HOST ALERT: epprogsql2;DOWN;SOFT;6;(Host check timed out after 30.00 seconds)
[1463486499] SERVICE ALERT: epprogsql2;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1463486500] SERVICE ALERT: epprogsql2;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1463486500] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1463486529] HOST ALERT: epprogsql2;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1463486554] SERVICE ALERT: epprogsql2;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463486570] SERVICE ALERT: epprogsql2;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463486584] HOST ALERT: epprogsql2;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1463486585] SERVICE ALERT: epprogsql2;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463486587] SERVICE ALERT: epprogsql2;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1463486589] SERVICE ALERT: epprogsql2;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463486607] SERVICE ALERT: epprogsql2;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463486615] HOST ALERT: epprogsql2;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1463486673] SERVICE ALERT: epprogsql2;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463486673] SERVICE ALERT: epprogsql2;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463486674] SERVICE ALERT: epprogsql2;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1463486690] SERVICE ALERT: epprogsql2;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463486703] HOST ALERT: epprogsql2;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1463486705] SERVICE ALERT: epprogsql2;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1463486707] SERVICE ALERT: epprogsql2;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463486709] SERVICE ALERT: epprogsql2;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463486727] SERVICE ALERT: epprogsql2;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463486780] SERVICE ALERT: epprogsql2;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463486793] SERVICE ALERT: epprogsql2;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463486793] SERVICE ALERT: epprogsql2;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463486810] SERVICE ALERT: epprogsql2;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1463486827] SERVICE ALERT: epprogsql2;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463486829] SERVICE ALERT: epprogsql2;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1463486847] SERVICE ALERT: epprogsql2;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1463486900] SERVICE ALERT: epprogsql2;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463486913] SERVICE ALERT: epprogsql2;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1463486913] SERVICE ALERT: epprogsql2;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1463486947] SERVICE ALERT: epprogsql2;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1463487020] SERVICE ALERT: epprogsql2;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1463487313] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 67836 found since last run. oneLevelSearchOps OK. searchOps is too high! 67837 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463487404] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463487434] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 28926 found since last run. oneLevelSearchOps OK. searchOps is too high! 28927 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463487524] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 92% (15 min average)
[1463487553] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 31570 found since last run. oneLevelSearchOps OK. searchOps is too high! 31571 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463488394] HOST DOWNTIME ALERT: ADAUTHDC1;STARTED; Host has entered a period of scheduled downtime
[1463489353] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463492324] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463492444] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463492564] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463493639] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;SOFT;1;No data was received from host!
[1463493740] SERVICE ALERT: BOEAPP;C:\ Drive Space;OK;SOFT;2;c:\ - total: 219.90 Gb - used: 162.11 Gb (74%) - free 57.79 Gb (26%)
[1463494964] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1463495564] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1463495684] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1463495804] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463497153] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 46847 found since last run. oneLevelSearchOps OK. searchOps is too high! 46848 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463497275] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 28140 found since last run. oneLevelSearchOps OK. searchOps is too high! 28141 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463497393] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 29060 found since last run. oneLevelSearchOps OK. searchOps is too high! 29061 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463499194] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463499194] SERVICE FLAPPING ALERT: idm1.emich.edu;LDAP Stats;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1463502648] SERVICE ALERT: cava2;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.7.192 and port 12489: Connection refused
[1463502678] HOST ALERT: cava2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1463502683] SERVICE ALERT: cava2;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463502683] SERVICE ALERT: cava2;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463502687] SERVICE ALERT: cava2;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463502697] SERVICE ALERT: cava2;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463502701] SERVICE ALERT: cava2;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1463502713] HOST ALERT: cava2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1463502765] SERVICE ALERT: cava2;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1463502769] HOST ALERT: cava2;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1463502783] SERVICE ALERT: cava2;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 2 minute(s)
[1463502783] SERVICE ALERT: cava2;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1463502787] SERVICE ALERT: cava2;VMWare Tools;OK;SOFT;2;VMTools: Started
[1463502797] SERVICE ALERT: cava2;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1463502820] SERVICE ALERT: cava2;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1463502839] SERVICE ALERT: cava2;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463502959] SERVICE ALERT: cava2;CPU Load;OK;SOFT;2;CPU Load 0% (15 min average)
[1463503173] HOST ALERT: cava2;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1463503226] SERVICE ALERT: cava2;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463503256] HOST ALERT: cava2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1463503257] HOST ALERT: crowd2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1463503320] HOST ALERT: cava2;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1463503321] HOST ALERT: crowd2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1463503327] SERVICE ALERT: cava2;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1463503819] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1463503874] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1463503879] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1463503934] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1463503934] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1463503939] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1463503942] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1463503985] SERVICE ALERT: cava2;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463503994] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1463503994] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1463504002] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1463504003] SERVICE ALERT: cava2;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463504003] SERVICE ALERT: cava2;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463504007] SERVICE ALERT: cava2;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463504015] HOST ALERT: cava2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1463504016] SERVICE ALERT: cava2;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463504021] SERVICE ALERT: cava2;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1463504046] HOST ALERT: cava2;DOWN;SOFT;2;(Host check timed out after 30.00 seconds)
[1463504054] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1463504055] HOST ALERT: cava2;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1463504062] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1463504063] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1463504085] SERVICE ALERT: cava2;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1463504099] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1463504103] SERVICE ALERT: cava2;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1463504103] SERVICE ALERT: cava2;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 2 minute(s)
[1463504107] SERVICE ALERT: cava2;VMWare Tools;OK;SOFT;2;VMTools: Started
[1463504109] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463504117] SERVICE ALERT: cava2;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1463504123] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1463504140] SERVICE ALERT: cava2;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1463504150] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1463504159] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1463504169] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463504183] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1463504189] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463504203] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1463504210] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1463504219] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1463504229] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463504249] SERVICE ALERT: hallefirstcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463504249] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463504270] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1463504282] SERVICE ALERT: hallefirst3;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1463504306] SERVICE ALERT: hallefirstcolor;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463504309] SERVICE ALERT: hallefirstcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463504309] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463504342] SERVICE ALERT: hallefirst3;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1463504362] SERVICE ALERT: hallefirst3;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1463504366] SERVICE ALERT: hallefirstcolor;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463504369] SERVICE ALERT: hallefirstcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463504402] SERVICE ALERT: hallefirst3;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1463504422] SERVICE ALERT: hallefirst3;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1463504426] SERVICE ALERT: hallefirstcolor;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463504434] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463504439] SERVICE ALERT: hallefirstcolor;Toner Collection Unit HP CE265A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463504465] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463504473] SERVICE ALERT: hallefirstcolor;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463504482] SERVICE ALERT: hallefirst3;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1463504494] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463504499] SERVICE ALERT: hallefirstcolor;Toner Collection Unit HP CE265A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463504525] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463504532] SERVICE ALERT: hallefirstcolor;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463504539] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463504554] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463504559] SERVICE ALERT: hallefirstcolor;Toner Collection Unit HP CE265A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463504585] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463504592] SERVICE ALERT: hallefirstcolor;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463504599] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463504659] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463504684] HOST ALERT: hallefirst1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1463504748] HOST ALERT: hallefirst1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.53 ms
[1463505459] HOST ALERT: hallefirstcolor;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1463505513] HOST ALERT: hallefirstcolor;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.44 ms
[1463505683] SERVICE ALERT: FITNESSGRAM;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463505713] HOST ALERT: FITNESSGRAM;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1463505727] SERVICE ALERT: FITNESSGRAM;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463505738] SERVICE ALERT: FITNESSGRAM;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463505757] HOST ALERT: FITNESSGRAM;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1463505803] SERVICE ALERT: FITNESSGRAM;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463505833] HOST ALERT: FITNESSGRAM;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1463505847] SERVICE ALERT: FITNESSGRAM;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463505848] SERVICE ALERT: CEDARON;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463505857] SERVICE ALERT: FITNESSGRAM;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463505867] SERVICE ALERT: CEDARON;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463505877] HOST ALERT: FITNESSGRAM;DOWN;SOFT;4;(Host check timed out after 30.00 seconds)
[1463505878] HOST ALERT: CEDARON;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1463505902] HOST ALERT: FITNESSGRAM;UP;SOFT;5;PING OK - Packet loss = 66%, RTA = 0.38 ms
[1463505903] SERVICE ALERT: FITNESSGRAM;Service: WMI;OK;SOFT;3;Winmgmt: Started
[1463505942] HOST ALERT: CEDARON;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1463505947] SERVICE ALERT: FITNESSGRAM;Service: Windows Time;OK;SOFT;3;W32Time: Started
[1463505948] SERVICE ALERT: CEDARON;C:\ Drive Space;OK;SOFT;2;c:\ - total: 249.66 Gb - used: 35.01 Gb (14%) - free 214.65 Gb (86%)
[1463505957] SERVICE ALERT: FITNESSGRAM;CPU Load;OK;SOFT;3;CPU Load 87% (15 min average)
[1463505967] SERVICE ALERT: CEDARON;VMWare Tools;OK;SOFT;2;VMTools: Started
[1463506079] HOST ALERT: tripwire;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1463506143] HOST ALERT: tripwire;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.39 ms
[1463506603] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463506723] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1463506843] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 88% (15 min average)
[1463507223] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;OK;HARD;3;Yellow Cartridge HP CE262A is at 69% - OK!
[1463507304] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;OK;HARD;3;Black Cartridge HP CE260X is at 73% - OK!
[1463507363] SERVICE ALERT: hallefirstcolor;Image Fuser Kit HP 110V-CE246A;OK;HARD;3;Image Fuser Kit HP 110V-CE246A, 220V-CE247A is at 100% - OK!
[1463507420] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1463507432] HOST ALERT: hallefirst3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1463507443] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463507477] SERVICE ALERT: hallefirst3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1463507481] HOST ALERT: hallefirst3;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.49 ms
[1463507533] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;TRAY 4 EMPTY PLAIN LETTER   Ready
[1463507548] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;OK;HARD;3;Magenta Cartridge HP CE263A is at 85% - OK!
[1463507553] SERVICE ALERT: hallefirstcolor;Toner Collection Unit HP CE265A;OK;HARD;3;
[1463507563] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463507580] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;OK;HARD;3;Image Transfer Kit HP CE249A is at 99% - OK!
[1463507586] SERVICE ALERT: hallefirstcolor;Device Status;OK;HARD;3;Ready
[1463507589] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1463507648] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 51% - OK!
[1463507653] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;OK;HARD;3;Cyan Cartridge HP CE261A is at 77% - OK!
[1463507656] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Ready
[1463507683] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463507777] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 44% - OK!
[1463507813] SERVICE ALERT: hallefirst1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1463507864] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 42% - OK!
[1463507995] SERVICE ALERT: hallefirst3;Device Status;OK;HARD;3;PERFORM PRINTER MAINTENANCE 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1463508793] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 91266 found since last run. oneLevelSearchOps OK. searchOps is too high! 91267 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463508914] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 29617 found since last run. oneLevelSearchOps OK. searchOps is too high! 29618 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463509033] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 28365 found since last run. oneLevelSearchOps OK. searchOps is too high! 28366 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463510833] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463513403] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1463513403] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1463513403] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1463513403] HOST DOWNTIME ALERT: ADAUTHDC1;STARTED; Host has entered a period of scheduled downtime
[1463513403] SERVICE DOWNTIME ALERT: NETIDADPROXY;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1463513403] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1463513403] HOST DOWNTIME ALERT: GOVGIS;STARTED; Host has entered a period of scheduled downtime
[1463513403] HOST DOWNTIME ALERT: epprogsql2;STARTED; Host has entered a period of scheduled downtime
[1463513426] SERVICE ALERT: epprogsql2;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1463513430] HOST ALERT: epprogsql2;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1463513478] SERVICE ALERT: epprogsql2;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1463513479] SERVICE ALERT: epprogsql2;CPU Load;OK;HARD;3;CPU Load 10% (15 min average)
[1463513479] SERVICE ALERT: epprogsql2;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1463513566] SERVICE ALERT: epprogsql2;Service: Windows Time;OK;HARD;3;W32Time: Started
[1463513653] SERVICE ALERT: epprogsql2;Service: WMI;OK;HARD;3;Winmgmt: Started
[1463513662] SERVICE ALERT: epprogsql2;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1463513683] SERVICE ALERT: epprogsql2;Service: RPC;OK;HARD;3;RpcSs: Started
[1463513683] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463513787] SERVICE ALERT: epprogsql2;C:\ Drive Space;OK;HARD;3;c:\ - total: 249.66 Gb - used: 139.99 Gb (56%) - free 109.66 Gb (44%)
[1463513807] SERVICE ALERT: epprogsql2;VMWare Tools;OK;HARD;3;VMTools: Started
[1463513826] SERVICE ALERT: epprogsql2;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1463513890] SERVICE ALERT: epprogsql2;Uptime;OK;HARD;3;System Uptime - 38 day(s) 20 hour(s) 14 minute(s)
[1463513890] SERVICE ALERT: epprogsql2;Memory Usage;OK;HARD;3;Memory usage: total:20479.66 MB - used: 8106.30 MB (40%) - free: 12373.37 MB (60%)
[1463513925] SERVICE ALERT: epprogsql2;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1463514175] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1463514235] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1463514283] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1463514283] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1463514294] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1463514554] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1463514575] HOST ALERT: owen101lab3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1463514614] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1463514644] HOST ALERT: owen101lab3;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1463514674] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1463514704] HOST ALERT: owen101lab3;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1463514794] HOST ALERT: owen101lab3;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1463514883] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 53% (15 min average)
[1463514884] HOST ALERT: owen101lab3;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1463514884] SERVICE ALERT: owen101lab3;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1463514914] HOST ALERT: owen101lab3;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1463514920] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1463514944] SERVICE ALERT: owen101lab3;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1463514950] HOST ALERT: owen101lab3;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1463514980] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1463515004] SERVICE ALERT: owen101lab3;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1463515010] HOST ALERT: owen101lab3;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1463515040] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1463515070] HOST ALERT: owen101lab3;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1463515160] HOST ALERT: owen101lab3;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1463515483] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463515603] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463515634] HOST ALERT: owen101lab3;UP;HARD;10;PING WARNING - Packet loss = 90%, RTA = 0.64 ms
[1463515635] SERVICE ALERT: owen101lab3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1463515723] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463515868] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 7% - OK!
[1463516198] SERVICE ALERT: owen101lab3;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1463516323] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 77% (15 min average)
[1463516923] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1463517043] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 91% (15 min average)
[1463517236] HOST ALERT: arthas;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1463517258] HOST ALERT: crowd1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1463517311] HOST ALERT: arthas;UP;SOFT;2;PING OK - Packet loss = 68%, RTA = 0.35 ms
[1463517322] HOST ALERT: crowd1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1463517643] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463517763] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 88% (15 min average)
[1463518030] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463518089] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463518109] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463518149] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463518169] SERVICE ALERT: hallefirstcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463518169] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463518224] SERVICE ALERT: hallefirstcolor;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463518229] SERVICE ALERT: hallefirstcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463518229] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463518284] SERVICE ALERT: hallefirstcolor;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463518289] SERVICE ALERT: hallefirstcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463518344] SERVICE ALERT: hallefirstcolor;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463518354] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463518359] SERVICE ALERT: hallefirstcolor;Toner Collection Unit HP CE265A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463518363] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1463518385] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463518391] SERVICE ALERT: hallefirstcolor;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463518414] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463518419] SERVICE ALERT: hallefirstcolor;Toner Collection Unit HP CE265A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463518445] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463518451] SERVICE ALERT: hallefirstcolor;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463518459] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463518474] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463518479] SERVICE ALERT: hallefirstcolor;Toner Collection Unit HP CE265A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463518483] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463518494] SERVICE ALERT: marshallpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1463518505] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463518511] SERVICE ALERT: hallefirstcolor;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463518519] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463518579] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1463518598] SERVICE ALERT: owen101lab3;Device Status;OK;HARD;3;Warming up
[1463518603] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463519198] SERVICE ALERT: owen101lab3;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1463519203] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 69% (15 min average)
[1463519214] SERVICE ALERT: hallefirst1;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1463519258] SERVICE ALERT: owen101lab3;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1463519274] SERVICE ALERT: hallefirst1;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1463519318] SERVICE ALERT: owen101lab3;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1463519318] SERVICE FLAPPING ALERT: owen101lab3;Device Status;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1463519325] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;SOFT;1;e:\ - total: 599.87 Gb - used: 482.23 Gb (80%) - free 117.64 Gb (20%)
[1463519334] SERVICE ALERT: hallefirst1;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1463519343] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;OK;HARD;3;Yellow Cartridge HP CE262A is at 68% - OK!
[1463519423] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;OK;HARD;3;Black Cartridge HP CE260X is at 73% - OK!
[1463519445] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;SOFT;2;e:\ - total: 599.87 Gb - used: 482.23 Gb (80%) - free 117.64 Gb (20%)
[1463519483] SERVICE ALERT: hallefirstcolor;Image Fuser Kit HP 110V-CE246A;OK;HARD;3;Image Fuser Kit HP 110V-CE246A, 220V-CE247A is at 100% - OK!
[1463519538] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1463519565] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;HARD;3;e:\ - total: 599.87 Gb - used: 482.23 Gb (80%) - free 117.64 Gb (20%)
[1463519668] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;OK;HARD;3;Magenta Cartridge HP CE263A is at 84% - OK!
[1463519674] SERVICE ALERT: hallefirstcolor;Toner Collection Unit HP CE265A;OK;HARD;3;
[1463519699] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;OK;HARD;3;Image Transfer Kit HP CE249A is at 99% - OK!
[1463519705] SERVICE ALERT: hallefirstcolor;Device Status;OK;HARD;3;Ready
[1463519774] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;OK;HARD;3;Cyan Cartridge HP CE261A is at 77% - OK!
[1463519803] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463519922] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463519934] SERVICE ALERT: hallefirst1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1463520042] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463520607] SERVICE ALERT: PAPERCUT;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1463520727] SERVICE ALERT: PAPERCUT;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1463520847] SERVICE ALERT: PAPERCUT;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463521447] SERVICE ALERT: PAPERCUT;CPU Load;OK;HARD;3;CPU Load 53% (15 min average)
[1463521841] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463522318] SERVICE ALERT: owen101lab3;Device Status;OK;HARD;3;41.3 UNEXPECTED SIZE IN TRAY 4 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   LOAD TRAY 4 PLAIN LETTER To use another 74 72 61 79 20 70 72 65 73 73 20 1E
[1463522441] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463524241] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 85% (15 min average)
[1463524384] SERVICE ALERT: owen101lab1;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1463524444] SERVICE ALERT: owen101lab1;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1463524504] SERVICE ALERT: owen101lab1;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1463524841] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1463524961] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463525081] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463525095] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1463525154] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1463525215] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1463525814] SERVICE ALERT: marshallpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1463525814] SERVICE FLAPPING ALERT: marshallpk;Paper Tray Status;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1463526281] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1463526881] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1463527001] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463527121] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463528921] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 76% (15 min average)
[1463529521] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463529641] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463529761] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463530361] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (15 min average)
[1463530961] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463531081] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463531201] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463531318] SERVICE FLAPPING ALERT: owen101lab3;Device Status;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1463531801] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463532401] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1463534801] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463535402] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463537201] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463537214] SERVICE FLAPPING ALERT: marshallpk;Paper Tray Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1463537801] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463540307] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1463540427] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 84% (15 min average)