[1462422600] HOST DOWNTIME ALERT: ADWU1;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: FTK;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: ADDC3;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: ADAUTH-RMAD;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: ADDC22;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: Echo-01;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: MBAM;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: ADDC33;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: IMAGES;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: ADDC4;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: ADDC44;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: NESSUS;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: OFFICEKMS;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: NETENG-PA;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: ADDC5;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: ADAUTHDC1-old;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: LANSWEEPER;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: ADDC55;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: OPIE;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: PNCAPP-2012;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: PNCDB-2012;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: ADAUTHDC1;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: ADCA1;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: PNCWEB-2012;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: ADWSUSCLIENTS;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: PROCARE;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: R25-DELTA2012;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: AD-RMAD;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: R25-ALPHA2012;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: ADCA2;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: CBORD;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: R25-ZETA2012;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: SP2010DB;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: SANDBOX-DB;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: MBAMDB;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: REFWIN;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: SP2010;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: CEDARON;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: COE-WEBDAV;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: SCCM-PRIM01;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: ADDC1;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: ADARC;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: TABLEAU;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: EMUSQLEXPRESS;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: DIRECTACCESS;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: TI22-COE;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: HDBACKUPS;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: TI22;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: DPS-SECURE02;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: ADDC2;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: EMU-VALUE01;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: TI22-PSYCLINIC;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: XDM;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: UHSPHARM-2012;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: WINKMS;STOPPED; Host has exited from a period of scheduled downtime
[1462422600] HOST DOWNTIME ALERT: EMUSCRIPT;STOPPED; Host has exited from a period of scheduled downtime
[1462423292] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462423892] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462430482] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1462431092] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462431692] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462435892] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 85% (15 min average)
[1462436492] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1462436612] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 29% (15 min average)
[1462438412] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462438532] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462438652] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1462438674] SERVICE ALERT: SANDBOX-DB;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462438794] SERVICE ALERT: SANDBOX-DB;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1462438914] SERVICE ALERT: SANDBOX-DB;CPU Load;OK;SOFT;3;CPU Load 92% (15 min average)
[1462439252] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1462439852] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462439972] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 80% (15 min average)
[1462440682] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1462440802] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 89% (15 min average)
[1462442372] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1462442492] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462442612] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1462443202] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462443322] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 91% (15 min average)
[1462445612] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 68% (15 min average)
[1462446212] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1462446332] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462446452] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462448252] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1462448852] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462448972] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462449092] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1462449602] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (37.0% change >= 20.0% threshold)
[1462449602] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1462449602] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1462449602] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1462449602] HOST DOWNTIME ALERT: R25-BETA2012;STARTED; Host has entered a period of scheduled downtime
[1462449602] SERVICE DOWNTIME ALERT: NETIDADPROXY;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1462449602] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1462449602] HOST DOWNTIME ALERT: R25-GAMMA2012;STARTED; Host has entered a period of scheduled downtime
[1462449602] HOST DOWNTIME ALERT: R25-EPSILON2012;STARTED; Host has entered a period of scheduled downtime
[1462452692] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 59% (15 min average)
[1462453522] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1462453642] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1462453762] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462454362] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1462455092] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1462455212] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462455332] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462455562] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462455682] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462455802] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462456402] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1462456402] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1462457732] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462458202] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462458332] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462458603] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STARTED; Service appears to have started flapping (28.6% change >= 20.0% threshold)
[1462458603] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (35.2% change >= 20.0% threshold)
[1462458603] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1462458603] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1462458603] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1462458603] HOST DOWNTIME ALERT: R25-BETA2012;STARTED; Host has entered a period of scheduled downtime
[1462458603] SERVICE DOWNTIME ALERT: NETIDADPROXY;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1462458603] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1462458603] HOST DOWNTIME ALERT: R25-GAMMA2012;STARTED; Host has entered a period of scheduled downtime
[1462458603] HOST DOWNTIME ALERT: R25-EPSILON2012;STARTED; Host has entered a period of scheduled downtime
[1462458802] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1462459402] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1462459507] SERVICE ALERT: ADDC11;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.7.201 and port 12489: Connection refused
[1462459507] SERVICE ALERT: ADDC11;Service: Distributed File System Replication;CRITICAL;SOFT;1;connect to address 164.76.7.201 and port 12489: Connection refused
[1462459508] SERVICE ALERT: ADDC11;Service: User Profile;CRITICAL;SOFT;1;connect to address 164.76.7.201 and port 12489: Connection refused
[1462459522] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462459532] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 83% (15 min average)
[1462459624] SERVICE ALERT: ADDC11;Memory Usage;OK;SOFT;2;Memory usage: total:20479.73 MB - used: 1293.06 MB (6%) - free: 19186.66 MB (94%)
[1462459624] SERVICE ALERT: ADDC11;Service: Distributed File System Replication;OK;SOFT;2;DFSR: Started
[1462459625] SERVICE ALERT: ADDC11;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1462459642] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462459654] SERVICE ALERT: ADDC11;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1462459774] SERVICE ALERT: ADDC11;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1462460132] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1462460252] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 90% (15 min average)
[1462460852] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462460972] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1462460973] SERVICE ALERT: prayharroldpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1462461032] SERVICE ALERT: prayharroldpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1462461082] SERVICE ALERT: ADDC5;GPO Replication;UNKNOWN;SOFT;1;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1462461092] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462461092] SERVICE ALERT: prayharroldpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1462461382] SERVICE ALERT: ADDC5;GPO Replication;UNKNOWN;SOFT;2;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1462461682] SERVICE ALERT: ADDC5;GPO Replication;UNKNOWN;HARD;3;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1462461952] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 36363 found since last run. oneLevelSearchOps OK. searchOps is too high! 36364 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1462462042] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 85% (15 min average)
[1462462072] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1462462144] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;WARNING;SOFT;1;0
[1462462265] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;WARNING;SOFT;2;0
[1462462383] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;WARNING;HARD;3;0
[1462462725] SERVICE ALERT: AX-QUAL;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462462727] SERVICE ALERT: AX-QUAL;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462462825] SERVICE ALERT: AX-QUAL;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1462462827] SERVICE ALERT: AX-QUAL;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1462463242] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1462463362] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 88% (15 min average)
[1462463492] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 65% (15 min average)
[1462463734] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1462463795] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1462463855] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1462463932] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462463992] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462464052] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462464562] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1462464682] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1462464692] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462464802] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1462464812] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462464932] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462465532] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462465672] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 125434 found since last run. oneLevelSearchOps OK. searchOps is too high! 125435 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1462465793] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 33189 found since last run. oneLevelSearchOps OK. searchOps is too high! 33190 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1462465912] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 34026 found since last run. oneLevelSearchOps OK. searchOps is too high! 34027 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1462466132] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1462466451] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Ready
[1462466459] SERVICE ALERT: ADDC1;Service: Dell Backup Agent;CRITICAL;SOFT;1;ErdAgent: Starting
[1462466492] SERVICE ALERT: prayharroldpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1462466579] SERVICE ALERT: ADDC1;Service: Dell Backup Agent;OK;SOFT;2;ErdAgent: Started
[1462466602] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1462467332] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (15 min average)
[1462467711] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1462467802] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462467922] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 81% (15 min average)
[1462468054] SERVICE ALERT: marshallpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1462468532] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462468652] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1462468816] SERVICE ALERT: ADDC5;GPO Replication;OK;HARD;3;OK! All GPOs are properly replicated.
[1462468983] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;OK;HARD;3;0
[1462469252] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462469372] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462469491] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1462470091] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 90% (15 min average)
[1462470454] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1462470513] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1462470574] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1462470922] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462471042] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462471109] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1462471162] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 89% (15 min average)
[1462471229] SERVICE ALERT: BOEAPPTEST;CPU Load;OK;SOFT;2;CPU Load 71% (15 min average)
[1462471291] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462471411] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462471531] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1462471829] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1462471949] SERVICE ALERT: BOEAPPTEST;CPU Load;OK;SOFT;2;CPU Load 44% (15 min average)
[1462472131] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462472374] SERVICE ALERT: marshallpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1462472374] SERVICE FLAPPING ALERT: marshallpk;Paper Tray Status;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1462472511] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9743 found since last run. oneLevelSearchOps OK. searchOps is too high! 9744 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1462472632] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 42529 found since last run. oneLevelSearchOps OK. searchOps is too high! 42530 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1462472751] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 31883 found since last run. oneLevelSearchOps OK. searchOps is too high! 31884 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1462473331] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462474551] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1462474551] SERVICE FLAPPING ALERT: idm1.emich.edu;LDAP Stats;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1462474681] SERVICE ALERT: idm1.emich.edu;Loopback Driver;WARNING;SOFT;1;0
[1462474801] SERVICE ALERT: idm1.emich.edu;Loopback Driver;OK;SOFT;2;0
[1462474985] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;WARNING;SOFT;1;0
[1462475103] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;OK;SOFT;2;0
[1462475131] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462475362] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1462475403] SERVICE ALERT: idm1.emich.edu;Loopback Driver;WARNING;SOFT;1;0
[1462475482] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 90% (15 min average)
[1462475521] SERVICE ALERT: idm1.emich.edu;Loopback Driver;OK;SOFT;2;0
[1462475704] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;WARNING;SOFT;1;0
[1462475731] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462475823] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;OK;SOFT;2;0
[1462475970] HOST ALERT: CASPER;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1462476051] HOST ALERT: JSS9;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.98)
[1462476055] SERVICE ALERT: JSS9;CPU Load;CRITICAL;SOFT;1;connect to address 164.76.102.98 and port 12489: No route to host
[1462476055] SERVICE ALERT: JSS9;Uptime;CRITICAL;SOFT;1;connect to address 164.76.102.98 and port 12489: No route to host
[1462476058] HOST ALERT: JSS9;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.98)
[1462476060] HOST ALERT: CASPER;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1462476072] SERVICE ALERT: CASPER;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462476097] SERVICE ALERT: JSS9;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.102.98 and port 12489: No route to host
[1462476100] HOST ALERT: JSS9;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.102.98)
[1462476102] HOST ALERT: CASPER;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1462476111] SERVICE ALERT: CASPER;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462476120] SERVICE ALERT: idm1.emich.edu;Loopback Driver;WARNING;SOFT;1;0
[1462476141] HOST ALERT: CASPER;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1462476155] HOST ALERT: JSS9;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1462476167] HOST ALERT: CASPER;UP;SOFT;5;PING OK - Packet loss = 58%, RTA = 0.40 ms
[1462476172] SERVICE ALERT: CASPER;Memory Usage;OK;SOFT;2;Memory usage: total:21759.74 MB - used: 1094.11 MB (5%) - free: 20665.63 MB (95%)
[1462476172] SERVICE ALERT: JSS9;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1462476173] SERVICE ALERT: JSS9;CPU Load;OK;SOFT;2;CPU Load 78% (15 min average)
[1462476211] SERVICE ALERT: CASPER;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1462476214] SERVICE ALERT: JSS9;Memory Usage;OK;SOFT;2;Memory usage: total:10943.64 MB - used: 2964.53 MB (27%) - free: 7979.12 MB (73%)
[1462476240] SERVICE ALERT: idm1.emich.edu;Loopback Driver;OK;SOFT;2;0
[1462476424] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;WARNING;SOFT;1;0
[1462476543] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;OK;SOFT;2;0
[1462476841] SERVICE ALERT: idm1.emich.edu;Loopback Driver;WARNING;SOFT;1;0
[1462476903] SERVICE ALERT: JSS9;VMWare Tools;UNKNOWN;SOFT;1;Failed to open service VMTools: 1060: The specified service does not exist as an installed service.
[1462476927] SERVICE ALERT: JSS9;Service: Firewall;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1462476960] SERVICE ALERT: idm1.emich.edu;Loopback Driver;WARNING;SOFT;2;0
[1462477023] SERVICE ALERT: JSS9;VMWare Tools;OK;SOFT;2;VMTools: Started
[1462477047] SERVICE ALERT: JSS9;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1462477056] SERVICE ALERT: CASPER;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1462477079] SERVICE ALERT: idm1.emich.edu;Loopback Driver;OK;SOFT;3;0
[1462477143] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;WARNING;SOFT;1;0
[1462477173] SERVICE ALERT: CASPER;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1462477263] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;OK;SOFT;2;0
[1462477282] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1462477679] SERVICE ALERT: idm1.emich.edu;Loopback Driver;WARNING;SOFT;1;0
[1462477800] SERVICE ALERT: idm1.emich.edu;Loopback Driver;WARNING;SOFT;2;0
[1462477864] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;WARNING;SOFT;1;0
[1462477919] SERVICE ALERT: idm1.emich.edu;Loopback Driver;OK;SOFT;3;0
[1462477983] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;OK;SOFT;2;0
[1462478131] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462478403] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (49.2% change >= 20.0% threshold)
[1462478403] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1462478403] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1462478403] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1462478403] HOST DOWNTIME ALERT: R25-BETA2012;STARTED; Host has entered a period of scheduled downtime
[1462478403] SERVICE DOWNTIME ALERT: NETIDADPROXY;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1462478403] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1462478403] HOST DOWNTIME ALERT: R25-GAMMA2012;STARTED; Host has entered a period of scheduled downtime
[1462478403] HOST DOWNTIME ALERT: R25-EPSILON2012;STARTED; Host has entered a period of scheduled downtime
[1462478583] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;WARNING;SOFT;1;0
[1462478703] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;OK;SOFT;2;0
[1462478731] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462479303] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;WARNING;SOFT;1;0
[1462479424] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;OK;SOFT;2;0
[1462480023] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;WARNING;SOFT;1;0
[1462480142] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;OK;SOFT;2;0
[1462480319] SERVICE ALERT: idm1.emich.edu;Loopback Driver;WARNING;SOFT;1;0
[1462480439] SERVICE ALERT: idm1.emich.edu;Loopback Driver;OK;SOFT;2;0
[1462481041] SERVICE ALERT: idm1.emich.edu;Loopback Driver;WARNING;SOFT;1;0
[1462481159] SERVICE ALERT: idm1.emich.edu;Loopback Driver;OK;SOFT;2;0
[1462481730] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (15 min average)
[1462482330] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462482450] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1462482570] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1462484370] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 77% (15 min average)
[1462484970] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462485090] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1462485210] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 92% (15 min average)
[1462485682] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1462485801] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1462485810] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1462485921] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 77% (15 min average)
[1462485930] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462486050] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462491450] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 84% (15 min average)
[1462492050] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462492170] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1462492290] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1462493490] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 88% (15 min average)
[1462494090] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462494210] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1462494330] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 71% (15 min average)
[1462494930] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1462495050] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462495170] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 94% (15 min average)
[1462495770] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462495890] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 67% (15 min average)
[1462496490] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462496610] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1462496730] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462498530] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 91% (15 min average)
[1462499130] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462499250] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462499370] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462501426] SERVICE ALERT: CBORD;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462501429] SERVICE ALERT: CBORD;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462501430] SERVICE ALERT: CBORD;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462501526] SERVICE ALERT: CBORD;Memory Usage;OK;SOFT;2;Memory usage: total:10943.58 MB - used: 3205.55 MB (29%) - free: 7738.03 MB (71%)
[1462501529] SERVICE ALERT: CBORD;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1462501530] SERVICE ALERT: CBORD;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1462501693] SERVICE ALERT: INTLTESTAPP;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462501793] SERVICE ALERT: INTLTESTAPP;C:\ Drive Space;OK;SOFT;2;c:\ - total: 63.90 Gb - used: 47.97 Gb (75%) - free 15.93 Gb (25%)
[1462501990] SERVICE ALERT: LANSWEEPER;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462502090] SERVICE ALERT: LANSWEEPER;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1462502720] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462502840] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1462503440] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462503560] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 91% (15 min average)
[1462504160] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462504280] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462504400] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462505600] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 91% (15 min average)
[1462505970] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462506200] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1462506320] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1462506440] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1462506570] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1462507040] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1462507040] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)