[1473825609] HOST ALERT: eateriespk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1473825699] HOST ALERT: eateriespk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1473825729] SERVICE ALERT: eateriespk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.59.141. Is the printer off?  Is the network down?
[1473825741] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;Open right door.
[1473825759] HOST ALERT: eateriespk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1473825789] SERVICE ALERT: eateriespk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.59.141. Is the printer off?  Is the network down?
[1473825819] HOST ALERT: eateriespk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1473825849] SERVICE ALERT: eateriespk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.59.141. Is the printer off?  Is the network down?
[1473825879] HOST ALERT: eateriespk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1473825895] SERVICE ALERT: eateriespk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.59.141. Is the printer off?  Is the network down?
[1473825925] HOST ALERT: eateriespk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1473825955] SERVICE ALERT: eateriespk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.59.141. Is the printer off?  Is the network down?
[1473825985] HOST ALERT: eateriespk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1473826015] SERVICE ALERT: eateriespk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.59.141. Is the printer off?  Is the network down?
[1473826341] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473826401] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473826461] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473826723] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1473826843] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1473826963] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1473827061] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;Open right door.
[1473827365] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473828262] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473828321] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473828374] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473828381] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473828504] SERVICE FLAPPING ALERT: AD-RMAD;E:\ Drive Space;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1473828565] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 2393 found since last run. oneLevelSearchOps OK. searchOps is too high! 2398 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1473828685] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473829278] HOST ALERT: emu-fndn-fe7;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1473829294] SERVICE ALERT: emu-fndn-fe7;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1473829303] SERVICE ALERT: emu-fndn-fe7;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1473829307] SERVICE ALERT: emu-fndn-fe7;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1473829324] HOST ALERT: emu-fndn-fe7;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1473829335] SERVICE ALERT: emu-fndn-fe7;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1473829338] SERVICE ALERT: emu-fndn-fe7;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1473829343] SERVICE ALERT: emu-fndn-fe7;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1473829343] SERVICE ALERT: emu-fndn-fe7;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1473829343] SERVICE ALERT: emu-fndn-fe7;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1473829365] HOST ALERT: emu-fndn-fe7;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1473829403] SERVICE ALERT: emu-fndn-fe7;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1473829405] SERVICE ALERT: emu-fndn-fe7;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1473829414] SERVICE ALERT: emu-fndn-fe7;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1473829423] SERVICE ALERT: emu-fndn-fe7;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1473829427] SERVICE ALERT: emu-fndn-fe7;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1473829433] HOST ALERT: emu-fndn-fe7;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1473829455] SERVICE ALERT: emu-fndn-fe7;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1473829457] SERVICE ALERT: emu-fndn-fe7;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1473829463] SERVICE ALERT: emu-fndn-fe7;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1473829463] SERVICE ALERT: emu-fndn-fe7;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1473829463] SERVICE ALERT: emu-fndn-fe7;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1473829467] SERVICE ALERT: emu-fndn-fe7;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1473829485] HOST ALERT: emu-fndn-fe7;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1473829520] SERVICE ALERT: emu-fndn-fe7;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1473829523] SERVICE ALERT: emu-fndn-fe7;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1473829523] SERVICE ALERT: emu-fndn-fe7;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1473829525] SERVICE ALERT: emu-fndn-fe7;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1473829534] SERVICE ALERT: emu-fndn-fe7;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1473829543] SERVICE ALERT: emu-fndn-fe7;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1473829547] SERVICE ALERT: emu-fndn-fe7;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1473829550] HOST ALERT: emu-fndn-fe7;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1473829574] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1896 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1473829575] SERVICE ALERT: emu-fndn-fe7;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1473829577] SERVICE ALERT: emu-fndn-fe7;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1473829582] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;Open right door.
[1473829583] SERVICE ALERT: emu-fndn-fe7;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1473829583] SERVICE ALERT: emu-fndn-fe7;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1473829583] SERVICE ALERT: emu-fndn-fe7;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1473829587] SERVICE ALERT: emu-fndn-fe7;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1473829605] HOST ALERT: emu-fndn-fe7;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1473829610] SERVICE ALERT: eateriespk;Device Status;OK;HARD;3;Genuine HP cartridge installed
[1473829614] HOST ALERT: eateriespk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.61 ms
[1473829640] SERVICE ALERT: emu-fndn-fe7;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1473829643] SERVICE ALERT: emu-fndn-fe7;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1473829643] SERVICE ALERT: emu-fndn-fe7;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1473829645] SERVICE ALERT: emu-fndn-fe7;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1473829670] HOST ALERT: emu-fndn-fe7;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1473829684] SERVICE ALERT: eateriespk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1473829694] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473829707] SERVICE ALERT: emu-fndn-fe7;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1473829737] HOST ALERT: emu-fndn-fe7;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1473829760] SERVICE ALERT: emu-fndn-fe7;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1473829763] SERVICE ALERT: emu-fndn-fe7;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1473829790] HOST ALERT: emu-fndn-fe7;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1473829901] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1473830021] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;2;Warning:  idm1  are more than 120 seconds behind!
[1473830067] HOST ALERT: emu-fndn-fe7;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1473830114] SERVICE ALERT: emu-fndn-fe7;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 4 minute(s)
[1473830123] SERVICE ALERT: emu-fndn-fe7;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1473830127] SERVICE ALERT: emu-fndn-fe7;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1473830141] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1473830155] SERVICE ALERT: emu-fndn-fe7;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1473830157] SERVICE ALERT: emu-fndn-fe7;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1473830163] SERVICE ALERT: emu-fndn-fe7;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1473830163] SERVICE ALERT: emu-fndn-fe7;Service: WMI;OK;HARD;3;Winmgmt: Started
[1473830163] SERVICE ALERT: emu-fndn-fe7;C:\ Drive Space;OK;HARD;3;c:\ - total: 698.53 Gb - used: 42.93 Gb (6%) - free 655.61 Gb (94%)
[1473830182] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473830223] SERVICE ALERT: emu-fndn-fe7;Service: RPC;OK;HARD;3;RpcSs: Started
[1473830225] SERVICE ALERT: emu-fndn-fe7;CPU Load;OK;HARD;3;CPU Load 7% (15 min average)
[1473830241] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473830287] SERVICE ALERT: emu-fndn-fe7;Service: Windows Time;OK;HARD;3;W32Time: Started
[1473830294] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 2268 found since last run. oneLevelSearchOps OK. searchOps is too high! 2296 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1473830301] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473830340] SERVICE ALERT: emu-fndn-fe7;Memory Usage;OK;HARD;3;Memory usage: total:18781.43 MB - used: 1928.45 MB (10%) - free: 16852.98 MB (90%)
[1473830343] SERVICE ALERT: emu-fndn-fe7;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1473830414] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473830563] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 88% (15 min average)
[1473830643] SERVICE ALERT: eateriespk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 23% - OK!
[1473830742] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1473831085] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 2291 found since last run. oneLevelSearchOps OK. searchOps is too high! 2299 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1473831205] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473831342] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1473831501] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;Open right door.
[1473832102] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473832161] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473832221] SERVICE ALERT: hallegroundcolor;Device Status;OK;SOFT;3;Open right door.
[1473832363] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1473832483] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 91% (15 min average)
[1473832821] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473832881] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473832941] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473833141] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1473833141] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1473833541] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;Open right door.
[1473833683] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1473833742] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1473833803] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 89% (15 min average)
[1473834141] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473834201] SERVICE ALERT: hallegroundcolor;Device Status;OK;SOFT;2;Open right door.
[1473834324] SERVICE ALERT: FITNESSGRAM;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1473834403] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1473834444] SERVICE ALERT: FITNESSGRAM;CPU Load;OK;SOFT;2;CPU Load 0% (15 min average)
[1473834523] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1473834643] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1473834802] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473834861] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473834921] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473835244] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1473836121] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;Open right door.
[1473836131] SERVICE ALERT: INTLDB;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1473836251] SERVICE ALERT: INTLDB;CPU Load;OK;SOFT;2;CPU Load 2% (15 min average)
[1473836721] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473836740] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1473836781] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473836841] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473837173] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1473837293] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1473837413] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1473837441] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;Open right door.
[1473837613] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1637 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1473837733] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473837940] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1473838540] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1473838641] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473838660] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1473838701] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473838761] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473839005] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1560 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1473839125] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473839443] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1473839961] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;Open right door.
[1473840043] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1473840163] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1473840283] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1473840561] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473840621] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473840681] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473840883] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 80% (15 min average)
[1473841281] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;Open right door.
[1473841933] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 535 found since last run.
[1473842053] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473842481] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473842541] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473842601] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473843283] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1473843403] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 19% (15 min average)
[1473843413] SERVICE ALERT: epprogsql2;CPU Load;OK;HARD;3;CPU Load 86% (15 min average)
[1473843801] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;Open right door.
[1473844003] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1473844123] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 84% (15 min average)
[1473844401] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473844461] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473844521] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473845323] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1473845403] SERVICE FLAPPING ALERT: hallegroundcolor;Device Status;STARTED; Service appears to have started flapping (64.1% change >= 20.0% threshold)
[1473845403] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1473845403] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1473845403] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1473845403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1473845403] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1473845403] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1473845403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1473845403] HOST DOWNTIME ALERT: NETIDADPROXY;STARTED; Host has entered a period of scheduled downtime
[1473845403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1473845403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1473845403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1473845403] HOST DOWNTIME ALERT: VUM;STARTED; Host has entered a period of scheduled downtime
[1473845403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1473845403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1473845403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1473845403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1473845403] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1473845403] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1473845443] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 86% (15 min average)
[1473845443] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1473845721] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;Open right door.
[1473846043] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1473846163] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 60% (15 min average)
[1473846253] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 527 found since last run.
[1473846321] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473846325] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 191 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 612 found since last run.
[1473846373] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473846381] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473846441] SERVICE ALERT: hallegroundcolor;Device Status;OK;SOFT;3;Open right door.
[1473846445] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473846973] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1709 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 505 found since last run.
[1473847041] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473847045] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 180 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 579 found since last run.
[1473847093] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473847101] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473847161] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473847165] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473847761] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;Open right door.
[1473847765] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 126 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 520 found since last run.
[1473847885] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473848362] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473848421] SERVICE ALERT: hallegroundcolor;Device Status;OK;SOFT;2;Open right door.
[1473848561] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1473848681] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1473848801] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1473849021] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473849081] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473849085] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 826 found since last run.
[1473849141] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473849205] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473849493] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 724 found since last run.
[1473849613] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473849741] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;Open right door.
[1473850803] SERVICE FLAPPING ALERT: hallegroundcolor;Device Status;STARTED; Service appears to have started flapping (60.1% change >= 20.0% threshold)
[1473850803] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1473850803] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1473850803] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1473850803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1473850803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1473850803] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1473850803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1473850803] HOST DOWNTIME ALERT: NETIDADPROXY;STARTED; Host has entered a period of scheduled downtime
[1473850803] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1473850803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1473850803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1473850803] HOST DOWNTIME ALERT: VUM;STARTED; Host has entered a period of scheduled downtime
[1473850803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1473850803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1473850803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1473850803] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1473850803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1473850803] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1473850813] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 547 found since last run.
[1473850933] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473850941] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473851002] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473851005] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 908 found since last run.
[1473851061] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473851125] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473851533] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1512 found since last run.
[1473851653] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473851661] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;Open right door.
[1473851725] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1543 found since last run.
[1473851845] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473852061] SERVICE ALERT: ldap.emich.edu;LDAP Stats;CRITICAL;SOFT;1;07:21:01  - (FAILURE!) LDAP Failed
[1473852121] SERVICE ALERT: ldap.emich.edu;LDAP Stats;OK;SOFT;2;07:22:01  - Destination LDAP online and available.
[1473852253] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1552 found since last run.
[1473852373] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473852401] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1473852445] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1552 found since last run.
[1473852565] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473852601] SERVICE ALERT: ldap.emich.edu;LDAP Stats;CRITICAL;SOFT;1;07:30:01  - (FAILURE!) LDAP Failed
[1473852661] SERVICE ALERT: ldap.emich.edu;LDAP Stats;OK;SOFT;2;07:31:01  - Destination LDAP online and available.
[1473852861] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473852920] SERVICE ALERT: hallegroundcolor;Device Status;OK;SOFT;2;Open right door.
[1473852973] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1553 found since last run.
[1473853001] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1473853001] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1473853093] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473853165] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1531 found since last run.
[1473853285] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473853299] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1473853520] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473853580] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473853603] SERVICE ALERT: prayharroldpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1473853621] SERVICE ALERT: ldap.emich.edu;LDAP Stats;CRITICAL;SOFT;1;07:47:01  - (FAILURE!) LDAP Failed
[1473853640] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473853681] SERVICE ALERT: ldap.emich.edu;LDAP Stats;OK;SOFT;2;07:48:01  - Destination LDAP online and available.
[1473853693] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1632 found since last run.
[1473853813] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473853885] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1714 found since last run.
[1473853899] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1473854005] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473854201] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1473854240] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;Open right door.
[1473854402] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1473854402] SERVICE FLAPPING ALERT: hallegroundcolor;Device Status;STARTED; Service appears to have started flapping (60.8% change >= 20.0% threshold)
[1473854402] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1473854402] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1473854402] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1473854402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1473854402] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1473854402] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1473854402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1473854402] HOST DOWNTIME ALERT: NETIDADPROXY;STARTED; Host has entered a period of scheduled downtime
[1473854402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1473854402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1473854402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1473854402] HOST DOWNTIME ALERT: VUM;STARTED; Host has entered a period of scheduled downtime
[1473854402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1473854402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1473854402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1473854402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1473854402] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1473854402] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1473854413] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1526 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 1701 found since last run.
[1473854533] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473854605] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1317 found since last run.
[1473854725] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473854801] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1473854921] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 65% (15 min average)
[1473855133] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5152 found since last run. oneLevelSearchOps OK. searchOps is too high! 5506 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 699 found since last run.
[1473855253] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473855325] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 3012 found since last run. oneLevelSearchOps OK. searchOps is too high! 3213 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 702 found since last run.
[1473855441] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473855445] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473855500] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473855560] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473855853] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 3615 found since last run. oneLevelSearchOps OK. searchOps is too high! 3892 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 642 found since last run.
[1473855973] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473856045] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 3193 found since last run. oneLevelSearchOps OK. searchOps is too high! 3354 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 727 found since last run.
[1473856121] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1473856160] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;Open right door.
[1473856165] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473856241] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 84% (15 min average)
[1473856367] HOST ALERT: halleground2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1473856369] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1473856399] HOST ALERT: halleground2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1473856429] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1473856459] HOST ALERT: halleground2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1473856486] SERVICE ALERT: halleground2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1473856489] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1473856516] HOST ALERT: halleground2;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1473856546] SERVICE ALERT: halleground2;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1473856569] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1473856573] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 2474 found since last run. oneLevelSearchOps OK. searchOps is too high! 2721 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 644 found since last run.
[1473856576] HOST ALERT: halleground2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1473856606] SERVICE ALERT: halleground2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1473856629] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1473856636] HOST ALERT: halleground2;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1473856689] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1473856693] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473856719] HOST ALERT: halleground2;DOWN;SOFT;7;(Host check timed out after 30.01 seconds)
[1473856760] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473856765] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5449 found since last run. oneLevelSearchOps OK. searchOps is too high! 5558 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 789 found since last run.
[1473856809] HOST ALERT: halleground2;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1473856820] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473856841] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1473856880] SERVICE ALERT: hallegroundcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473856885] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473856899] HOST ALERT: halleground2;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1473856961] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 89% (15 min average)
[1473856989] HOST ALERT: halleground2;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1473857112] HOST ALERT: halleground2;UP;HARD;10;PING OK - Packet loss = 79%, RTA = 0.73 ms
[1473857200] SERVICE ALERT: halleground2;Device Status;OK;HARD;3;Ready 164.76.106.52
[1473857283] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 77% - OK!
[1473857293] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 2249 found since last run. oneLevelSearchOps OK. searchOps is too high! 2709 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 630 found since last run.
[1473857413] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473857480] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;Remove all paper found.
[1473857485] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6490 found since last run. oneLevelSearchOps OK. searchOps is too high! 6647 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 769 found since last run.
[1473857605] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1884 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1473857683] SERVICE ALERT: halleground2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1473857725] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1830 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1473858014] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1957 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 715 found since last run.
[1473858133] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1473858161] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1473858281] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 67% (15 min average)
[1473858732] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7246 found since last run. oneLevelSearchOps OK. searchOps is too high! 7492 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 820 found since last run.
[1473858852] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 2763 found since last run. oneLevelSearchOps OK. searchOps is too high! 2804 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1473858972] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 2445 found since last run. oneLevelSearchOps OK. searchOps is too high! 2500 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1473859408] SERVICE ALERT: prayharroldpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473859467] SERVICE ALERT: prayharroldpk;Device Status;OK;SOFT;2;Open left door. Fuser is HOT.
[1473860081] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1473860201] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1473860321] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 79% (15 min average)
[1473860921] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1473861029] HOST ALERT: vandal;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1473861041] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1473861119] HOST ALERT: vandal;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1473861161] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1473861209] HOST ALERT: vandal;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1473861299] HOST ALERT: vandal;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1473861389] HOST ALERT: vandal;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1473861459] HOST ALERT: vandal;UP;SOFT;6;PING OK - Packet loss = 54%, RTA = 0.28 ms
[1473862361] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 60% (15 min average)
[1473864161] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1473864281] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1473864401] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1473865777] SERVICE ALERT: hallefirst3;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473865837] SERVICE ALERT: hallefirst3;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473865897] SERVICE ALERT: hallefirst3;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473866497] SERVICE ALERT: hallefirst3;Device Status;OK;HARD;3;TRAY 2 EMPTY PLAIN LETTER   Processing job from tray 4
[1473866801] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1473867401] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1473867521] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 91% (15 min average)
[1473868121] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1473868241] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1473868361] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1473868681] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1473868741] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1473868800] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1473868879] SERVICE FLAPPING ALERT: hallegroundcolor;Device Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1473869401] SERVICE ALERT: marshallpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1473870243] SERVICE ALERT: eateriespk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1473870303] SERVICE ALERT: eateriespk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1473870363] SERVICE ALERT: eateriespk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1473870696] SERVICE ALERT: hallefirst3;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473870756] SERVICE ALERT: hallefirst3;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473870757] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 80% (15 min average)
[1473870816] SERVICE ALERT: hallefirst3;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473871357] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1473871416] SERVICE ALERT: hallefirst3;Device Status;OK;HARD;3;TRAY 2 EMPTY PLAIN LETTER   Processing job from tray 4
[1473871416] SERVICE FLAPPING ALERT: hallefirst3;Device Status;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1473871477] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1473871597] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1473871783] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1473871843] SERVICE ALERT: owen101lab2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473871843] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1473871902] SERVICE ALERT: owen101lab2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1473871902] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1473871962] SERVICE ALERT: owen101lab2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1473872197] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1473872763] SERVICE ALERT: eateriespk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 23% - OK!
[1473872797] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1473873704] SERVICE ALERT: hallegroundcolor;Image Transfer Kit HP CE249A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1473873758] SERVICE ALERT: hallegroundcolor;Image Transfer Kit HP CE249A;OK;SOFT;2;Image Transfer Kit HP CE249A is at 58% - OK!
[1473873997] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1473874596] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1473874612] SERVICE ALERT: epprogsql2;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1473874732] SERVICE ALERT: epprogsql2;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1473874801] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1473874852] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1473874861] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1473874903] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1473874920] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1473876652] SERVICE ALERT: epprogsql2;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1473876762] SERVICE ALERT: owen101lab2;Device Status;OK;HARD;3;Clearing paper path
[1473877026] SERVICE ALERT: owen101lab3;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1473877086] SERVICE ALERT: owen101lab3;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! 1 of 3 Trays Still Full
[1473877252] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1473877372] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1473877417] HOST ALERT: studentcenterpk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.58 ms
[1473877433] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 6% - OK!
[1473877530] SERVICE ALERT: studentcenterpk;Device Status;OK;HARD;3;Processing job from tray 4
[1473877596] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1473877692] SERVICE ALERT: studentcenterpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1473877707] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 65% - OK!
[1473877852] SERVICE ALERT: epprogsql2;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1473877852] SERVICE FLAPPING ALERT: epprogsql2;CPU Load;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1473878196] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1473879396] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 85% (15 min average)
[1473879996] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1473880116] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 91% (15 min average)
[1473880852] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1473880954] SERVICE ALERT: hallefirst1;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1473881014] SERVICE ALERT: hallefirst1;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! All Trays Still Full
[1473881316] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1473881437] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1473881452] SERVICE ALERT: epprogsql2;CPU Load;OK;HARD;3;CPU Load 1% (15 min average)
[1473881521] SERVICE ALERT: marshallpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1473881556] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1473882815] SERVICE FLAPPING ALERT: hallefirst3;Device Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1473883080] SERVICE ALERT: owen101lab3;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1473883140] SERVICE ALERT: owen101lab3;Device Status;OK;SOFT;2;Clearing paper path
[1473883956] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 59% (15 min average)
[1473884556] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1473884676] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1473884796] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1473884979] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1473885009] HOST ALERT: hallefirst2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1473885039] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1473885068] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1473885069] HOST ALERT: hallefirst2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1473885093] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;SOFT;3;Black Print Cartridge HP C8543X is at 11% - OK!
[1473885097] HOST ALERT: hallefirst2;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.45 ms
[1473885122] SERVICE ALERT: hallefirst2;Device Status;OK;SOFT;2;Processing job from tray 4
[1473885396] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1473885693] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at WARNING level!
[1473885753] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at WARNING level!
[1473885813] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at WARNING level!
[1473885844] SERVICE ALERT: alexanderpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.142.94. Is the printer off?  Is the network down?
[1473885898] SERVICE ALERT: alexanderpk;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 42% - OK!
[1473885996] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1473886596] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1473888396] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1473888996] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1473892683] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1473892743] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1473892803] SERVICE ALERT: eateriespk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1473892852] SERVICE FLAPPING ALERT: epprogsql2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1473896792] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1473897392] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1473899613] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1473901335] SERVICE ALERT: SP2010;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 79.90 Gb - used: 76.07 Gb (95%) - free 3.83 Gb (5%)
[1473901700] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;SOFT;1;e:\ - total: 699.87 Gb - used: 569.39 Gb (81%) - free 130.49 Gb (19%)
[1473901728] SERVICE ALERT: hallesouthcommons;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1473901787] SERVICE ALERT: hallesouthcommons;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1473901820] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;SOFT;2;e:\ - total: 699.87 Gb - used: 569.39 Gb (81%) - free 130.49 Gb (19%)
[1473901848] SERVICE ALERT: hallesouthcommons;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1473901935] SERVICE ALERT: SP2010;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 79.90 Gb - used: 75.77 Gb (95%) - free 4.13 Gb (5%)
[1473901940] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;HARD;3;e:\ - total: 699.87 Gb - used: 569.39 Gb (81%) - free 130.49 Gb (19%)
[1473902108] SERVICE ALERT: TARGETX-2012;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1473902228] SERVICE ALERT: TARGETX-2012;CPU Load;OK;SOFT;2;CPU Load 87% (15 min average)
[1473902792] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1473903047] SERVICE ALERT: hallesouthcommons;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1473903392] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1473903740] SERVICE ALERT: AD-RMAD;E:\ Drive Space;OK;HARD;3;e:\ - total: 699.87 Gb - used: 547.73 Gb (78%) - free 152.14 Gb (22%)
[1473906600] HOST DOWNTIME ALERT: SCCM-PRIM01;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: ADWU1;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: CBORD;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: SP2010;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: TI22-COE;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: TI22-PSYCLINIC;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: COE-WEBDAV;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: UHSPHARM-2012;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: WINKMS;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: DIRECTACCESS;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: XDM;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: DPS-SECURE02;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: EMU-VALUE01;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: 4WINDSMANAGER;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: EMUSCRIPT;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: EMUSQLEXPRESS;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: AD-RMAD;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: Echo-01;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: FTK;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: ADAUTH-RMAD;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: HDBACKUPS;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: IMAGES;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: ADAUTHDC1;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: LANSWEEPER;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: CEDARON;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: MBAM;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: ADAUTHDC1-OLD;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: MBAMDB;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: SP2010DB;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: NESSUS;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: ADCA1;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: NETENG-PA;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: OFFICEKMS;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: ADCA2;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: OPIE;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: PNCAPP-2012;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: ADDC1;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: PNCDB-2012;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: PNCWEB-2012;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: ADDC2;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: PROCARE;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: R25-ALPHA2012;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: ADDC3;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: R25-DELTA2012;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: R25-INTERFACE;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: ADDC4;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: R25-WEBSERVICES;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: R25-WEBVIEWER;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: ADDC5;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: R25-ZETA2012;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: REFWIN;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: ADWSUSCLIENTS;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: TABLEAU;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: SANDBOX-DB;STARTED; Host has entered a period of scheduled downtime
[1473906600] HOST DOWNTIME ALERT: TI22;STARTED; Host has entered a period of scheduled downtime
[1473906992] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 77% (15 min average)
[1473907592] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1473907712] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1473907832] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1473908535] SERVICE ALERT: SP2010;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 79.90 Gb - used: 75.93 Gb (95%) - free 3.97 Gb (5%)
[1473908707] SERVICE ALERT: XDM;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1473908808] SERVICE ALERT: PROCARE;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1473908827] SERVICE ALERT: XDM;CPU Load;OK;SOFT;2;CPU Load 35% (15 min average)
[1473908863] SERVICE ALERT: XDM;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.2.79 and port 12489: Connection refused
[1473908872] SERVICE ALERT: XDM;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1473908896] SERVICE ALERT: ADCA2;Uptime;CRITICAL;SOFT;1;connect to address 164.76.7.151 and port 12489: Connection refused
[1473908932] SERVICE ALERT: PROCARE;CPU Load;OK;SOFT;2;CPU Load 23% (15 min average)
[1473908972] SERVICE ALERT: XDM;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1473908974] SERVICE ALERT: XDM;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1473909016] SERVICE ALERT: ADCA2;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 2 minute(s)
[1473909022] SERVICE ALERT: NESSUS;Uptime;CRITICAL;SOFT;1;connect to address 164.76.2.69 and port 12489: Connection refused
[1473909060] SERVICE ALERT: NETENG-PA;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1473909133] SERVICE ALERT: NESSUS;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 2 minute(s)
[1473909180] SERVICE ALERT: NETENG-PA;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1473909256] SERVICE ALERT: EMUSQLEXPRESS;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.7.169 and port 12489: Connection refused
[1473909258] SERVICE ALERT: ADDC4;Service: Dell Backup Agent;CRITICAL;SOFT;1;ErdAgent: Starting
[1473909258] SERVICE ALERT: ADDC4;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1473909295] SERVICE ALERT: TI22-COE;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.4.108 and port 12489: Connection refused
[1473909295] SERVICE ALERT: ADDC3;Service: Dell Backup Agent;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1473909296] SERVICE ALERT: ADDC3;Service: Windows Time;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1473909300] SERVICE ALERT: NETENG-PA;CPU Load;OK;SOFT;3;CPU Load 10% (15 min average)
[1473909305] SERVICE ALERT: ADDC2;Service: Distributed File System Replication;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1473909308] SERVICE ALERT: FTK;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1473909312] SERVICE ALERT: ADWU1;Windows Update Services Event 10042;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1473909313] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1473909315] SERVICE ALERT: FTK;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1473909319] SERVICE ALERT: ADDC3;VMWare Tools;CRITICAL;SOFT;1;connect to address 164.76.7.203 and port 12489: Connection refused
[1473909322] SERVICE ALERT: ADDC2;AD Replication;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1473909323] SERVICE ALERT: FTK;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1473909325] SERVICE ALERT: ADDC5;Service: DFS Namespace;CRITICAL;SOFT;1;connect to address 164.76.7.205 and port 12489: Connection refused
[1473909325] SERVICE ALERT: ADDC5;Service: User Profile;CRITICAL;SOFT;1;connect to address 164.76.7.205 and port 12489: Connection refused
[1473909326] SERVICE ALERT: ADDC3;Service: DNS Service;CRITICAL;SOFT;1;connect to address 164.76.7.203 and port 12489: Connection refused
[1473909327] SERVICE ALERT: ADDC5;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.7.205 and port 12489: Connection refused
[1473909328] SERVICE ALERT: FTK;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1473909329] SERVICE ALERT: ADDC5;Service: Intersite Messaging;CRITICAL;SOFT;1;connect to address 164.76.7.205 and port 12489: Connection refused
[1473909332] SERVICE ALERT: ADAUTHDC1;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1473909338] HOST ALERT: FTK;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1473909341] HOST ALERT: OFFICEKMS;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1473909344] SERVICE ALERT: FTK;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1473909345] SERVICE ALERT: FTK;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1473909345] SERVICE ALERT: ADAUTHDC1-OLD;Service: AD Domain Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1473909347] SERVICE ALERT: FTK;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1473909349] SERVICE ALERT: ADDC3;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1473909350] SERVICE ALERT: FTK;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1473909355] SERVICE ALERT: HDBACKUPS;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1473909360] HOST ALERT: HDBACKUPS;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.109)
[1473909365] SERVICE ALERT: HDBACKUPS;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1473909368] HOST ALERT: HDBACKUPS;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.4.109)
[1473909368] SERVICE ALERT: ADDC5;Service: DNS Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1473909370] SERVICE ALERT: FTK;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1473909373] SERVICE ALERT: FTK;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1473909374] HOST ALERT: FTK;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1473909376] SERVICE ALERT: EMUSQLEXPRESS;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1473909378] SERVICE ALERT: ADDC4;Service: Dell Backup Agent;CRITICAL;SOFT;2;ErdAgent: Starting
[1473909378] SERVICE ALERT: ADDC4;Service: Windows Time;CRITICAL;SOFT;2;W32Time: Stopped
[1473909392] SERVICE ALERT: PNCWEB-2012;Service: RPC;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1473909395] SERVICE ALERT: TI22-PSYCLINIC;Service: NetLogin;CRITICAL;SOFT;1;connect to address 164.76.4.203 and port 12489: Connection refused
[1473909401] SERVICE ALERT: CBORD;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.4.184 and port 12489: Connection refused
[1473909405] HOST ALERT: OFFICEKMS;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1473909406] SERVICE ALERT: PNCDB-2012;Service: Group Policy;CRITICAL;SOFT;1;gpsvc: Stopped
[1473909412] SERVICE ALERT: CBORD;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1473909415] SERVICE ALERT: TI22-COE;C:\ Drive Space;OK;SOFT;2;c:\ - total: 79.66 Gb - used: 22.25 Gb (28%) - free 57.41 Gb (72%)
[1473909415] SERVICE ALERT: ADDC3;Service: Dell Backup Agent;CRITICAL;SOFT;2;ErdAgent: Starting
[1473909416] SERVICE ALERT: ADDC3;Service: Windows Time;CRITICAL;SOFT;2;W32Time: Stopped
[1473909422] SERVICE ALERT: ADWU1;Windows Update Services Event 10042;OK;SOFT;2;OK: No entries found
[1473909425] SERVICE ALERT: ADDC2;Service: Distributed File System Replication;OK;SOFT;2;DFSR: Started
[1473909428] SERVICE ALERT: FTK;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1473909431] HOST ALERT: HDBACKUPS;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.4.109)
[1473909432] SERVICE ALERT: ADAUTHDC1;Service: Windows Time;CRITICAL;SOFT;2;W32Time: Stopped
[1473909433] SERVICE ALERT: ADWU1;Windows Update Services Event 10032;OK;SOFT;2;OK: No entries found
[1473909435] SERVICE ALERT: FTK;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1473909439] SERVICE ALERT: ADDC3;VMWare Tools;OK;SOFT;2;VMTools: Started
[1473909441] SERVICE ALERT: ADDC2;AD Replication;OK;SOFT;2;OK: replication is up and running
[1473909443] SERVICE ALERT: FTK;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1473909444] SERVICE ALERT: FTK;Uptime;CRITICAL;SOFT;2;connect to address 164.76.7.103 and port 12489: Connection refused
[1473909445] SERVICE ALERT: ADDC5;Service: DFS Namespace;OK;SOFT;2;Dfs: Started
[1473909445] SERVICE ALERT: ADDC5;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1473909445] SERVICE ALERT: FTK;Service: NetLogin;CRITICAL;SOFT;2;connect to address 164.76.7.103 and port 12489: Connection refused
[1473909445] SERVICE ALERT: ADAUTHDC1-OLD;Service: AD Domain Service;OK;SOFT;2;NTDS: Started
[1473909446] SERVICE ALERT: FTK;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.7.103 and port 12489: Connection refused
[1473909446] SERVICE ALERT: ADDC3;Service: DNS Service;OK;SOFT;2;DNS: Started
[1473909447] SERVICE ALERT: FTK;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1473909447] SERVICE ALERT: FTK;Service: Firewall;CRITICAL;SOFT;2;connect to address 164.76.7.103 and port 12489: Connection refused
[1473909447] SERVICE ALERT: ADDC5;Memory Usage;OK;SOFT;2;Memory usage: total:20479.73 MB - used: 1233.89 MB (6%) - free: 19245.84 MB (94%)
[1473909448] SERVICE ALERT: FTK;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1473909448] SERVICE ALERT: FTK;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1473909449] SERVICE ALERT: ADDC3;CPU Load;OK;SOFT;2;CPU Load 0% (15 min average)
[1473909449] SERVICE ALERT: ADDC5;Service: Intersite Messaging;OK;SOFT;2;IsmServ: Started
[1473909450] SERVICE ALERT: FTK;Service: WMI;CRITICAL;SOFT;2;connect to address 164.76.7.103 and port 12489: Connection refused
[1473909450] SERVICE ALERT: ADDC5;Service: Dell Backup Agent;CRITICAL;SOFT;1;ErdAgent: Starting
[1473909450] SERVICE ALERT: ADDC5;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1473909454] SERVICE ALERT: DPS-SECURE02;Service: Group Policy;CRITICAL;SOFT;1;gpsvc: Stopped
[1473909458] HOST ALERT: FTK;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1473909458] SERVICE ALERT: HDBACKUPS;Memory Usage;CRITICAL;SOFT;2;connect to address 164.76.4.109 and port 12489: No route to host
[1473909461] HOST ALERT: HDBACKUPS;DOWN;SOFT;4;CRITICAL - Host Unreachable (164.76.4.109)
[1473909468] SERVICE ALERT: HDBACKUPS;Service: Group Policy;CRITICAL;SOFT;2;connect to address 164.76.4.109 and port 12489: No route to host
[1473909468] SERVICE ALERT: ADDC5;Service: DNS Service;OK;SOFT;2;DNS: Started
[1473909477] SERVICE ALERT: FTK;Service: RPC;OK;SOFT;2;RpcSs: Started
[1473909479] SERVICE ALERT: FTK;CPU Load;OK;SOFT;2;CPU Load 4% (15 min average)
[1473909495] HOST ALERT: FTK;UP;SOFT;4;PING OK - Packet loss = 73%, RTA = 0.29 ms
[1473909498] SERVICE ALERT: ADDC4;Service: Dell Backup Agent;OK;SOFT;3;ErdAgent: Started
[1473909498] SERVICE ALERT: ADDC4;Service: Windows Time;OK;SOFT;3;W32Time: Started
[1473909512] SERVICE ALERT: CBORD;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1473909512] SERVICE ALERT: CBORD;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1473909512] SERVICE ALERT: TI22-PSYCLINIC;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1473909512] SERVICE ALERT: PNCWEB-2012;Service: RPC;OK;SOFT;2;RpcSs: Started
[1473909526] SERVICE ALERT: PNCDB-2012;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1473909528] SERVICE ALERT: FTK;C:\ Drive Space;OK;SOFT;3;c:\ - total: 408.28 Gb - used: 139.84 Gb (34%) - free 268.43 Gb (66%)
[1473909535] SERVICE ALERT: FTK;Service: User Profile;OK;SOFT;3;ProfSvc: Started
[1473909535] SERVICE ALERT: ADDC3;Service: Dell Backup Agent;OK;SOFT;3;ErdAgent: Started
[1473909536] SERVICE ALERT: ADDC3;Service: Windows Time;OK;SOFT;3;W32Time: Started
[1473909543] SERVICE ALERT: FTK;Service: DNS Client;OK;SOFT;3;Dnscache: Started
[1473909547] SERVICE ALERT: FTK;Memory Usage;OK;SOFT;2;Memory usage: total:65524.62 MB - used: 19235.52 MB (29%) - free: 46289.10 MB (71%)
[1473909548] SERVICE ALERT: FTK;NSClient++ Version;OK;SOFT;3;NSClient++ 0.4.4.15 2015-11-25
[1473909548] SERVICE ALERT: FTK;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1473909552] SERVICE ALERT: ADAUTHDC1;Service: Windows Time;OK;SOFT;3;W32Time: Started
[1473909564] SERVICE ALERT: FTK;Uptime;OK;SOFT;3;System Uptime - 0 day(s) 0 hour(s) 2 minute(s)
[1473909565] SERVICE ALERT: FTK;Service: NetLogin;OK;SOFT;3;Netlogon: Started
[1473909566] SERVICE ALERT: FTK;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1473909567] SERVICE ALERT: FTK;Service: Firewall;OK;SOFT;3;MpsSvc: Started
[1473909570] SERVICE ALERT: FTK;Service: WMI;OK;SOFT;3;Winmgmt: Started
[1473909570] SERVICE ALERT: ADDC5;Service: Dell Backup Agent;OK;SOFT;2;ErdAgent: Started
[1473909570] SERVICE ALERT: ADDC5;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1473909575] SERVICE ALERT: HDBACKUPS;Memory Usage;OK;SOFT;3;Memory usage: total:37532.83 MB - used: 1322.81 MB (4%) - free: 36210.02 MB (96%)
[1473909579] HOST ALERT: HDBACKUPS;UP;SOFT;5;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1473909585] SERVICE ALERT: HDBACKUPS;Service: Group Policy;OK;SOFT;3;gpsvc: Started
[1473909594] SERVICE ALERT: DPS-SECURE02;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1473909624] HOST ALERT: DPS-SECURE02;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1473909694] SERVICE ALERT: DPS-SECURE02;Service: Group Policy;CRITICAL;HARD;3;connect to address 164.76.2.46 and port 12489: Connection refused
[1473909714] HOST ALERT: DPS-SECURE02;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1473909726] SERVICE ALERT: DPS-SECURE02;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.2.46 and port 12489: Connection refused
[1473909735] SERVICE ALERT: SP2010;C:\ Drive Space;OK;HARD;3;c:\ - total: 79.90 Gb - used: 62.39 Gb (78%) - free 17.51 Gb (22%)
[1473909735] SERVICE FLAPPING ALERT: SP2010;C:\ Drive Space;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1473909756] HOST ALERT: DPS-SECURE02;UP;SOFT;3;PING WARNING - Packet loss = 86%, RTA = 0.28 ms
[1473909779] SERVICE ALERT: EMU-VALUE01;CPU Load;CRITICAL;SOFT;1;connect to address 164.76.4.85 and port 12489: No route to host
[1473909779] HOST ALERT: EMU-VALUE01;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.85)
[1473909836] SERVICE ALERT: EMU-VALUE01;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.4.85 and port 12489: No route to host
[1473909839] HOST ALERT: EMU-VALUE01;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.4.85)
[1473909842] SERVICE ALERT: MBAMDB;Service: Firewall;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1473909842] SERVICE ALERT: EMU-VALUE01;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.4.85 and port 12489: No route to host
[1473909843] SERVICE ALERT: ADAUTHDC1;AD Replication;WARNING;SOFT;1;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1473909846] SERVICE ALERT: DPS-SECURE02;Memory Usage;OK;SOFT;2;Memory usage: total:4790.25 MB - used: 894.26 MB (19%) - free: 3895.99 MB (81%)
[1473909857] SERVICE ALERT: MBAMDB;Uptime;CRITICAL;SOFT;1;connect to address 164.76.7.97 and port 12489: Connection refused
[1473909860] SERVICE ALERT: EMU-VALUE01;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.4.85 and port 12489: Connection refused
[1473909862] SERVICE ALERT: EMU-VALUE01;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.4.85 and port 12489: Connection refused
[1473909870] SERVICE ALERT: EMU-VALUE01;Service: User Profile;CRITICAL;SOFT;1;connect to address 164.76.4.85 and port 12489: Connection refused
[1473909874] SERVICE ALERT: EMU-VALUE01;Uptime;CRITICAL;SOFT;1;connect to address 164.76.4.85 and port 12489: Connection refused
[1473909887] SERVICE ALERT: EMU-VALUE01;CPU Load;CRITICAL;SOFT;2;connect to address 164.76.4.85 and port 12489: Connection refused
[1473909890] HOST ALERT: EMU-VALUE01;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1473909896] SERVICE ALERT: EMU-VALUE01;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.4.85 and port 12489: Connection refused
[1473909896] SERVICE ALERT: EMU-VALUE01;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.4.85 and port 12489: Connection refused
[1473909900] SERVICE ALERT: NETENG-PA;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1473909902] SERVICE ALERT: EMU-VALUE01;Service: NetLogin;CRITICAL;SOFT;1;connect to address 164.76.4.85 and port 12489: Connection refused
[1473909926] HOST ALERT: EMU-VALUE01;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1473909926] SERVICE ALERT: EMU-VALUE01;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.4.85 and port 12489: Connection refused
[1473909953] SERVICE ALERT: EMU-VALUE01;C:\ Drive Space;CRITICAL;SOFT;2;connect to address 164.76.4.85 and port 12489: Connection refused
[1473909956] HOST ALERT: EMU-VALUE01;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1473909962] SERVICE ALERT: MBAMDB;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1473909963] SERVICE ALERT: ADAUTHDC1;AD Replication;WARNING;SOFT;2;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1473909973] SERVICE ALERT: EMU-VALUE01;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1473909977] SERVICE ALERT: MBAMDB;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 2 minute(s)
[1473909979] SERVICE ALERT: EMU-VALUE01;Service: Windows Time;CRITICAL;SOFT;2;W32Time: Stopped
[1473909982] SERVICE ALERT: EMU-VALUE01;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1473909990] SERVICE ALERT: EMU-VALUE01;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1473909994] HOST ALERT: EMU-VALUE01;UP;SOFT;6;PING OK - Packet loss = 77%, RTA = 0.37 ms
[1473909994] SERVICE ALERT: EMU-VALUE01;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 4 minute(s)
[1473910007] SERVICE ALERT: EMU-VALUE01;CPU Load;OK;SOFT;3;CPU Load 13% (15 min average)
[1473910016] SERVICE ALERT: EMU-VALUE01;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1473910016] SERVICE ALERT: EMU-VALUE01;Memory Usage;OK;SOFT;2;Memory usage: total:9444.43 MB - used: 2125.39 MB (23%) - free: 7319.04 MB (77%)
[1473910020] SERVICE ALERT: NETENG-PA;CPU Load;OK;SOFT;2;CPU Load 0% (15 min average)
[1473910022] SERVICE ALERT: EMU-VALUE01;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1473910046] SERVICE ALERT: EMU-VALUE01;Service: RPC;OK;SOFT;2;RpcSs: Started
[1473910074] SERVICE ALERT: EMU-VALUE01;C:\ Drive Space;OK;SOFT;3;c:\ - total: 465.42 Gb - used: 15.52 Gb (3%) - free 449.90 Gb (97%)
[1473910083] SERVICE ALERT: ADAUTHDC1;AD Replication;WARNING;HARD;3;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1473910099] SERVICE ALERT: EMU-VALUE01;Service: Windows Time;OK;SOFT;3;W32Time: Started
[1473910294] SERVICE ALERT: DPS-SECURE02;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1473911287] SERVICE ALERT: ADAUTHDC1;AD Replication;OK;HARD;3;OK: replication is up and running