[1460521983] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 90% (5 min average)
[1460522583] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1460522583] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1460523296] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 94% (5 min average)
[1460523416] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 81% (5 min average)
[1460526416] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1460526536] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 77% (5 min average)
[1460527136] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 94% (5 min average)
[1460527211] SERVICE ALERT: DB01-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460527256] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 94% (5 min average)
[1460527331] SERVICE ALERT: DB01-2012;CPU Load;OK;SOFT;2;CPU Load 1% (5 min average)
[1460527376] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 94% (5 min average)
[1460527976] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 97% (5 min average)
[1460528576] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1460529176] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 71% (5 min average)
[1460529176] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1460529776] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1460529896] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 86% (5 min average)
[1460530331] SERVICE ALERT: NETENG-PA;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460530451] SERVICE ALERT: NETENG-PA;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1460530496] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1460530616] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 94% (5 min average)
[1460530736] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1460531583] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (5 min average)
[1460532183] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460532303] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1460532423] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460532535] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 73% (5 min average)
[1460533135] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 91% (5 min average)
[1460533255] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 93% (5 min average)
[1460533375] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 77% (5 min average)
[1460533773] SERVICE ALERT: halleground1;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460533833] SERVICE ALERT: halleground1;Printer Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460533894] SERVICE ALERT: halleground1;Printer Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460534248] HOST ALERT: halleground2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460534312] HOST ALERT: halleground2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.55 ms
[1460534493] SERVICE ALERT: halleground1;Printer Status;OK;HARD;3;ORDER CARTRIDGE LESS THAN 1500 PAGES 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1460535423] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 39% (5 min average)
[1460536623] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460536743] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 95% (5 min average)
[1460536863] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 55% (5 min average)
[1460537463] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460537583] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 85% (5 min average)
[1460540583] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460540703] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 97% (5 min average)
[1460540823] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 97% (5 min average)
[1460541423] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 56% (5 min average)
[1460541729] SERVICE ALERT: DB01-2012;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1460541849] SERVICE ALERT: DB01-2012;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1460541969] SERVICE ALERT: DB01-2012;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1460542023] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460542143] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1460542263] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1460542569] SERVICE ALERT: DB01-2012;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1460542975] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1460543169] SERVICE ALERT: DB01-2012;CPU Load;OK;HARD;3;CPU Load 4% (5 min average)
[1460545263] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 73% (5 min average)
[1460547063] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460547183] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 97% (5 min average)
[1460547303] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460548586] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1460548647] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1460548707] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1460549070] SERVICE ALERT: owen101labcolor;Magenta Cartridge HP CE263A;OK;HARD;3;
[1460549483] SERVICE ALERT: owen101labcolor;Cyan Cartridge HP CE261A;OK;HARD;3;
[1460549718] HOST ALERT: porterpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460549782] HOST ALERT: porterpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1460549856] SERVICE ALERT: porterpk;Black Cartridge HP CC364X;OK;HARD;3;Black Cartridge HP CC364X is at 27% - OK!
[1460549907] SERVICE ALERT: marshallpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1460549978] SERVICE ALERT: porterpk;Printer Status;OK;HARD;3;Ready
[1460549994] SERVICE ALERT: porterpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1460550303] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 47% (5 min average)
[1460551503] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460551623] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 97% (5 min average)
[1460551743] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 96% (5 min average)
[1460552574] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1460552694] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 74% (5 min average)
[1460553543] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1460553994] HOST ALERT: hallethirdsouth;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1460554143] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 78% (5 min average)
[1460555343] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460555463] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 89% (5 min average)
[1460556063] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1460556104] SERVICE ALERT: idm3.emich.edu;eDirectory Service;WARNING;SOFT;1;(No output on stdout) stderr: Could not resolve hostname idm3.emich.edu: Name or service not known
[1460556183] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1460556224] SERVICE ALERT: idm3.emich.edu;eDirectory Service;OK;SOFT;2;OK: ndsd status is good.
[1460556303] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1460556423] SERVICE ALERT: idm3.emich.edu;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460556431] HOST ALERT: idm3.emich.edu;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460556521] HOST ALERT: idm3.emich.edu;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460556543] SERVICE ALERT: idm3.emich.edu;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460556573] HOST ALERT: idm3.emich.edu;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460556605] SERVICE ALERT: idm3.emich.edu;LDAP Stats;UNKNOWN;SOFT;1;Difference is negative. LDAP recently rebooted.
[1460556635] HOST ALERT: idm3.emich.edu;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460556663] SERVICE ALERT: idm3.emich.edu;PING;CRITICAL;HARD;3;PING CRITICAL - Packet loss = 100%
[1460556693] HOST ALERT: idm3.emich.edu;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460556725] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1460556755] HOST ALERT: idm3.emich.edu;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1460556845] HOST ALERT: idm3.emich.edu;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1460556935] HOST ALERT: idm3.emich.edu;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1460557025] HOST ALERT: idm3.emich.edu;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1460557115] HOST ALERT: idm3.emich.edu;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1460557503] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 72% (5 min average)
[1460557594] HOST ALERT: hallethirdsouth;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460557684] HOST ALERT: hallethirdsouth;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460557730] HOST ALERT: hallethirdsouth;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460557733] SERVICE ALERT: hallefirst2;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460557793] SERVICE ALERT: hallefirst2;Printer Status;OK;SOFT;2;RETRIEVE JOB  NO STORED JOBS
[1460557793] HOST ALERT: idm2.emich.edu;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460557820] HOST ALERT: hallethirdsouth;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460557883] HOST ALERT: idm2.emich.edu;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460557890] HOST ALERT: hallethirdsouth;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460557973] HOST ALERT: idm2.emich.edu;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460557980] HOST ALERT: hallethirdsouth;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1460558063] HOST ALERT: idm2.emich.edu;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460558070] HOST ALERT: hallethirdsouth;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1460558083] SERVICE ALERT: idm2.emich.edu;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460558103] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460558113] HOST ALERT: idm2.emich.edu;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460558135] SERVICE ALERT: idm2.emich.edu;LDAP Stats;UNKNOWN;SOFT;1;Difference is negative. LDAP recently rebooted.
[1460558160] HOST ALERT: hallethirdsouth;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1460558164] HOST ALERT: idm2.emich.edu;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1460558203] SERVICE ALERT: idm2.emich.edu;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460558223] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460558233] HOST ALERT: idm2.emich.edu;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1460558250] HOST ALERT: hallethirdsouth;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1460558255] SERVICE ALERT: idm2.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1460558285] HOST ALERT: idm2.emich.edu;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1460558323] SERVICE ALERT: idm2.emich.edu;PING;CRITICAL;HARD;3;PING CRITICAL - Packet loss = 100%
[1460558330] HOST ALERT: hallethirdsouth;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1460558343] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 96% (5 min average)
[1460558353] HOST ALERT: idm2.emich.edu;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1460558443] HOST ALERT: idm2.emich.edu;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1460559125] SERVICE ALERT: idm3.emich.edu;LDAP Stats;UNKNOWN;SOFT;1;Difference is negative. LDAP recently rebooted.
[1460559245] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1460559294] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1460559410] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 43% - OK!
[1460559414] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 87% (5 min average)
[1460559743] SERVICE ALERT: halleground2;Printer Status;OK;HARD;3;Ready
[1460559813] SERVICE ALERT: halleground2;Maintenance Kit HP 110V-C9152A;OK;HARD;3;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 44% - OK!
[1460559963] SERVICE ALERT: halleground2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1460560279] SERVICE ALERT: idm1.emich.edu;LDAP Stats;CRITICAL;SOFT;1;CRITICAL: Failed to connect to LDAP!
[1460560402] SERVICE ALERT: idm1.emich.edu;LDAP Stats;CRITICAL;SOFT;2;CRITICAL: Failed to connect to LDAP!
[1460560432] HOST ALERT: idm1.emich.edu;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460560519] SERVICE ALERT: idm1.emich.edu;LDAP Stats;CRITICAL;HARD;3;CRITICAL: Failed to connect to LDAP!
[1460560522] HOST ALERT: idm1.emich.edu;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460560612] HOST ALERT: idm1.emich.edu;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460560702] HOST ALERT: idm1.emich.edu;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460560792] HOST ALERT: idm1.emich.edu;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460560882] HOST ALERT: idm1.emich.edu;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1460560943] SERVICE ALERT: idm1.emich.edu;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460560972] HOST ALERT: idm1.emich.edu;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1460561055] HOST ALERT: idm1.emich.edu;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1460561063] SERVICE ALERT: idm1.emich.edu;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460561093] HOST ALERT: idm1.emich.edu;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1460561119] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1460561149] HOST ALERT: idm1.emich.edu;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1460561183] SERVICE ALERT: idm1.emich.edu;PING;CRITICAL;HARD;3;PING CRITICAL - Packet loss = 100%
[1460561214] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1460561334] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 88% (5 min average)
[1460562963] SERVICE ALERT: hallesecondsouth;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460563023] SERVICE ALERT: hallesecondsouth;Printer Status;OK;SOFT;2;Clearing paper path
[1460563396] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (27.6% change >= 20.0% threshold)
[1460563396] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1460563396] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1460563396] HOST DOWNTIME ALERT: idm2.emich.edu;STARTED; Host has entered a period of scheduled downtime
[1460563396] HOST DOWNTIME ALERT: idm3.emich.edu;STARTED; Host has entered a period of scheduled downtime
[1460563396] HOST DOWNTIME ALERT: idm1.emich.edu;STARTED; Host has entered a period of scheduled downtime
[1460563396] HOST DOWNTIME ALERT: idmrbs.emich.edu;STARTED; Host has entered a period of scheduled downtime
[1460563396] HOST DOWNTIME ALERT: idmrbstest.emich.edu;STARTED; Host has entered a period of scheduled downtime
[1460563417] HOST ALERT: idmrbs.emich.edu;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1460563429] HOST ALERT: idm1.emich.edu;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1460563429] HOST ALERT: idm3.emich.edu;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.27 ms
[1460563457] HOST ALERT: idm2.emich.edu;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.92 ms
[1460563577] SERVICE ALERT: idm1.emich.edu;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.85 ms
[1460563583] SERVICE ALERT: idm1.emich.edu;NSDS Errors;CRITICAL;SOFT;1;NRPE: Command 'check_nsds_errors' not defined
[1460563630] SERVICE ALERT: idm2.emich.edu;NSDS Errors;CRITICAL;SOFT;1;NRPE: Command 'check_nsds_errors' not defined
[1460563677] SERVICE ALERT: idm3.emich.edu;NSDS Errors;CRITICAL;SOFT;1;NRPE: Command 'check_nsds_errors' not defined
[1460563703] SERVICE ALERT: idm1.emich.edu;NSDS Errors;CRITICAL;SOFT;2;NRPE: Command 'check_nsds_errors' not defined
[1460563717] SERVICE ALERT: idm2.emich.edu;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.42 ms
[1460563724] SERVICE ALERT: idmrbs.emich.edu;NSDS Errors;UNKNOWN;SOFT;1;NRPE: Unable to read output
[1460563738] SERVICE ALERT: idmrbs.emich.edu;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 1.56 ms
[1460563750] SERVICE ALERT: idm2.emich.edu;NSDS Errors;UNKNOWN;SOFT;2;NRPE: Unable to read output
[1460563771] SERVICE ALERT: idmrbstest.emich.edu;NSDS Errors;UNKNOWN;SOFT;1;NRPE: Unable to read output
[1460563797] SERVICE ALERT: idm3.emich.edu;NSDS Errors;UNKNOWN;SOFT;2;NRPE: Unable to read output
[1460563823] SERVICE ALERT: idm1.emich.edu;NSDS Errors;UNKNOWN;HARD;3;NRPE: Unable to read output
[1460563844] SERVICE ALERT: idmrbs.emich.edu;NSDS Errors;UNKNOWN;SOFT;2;NRPE: Unable to read output
[1460563857] SERVICE ALERT: idm3.emich.edu;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1460563870] SERVICE ALERT: idm2.emich.edu;NSDS Errors;UNKNOWN;HARD;3;NRPE: Unable to read output
[1460563891] SERVICE ALERT: crossroadspk;Black Cartridge HP CC364X;WARNING;SOFT;1;Black Cartridge HP CC364X is at 2% - WARNING!
[1460563891] SERVICE ALERT: idmrbstest.emich.edu;NSDS Errors;UNKNOWN;SOFT;2;NRPE: Unable to read output
[1460563917] SERVICE ALERT: idm3.emich.edu;NSDS Errors;UNKNOWN;HARD;3;NRPE: Unable to read output
[1460563950] SERVICE ALERT: crossroadspk;Black Cartridge HP CC364X;WARNING;SOFT;2;Black Cartridge HP CC364X is at 2% - WARNING!
[1460563964] SERVICE ALERT: idmrbs.emich.edu;NSDS Errors;UNKNOWN;HARD;3;NRPE: Unable to read output
[1460564010] SERVICE ALERT: crossroadspk;Black Cartridge HP CC364X;WARNING;HARD;3;Black Cartridge HP CC364X is at 2% - WARNING!
[1460564011] SERVICE ALERT: idmrbstest.emich.edu;NSDS Errors;UNKNOWN;HARD;3;NRPE: Unable to read output
[1460564060] SERVICE ALERT: idm1.emich.edu;NSDS Errors;OK;HARD;3;OK:
[1460564470] SERVICE ALERT: idm2.emich.edu;NSDS Errors;OK;HARD;3;OK:
[1460564517] SERVICE ALERT: idm3.emich.edu;NSDS Errors;OK;HARD;3;OK:
[1460564564] SERVICE ALERT: idmrbs.emich.edu;NSDS Errors;OK;HARD;3;OK:
[1460565243] SERVICE ALERT: owen101lab3;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1460565303] SERVICE ALERT: owen101lab3;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1460565363] SERVICE ALERT: owen101lab3;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1460565579] SERVICE ALERT: mckennypk;Black Cartridge HP CC364X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.152.26. Is the printer off?  Is the network down?
[1460565609] HOST ALERT: mckennypk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460565639] SERVICE ALERT: mckennypk;Black Cartridge HP CC364X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.152.26. Is the printer off?  Is the network down?
[1460565669] HOST ALERT: mckennypk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460565699] SERVICE ALERT: mckennypk;Black Cartridge HP CC364X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.152.26. Is the printer off?  Is the network down?
[1460565729] HOST ALERT: mckennypk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460565789] SERVICE ALERT: mckennypk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.152.26. Is the printer off?  Is the network down?
[1460565819] HOST ALERT: mckennypk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460565849] SERVICE ALERT: mckennypk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.152.26. Is the printer off?  Is the network down?
[1460565879] HOST ALERT: mckennypk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460565909] SERVICE ALERT: mckennypk;Printer Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.152.26. Is the printer off?  Is the network down?
[1460565909] SERVICE ALERT: mckennypk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.152.26. Is the printer off?  Is the network down?
[1460565939] HOST ALERT: mckennypk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1460565969] SERVICE ALERT: mckennypk;Printer Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.152.26. Is the printer off?  Is the network down?
[1460565999] HOST ALERT: mckennypk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1460566023] SERVICE ALERT: hallesecondsouth;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460566029] SERVICE ALERT: mckennypk;Printer Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.152.26. Is the printer off?  Is the network down?
[1460566059] HOST ALERT: mckennypk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1460566083] SERVICE ALERT: hallesecondsouth;Printer Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460566134] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1460566143] SERVICE ALERT: hallesecondsouth;Printer Status;OK;SOFT;3;MENUS  IPDS MENU  RETRIEVE JOB  INFORMATION
[1460566149] HOST ALERT: mckennypk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1460566239] HOST ALERT: mckennypk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1460566254] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 80% (5 min average)
[1460567162] SERVICE ALERT: owen101lab3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1460567514] SERVICE ALERT: idmrbs.emich.edu;Banner Driver;OK;HARD;3;0
[1460567561] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1460567561] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1460567561] HOST DOWNTIME ALERT: idmrbs.emich.edu;STARTED; Host has entered a period of scheduled downtime
[1460567561] HOST DOWNTIME ALERT: idm3.emich.edu;STARTED; Host has entered a period of scheduled downtime
[1460567561] HOST DOWNTIME ALERT: idmrbstest.emich.edu;STARTED; Host has entered a period of scheduled downtime
[1460567561] HOST DOWNTIME ALERT: idm2.emich.edu;STARTED; Host has entered a period of scheduled downtime
[1460567561] HOST DOWNTIME ALERT: idm1.emich.edu;STARTED; Host has entered a period of scheduled downtime
[1460567615] HOST ALERT: idmrbstest.emich.edu;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1460567900] SERVICE ALERT: idmrbstest.emich.edu;PING;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1460568252] SERVICE ALERT: ADWU1;Service: Firewall;CRITICAL;SOFT;1;(Return code of 13 is out of bounds)
[1460568252] SERVICE ALERT: ADRODC-LIV;Service: PXE Boot;CRITICAL;SOFT;1;(Return code of 13 is out of bounds)
[1460568252] SERVICE ALERT: SNAP;Uptime;CRITICAL;SOFT;1;(Return code of 13 is out of bounds)
[1460568252] SERVICE ALERT: INTLDB;Service: RPC;CRITICAL;SOFT;1;(Return code of 13 is out of bounds)
[1460568252] SERVICE ALERT: SQL1;Service: WMI;CRITICAL;SOFT;1;(Return code of 13 is out of bounds)
[1460568254] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1460568254] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1460568254] HOST DOWNTIME ALERT: idm2.emich.edu;STARTED; Host has entered a period of scheduled downtime
[1460568254] HOST DOWNTIME ALERT: idm3.emich.edu;STARTED; Host has entered a period of scheduled downtime
[1460568254] HOST DOWNTIME ALERT: idm1.emich.edu;STARTED; Host has entered a period of scheduled downtime
[1460568254] HOST DOWNTIME ALERT: idmrbs.emich.edu;STARTED; Host has entered a period of scheduled downtime
[1460568254] HOST DOWNTIME ALERT: idmrbstest.emich.edu;STARTED; Host has entered a period of scheduled downtime
[1460568363] SERVICE ALERT: owen101lab3;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1460568372] SERVICE ALERT: INTLDB;Service: RPC;OK;SOFT;2;RpcSs: Started
[1460568372] SERVICE ALERT: ADWU1;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1460568372] SERVICE ALERT: SNAP;Uptime;OK;SOFT;2;System Uptime - 4 day(s) 18 hour(s) 37 minute(s)
[1460568372] SERVICE ALERT: SQL1;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1460568372] SERVICE ALERT: ADRODC-LIV;Service: PXE Boot;OK;SOFT;2;WDSServer: Started
[1460568384] SERVICE ALERT: idmrbstest.emich.edu;Banner Driver;OK;HARD;3;0
[1460568422] SERVICE ALERT: owen101lab3;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1460568483] SERVICE ALERT: owen101lab3;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1460568628] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;CRITICAL;SOFT;1;CRITICAL! dxcmd.keystore not found.
[1460568675] SERVICE ALERT: idm1.emich.edu;Loopback Driver;CRITICAL;SOFT;1;CRITICAL! dxcmd.keystore not found.
[1460568748] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;CRITICAL;SOFT;2;CRITICAL! dxcmd.keystore not found.
[1460568795] SERVICE ALERT: idm1.emich.edu;Loopback Driver;CRITICAL;SOFT;2;CRITICAL! dxcmd.keystore not found.
[1460568868] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;CRITICAL;HARD;3;CRITICAL! dxcmd.keystore not found.
[1460568915] SERVICE ALERT: idm1.emich.edu;Loopback Driver;CRITICAL;HARD;3;CRITICAL! dxcmd.keystore not found.
[1460569141] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 90% (5 min average)
[1460569503] SERVICE ALERT: mckennypk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1460569507] HOST ALERT: mckennypk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.36 ms
[1460569623] SERVICE ALERT: mckennypk;Printer Status;OK;HARD;3;ORDER MAINTENANCE KIT 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1460569741] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460569793] SERVICE ALERT: hallefirst2;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460569853] SERVICE ALERT: hallefirst2;Printer Status;OK;SOFT;2;RETRIEVE JOB  NO STORED JOBS
[1460569893] SERVICE ALERT: mckennypk;Black Cartridge HP CC364X;OK;HARD;3;Black Cartridge HP CC364X is at 87% - OK!
[1460570098] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;OK;HARD;3;0
[1460570106] SERVICE ALERT: idm1.emich.edu;Loopback Driver;OK;HARD;3;0
[1460570128] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 60 seconds.
[1460570189] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;OK;SOFT;2;0
[1460570453] SERVICE ALERT: hallefirst2;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460570519] SERVICE ALERT: hallefirst2;Printer Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1460570574] SERVICE ALERT: hallefirst2;Printer Status;OK;SOFT;3;MENUS  RETRIEVE JOB  INFORMATION  PAPER HANDLING
[1460570582] SERVICE ALERT: idmrbstest.emich.edu;NSDS Errors;OK;HARD;3;OK:
[1460570610] SERVICE ALERT: crossroadspk;Black Cartridge HP CC364X;OK;HARD;3;Black Cartridge HP CC364X is at 100% - OK!
[1460570703] SERVICE ALERT: hallefirst2;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Maintenance Kit HP 110V-C9152A, 220V-C9153A'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1460570763] SERVICE ALERT: hallefirst2;Maintenance Kit HP 110V-C9152A;OK;SOFT;2;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 89% - OK!
[1460571224] SERVICE ALERT: prayharroldpk;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460571284] SERVICE ALERT: prayharroldpk;Printer Status;OK;SOFT;2;Clearing paper path
[1460571482] SERVICE ALERT: owen101lab3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1460571482] SERVICE FLAPPING ALERT: owen101lab3;Paper Tray Status;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1460572352] SERVICE ALERT: owen101lab1;Printer Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460572359] SERVICE ALERT: owen101lab1;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460572362] SERVICE ALERT: owen101lab1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460572368] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460572412] SERVICE ALERT: owen101lab1;Printer Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460572419] SERVICE ALERT: owen101lab1;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460572422] SERVICE ALERT: owen101lab1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460572428] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460572472] SERVICE ALERT: owen101lab1;Printer Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460572479] SERVICE ALERT: owen101lab1;Maintenance Kit HP 110V-C9152A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460572482] SERVICE ALERT: owen101lab1;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460572488] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460572962] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1460573016] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;OK;SOFT;2;Yellow Cartridge HP CE262A is at 31% - OK!
[1460573342] SERVICE ALERT: hallesecondsouth;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460573402] SERVICE ALERT: hallesecondsouth;Printer Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460573462] SERVICE ALERT: hallesecondsouth;Printer Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460573667] SERVICE ALERT: owen101lab1;Printer Status;OK;HARD;3;Ready
[1460573673] SERVICE ALERT: owen101lab1;Maintenance Kit HP 110V-C9152A;OK;HARD;3;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 55% - OK!
[1460573676] SERVICE ALERT: owen101lab1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1460573682] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 36% - OK!
[1460574054] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 94% (5 min average)
[1460574062] SERVICE ALERT: hallesecondsouth;Printer Status;OK;HARD;3;Pull then lift green tab.
[1460574174] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 74% (5 min average)
[1460575342] SERVICE ALERT: prayharroldpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1460575401] SERVICE ALERT: prayharroldpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1460575462] SERVICE ALERT: prayharroldpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1460576062] SERVICE ALERT: prayharroldpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1460576572] SERVICE ALERT: hallefirst2;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460576574] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1460576632] SERVICE ALERT: hallefirst2;Printer Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460576692] SERVICE ALERT: hallefirst2;Printer Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460576694] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 84% (5 min average)
[1460576905] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1460576966] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1460577026] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1460577147] HOST DOWNTIME ALERT: idm1.emich.edu;CANCELLED; Scheduled downtime for host has been cancelled.
[1460577151] HOST DOWNTIME ALERT: idm2.emich.edu;CANCELLED; Scheduled downtime for host has been cancelled.
[1460577156] HOST DOWNTIME ALERT: idm3.emich.edu;CANCELLED; Scheduled downtime for host has been cancelled.
[1460577164] HOST DOWNTIME ALERT: idmrbs.emich.edu;CANCELLED; Scheduled downtime for host has been cancelled.
[1460577168] HOST DOWNTIME ALERT: idmrbstest.emich.edu;CANCELLED; Scheduled downtime for host has been cancelled.
[1460577294] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 94% (5 min average)
[1460577414] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 87% (5 min average)
[1460577892] SERVICE ALERT: hallefirst2;Printer Status;OK;HARD;3;Remove any  jammed  paper.
[1460578141] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1460578291] SERVICE ALERT: halleground1;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460578351] SERVICE ALERT: halleground1;Printer Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460578411] SERVICE ALERT: halleground1;Printer Status;OK;SOFT;3;Open right  door.
[1460578741] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 82% (5 min average)
[1460578741] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1460578942] SERVICE ALERT: hallegroundcolor;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460579003] SERVICE ALERT: hallegroundcolor;Printer Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460579011] SERVICE ALERT: halleground1;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460579062] SERVICE ALERT: hallegroundcolor;Printer Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460579071] SERVICE ALERT: halleground1;Printer Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460579131] SERVICE ALERT: halleground1;Printer Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460579288] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1460579288] SERVICE ALERT: hallefirstcolor;Toner Collection Unit HP CE265A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1460579318] HOST ALERT: hallefirstcolor;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460579340] SERVICE ALERT: hallefirstcolor;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1460579341] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460579348] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1460579348] SERVICE ALERT: hallefirstcolor;Toner Collection Unit HP CE265A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1460579357] HOST ALERT: hallefirstcolor;UP;SOFT;2;PING OK - Packet loss = 72%, RTA = 0.39 ms
[1460579394] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! All Trays Still Full
[1460579402] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;OK;SOFT;3;
[1460579402] SERVICE ALERT: hallefirstcolor;Toner Collection Unit HP CE265A;OK;SOFT;3;
[1460579426] SERVICE ALERT: marshallpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1460579461] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1460579581] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 95% (5 min average)
[1460579661] SERVICE ALERT: hallegroundcolor;Printer Status;OK;HARD;3;MAGENTA CARTRIDGE LOW 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1460580482] SERVICE FLAPPING ALERT: owen101lab3;Paper Tray Status;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1460580781] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 65% (5 min average)
[1460581381] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460581501] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1460581531] SERVICE ALERT: halleground1;Printer Status;OK;HARD;3;ORDER CARTRIDGE LESS THAN 1600 PAGES 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Processing job
[1460581621] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 94% (5 min average)
[1460582221] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 82% (5 min average)
[1460582821] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460582941] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460582952] SERVICE ALERT: EMUFILES;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460583061] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460583072] SERVICE ALERT: EMUFILES;CPU Load;OK;SOFT;2;CPU Load 3% (5 min average)
[1460584260] SERVICE ALERT: halleground1;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1460584287] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1460584290] HOST ALERT: halleground1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460584314] SERVICE ALERT: halleground1;Maintenance Kit HP 110V-C9152A;OK;SOFT;2;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 28% - OK!
[1460584318] HOST ALERT: halleground1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.59 ms
[1460584341] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 100% - OK!
[1460584531] SERVICE ALERT: halleground1;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460584591] SERVICE ALERT: halleground1;Printer Status;OK;SOFT;2;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1460585528] SERVICE ALERT: owen101lab2;Printer Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1460585528] SERVICE ALERT: owen101lab2;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1460585558] HOST ALERT: owen101lab2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460585581] HOST ALERT: owen101lab1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460585582] SERVICE ALERT: owen101lab2;Printer Status;OK;SOFT;2;Genuine HP cartridge installed
[1460585585] SERVICE ALERT: owen101lab2;Maintenance Kit HP 110V-C9152A;OK;SOFT;2;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 58% - OK!
[1460585586] HOST ALERT: owen101lab2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1460585645] HOST ALERT: owen101lab1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.54 ms
[1460585672] SERVICE ALERT: owen101lab1;Printer Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460585677] SERVICE ALERT: owen101lab1;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460585681] SERVICE ALERT: owen101lab1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460585688] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460585718] HOST ALERT: owen101lab1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460585732] SERVICE ALERT: owen101lab1;Printer Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460585737] SERVICE ALERT: owen101lab1;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460585741] SERVICE ALERT: owen101lab1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460585748] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460585762] HOST ALERT: owen101lab1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460585792] SERVICE ALERT: owen101lab1;Printer Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460585796] HOST ALERT: owen101lab1;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.58 ms
[1460585796] HOST FLAPPING ALERT: owen101lab1;STARTED; Host appears to have started flapping (20.5% change > 20.0% threshold)
[1460585797] SERVICE ALERT: owen101lab1;Maintenance Kit HP 110V-C9152A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460585801] SERVICE ALERT: owen101lab1;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460585808] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1460585846] SERVICE ALERT: owen101lab3;Printer Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1460585876] HOST ALERT: owen101lab3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460585887] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1460585895] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1460585906] SERVICE ALERT: owen101lab3;Printer Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1460585915] HOST ALERT: owen101lab2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460585917] HOST ALERT: owen101lab3;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460585942] SERVICE ALERT: owen101lab3;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! 2 of 3 Trays Still Full
[1460585946] HOST ALERT: owen101lab3;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.63 ms
[1460585955] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1460585955] SERVICE ALERT: owen101lab2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1460585959] HOST ALERT: owen101lab2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1460585959] HOST FLAPPING ALERT: owen101lab2;STARTED; Host appears to have started flapping (22.9% change > 20.0% threshold)
[1460585966] SERVICE ALERT: owen101lab3;Printer Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1460586015] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1460586015] SERVICE ALERT: owen101lab2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1460586045] HOST ALERT: owen101lab2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460586075] SERVICE ALERT: owen101lab2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1460586085] HOST ALERT: owen101lab2;UP;SOFT;2;PING OK - Packet loss = 54%, RTA = 0.47 ms
[1460586188] SERVICE ALERT: owen101lab2;Printer Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1460586188] SERVICE ALERT: owen101lab2;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1460586248] SERVICE ALERT: owen101lab2;Printer Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1460586248] SERVICE ALERT: owen101lab2;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1460586300] HOST ALERT: owen101lab3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460586308] SERVICE ALERT: owen101lab2;Printer Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1460586308] SERVICE ALERT: owen101lab2;Maintenance Kit HP 110V-C9152A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1460586338] HOST ALERT: owen101lab2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460586364] HOST ALERT: owen101lab3;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.49 ms
[1460586364] HOST FLAPPING ALERT: owen101lab3;STARTED; Host appears to have started flapping (23.2% change > 20.0% threshold)
[1460586422] HOST ALERT: owen101lab1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460586428] HOST ALERT: owen101lab2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460586486] HOST ALERT: owen101lab1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.53 ms
[1460586518] HOST ALERT: owen101lab2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460586547] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1460586577] HOST ALERT: owen101lab3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460586607] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1460586608] HOST ALERT: owen101lab2;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460586626] HOST ALERT: owen101lab2;UP;SOFT;5;PING OK - Packet loss = 58%, RTA = 0.46 ms
[1460586627] HOST ALERT: owen101lab3;UP;SOFT;2;PING OK - Packet loss = 76%, RTA = 0.50 ms
[1460586667] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1460586903] SERVICE ALERT: owen101lab2;Printer Status;OK;HARD;3;Warming up
[1460586903] SERVICE ALERT: owen101lab2;Maintenance Kit HP 110V-C9152A;OK;HARD;3;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 58% - OK!
[1460587001] HOST ALERT: owen101lab3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460587009] HOST ALERT: owen101lab2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460587022] HOST ALERT: owen101lab1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460587091] HOST ALERT: owen101lab3;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460587099] HOST ALERT: owen101lab2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460587112] HOST ALERT: owen101lab1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460587181] HOST ALERT: owen101lab3;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460587189] HOST ALERT: owen101lab2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460587202] HOST ALERT: owen101lab1;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460587245] HOST ALERT: owen101lab2;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460587271] HOST ALERT: owen101lab3;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460587292] HOST ALERT: owen101lab1;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460587305] HOST ALERT: owen101lab2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460587352] HOST ALERT: owen101lab3;UP;SOFT;5;PING OK - Packet loss = 77%, RTA = 2.84 ms
[1460587357] HOST ALERT: owen101lab1;UP;SOFT;5;PING OK - Packet loss = 16%, RTA = 0.72 ms
[1460587369] HOST ALERT: owen101lab2;UP;SOFT;6;PING OK - Packet loss = 0%, RTA = 0.55 ms
[1460587508] SERVICE ALERT: owen101lab2;Printer Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1460587508] SERVICE ALERT: owen101lab2;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1460587568] SERVICE ALERT: owen101lab2;Printer Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1460587568] SERVICE ALERT: owen101lab2;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1460587622] HOST ALERT: owen101lab1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460587628] SERVICE ALERT: owen101lab2;Printer Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1460587628] SERVICE ALERT: owen101lab2;Maintenance Kit HP 110V-C9152A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1460587658] HOST ALERT: owen101lab2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460587686] HOST ALERT: owen101lab1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1460587722] HOST ALERT: owen101lab2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.45 ms
[1460587796] HOST ALERT: owen101lab3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460587886] HOST ALERT: owen101lab3;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460587976] HOST ALERT: owen101lab3;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460588040] HOST ALERT: owen101lab3;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.45 ms
[1460588461] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 75% (5 min average)
[1460589022] SERVICE ALERT: EMUSCRIPT;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.4.107 and port 12489: Connection refused
[1460589061] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460589075] HOST FLAPPING ALERT: owen101lab2;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1460589142] SERVICE ALERT: halleground2;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460589142] SERVICE ALERT: EMUSCRIPT;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1460589181] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1460589202] SERVICE ALERT: halleground2;Printer Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460589262] SERVICE ALERT: halleground2;Printer Status;OK;SOFT;3;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1460589301] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460589320] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1460589335] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1460589346] SERVICE ALERT: hallefirstcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1460589347] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1460589350] HOST ALERT: hallefirstcolor;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460589374] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Image Transfer Kit HP CE249A'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1460589378] HOST ALERT: hallefirstcolor;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1460589389] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;OK;SOFT;2;Black Cartridge HP CE260X is at 49% - OK!
[1460589401] SERVICE ALERT: hallefirstcolor;Image Fuser Kit HP 110V-CE246A;OK;SOFT;2;Image Fuser Kit HP 110V-CE246A, 220V-CE247A is at 100% - OK!
[1460589401] HOST ALERT: owen101lab3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460589402] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;OK;SOFT;2;Cyan Cartridge HP CE261A is at 79% - OK!
[1460589408] HOST FLAPPING ALERT: owen101lab1;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1460589409] HOST ALERT: owen101lab2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460589422] HOST ALERT: owen101lab1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460589435] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;OK;SOFT;3;Image Transfer Kit HP CE249A is at 11% - OK!
[1460589458] HOST ALERT: owen101lab2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460589491] HOST ALERT: owen101lab3;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460589512] HOST ALERT: owen101lab1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460589548] HOST ALERT: owen101lab2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460589581] HOST ALERT: owen101lab3;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460589602] HOST ALERT: owen101lab1;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460589638] HOST ALERT: owen101lab2;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460589671] HOST ALERT: owen101lab3;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460589692] HOST ALERT: owen101lab1;DOWN;SOFT;4;(Host check timed out after 30.01 seconds)
[1460589705] HOST ALERT: owen101lab2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460589761] HOST ALERT: owen101lab3;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460589782] HOST ALERT: owen101lab1;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460589795] HOST ALERT: owen101lab2;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1460589851] HOST ALERT: owen101lab3;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1460589872] HOST ALERT: owen101lab1;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1460589885] HOST ALERT: owen101lab2;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1460589941] HOST ALERT: owen101lab3;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1460589962] HOST ALERT: owen101lab1;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1460589975] HOST ALERT: owen101lab2;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1460590022] HOST ALERT: owen101lab1;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1460590029] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1460590031] HOST ALERT: owen101lab3;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1460590058] HOST ALERT: owen101lab2;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1460590089] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1460590112] HOST ALERT: owen101lab1;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1460590121] HOST ALERT: owen101lab3;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1460590121] HOST FLAPPING ALERT: owen101lab3;STOPPED; Host appears to have stopped flapping (4.9% change < 5.0% threshold)
[1460590148] HOST ALERT: owen101lab2;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1460590148] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1460590196] HOST ALERT: owen101lab3;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1460590202] HOST ALERT: owen101lab1;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1460590922] HOST ALERT: ADDC2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460590934] SERVICE ALERT: ADDC2;Service: Distributed File System Replication;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460590964] HOST ALERT: ADDC2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460590969] SERVICE ALERT: ADDC2;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460590992] SERVICE ALERT: ADDC2;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460590999] HOST ALERT: ADDC2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460591052] SERVICE ALERT: ADDC2;Service: Intersite Messaging;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460591052] SERVICE ALERT: ADDC2;AD Replication;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 30 seconds.
[1460591053] SERVICE ALERT: ADDC2;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460591054] SERVICE ALERT: ADDC2;Service: Distributed File System Replication;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460591082] HOST ALERT: ADDC2;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460591089] SERVICE ALERT: ADDC2;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460591098] SERVICE ALERT: ADDC2;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460591100] SERVICE ALERT: ADDC2;Service: Key Distribution Center;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460591112] SERVICE ALERT: ADDC2;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460591119] HOST ALERT: ADDC2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460591147] SERVICE ALERT: ADDC2;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460591152] SERVICE ALERT: ADDC2;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460591172] SERVICE ALERT: ADDC2;Service: Intersite Messaging;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460591172] SERVICE ALERT: ADDC2;AD Replication;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 30 seconds.
[1460591173] SERVICE ALERT: ADDC2;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460591174] SERVICE ALERT: ADDC2;Service: Distributed File System Replication;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460591177] HOST ALERT: ADDC2;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1460591194] SERVICE ALERT: ADDC2;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460591194] SERVICE ALERT: ADDC3;AD Replication;WARNING;SOFT;1;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1460591196] SERVICE ALERT: ADDC2;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460591209] SERVICE ALERT: ADDC2;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460591218] SERVICE ALERT: ADDC2;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460591220] SERVICE ALERT: ADDC2;Service: Key Distribution Center;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460591224] HOST ALERT: ADDC2;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1460591232] SERVICE ALERT: ADDC2;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460591242] SERVICE ALERT: ADDC2;Service: AD Domain Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460591242] SERVICE ALERT: ADDC2;Service: Remote Management;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460591262] HOST ALERT: ADDC2;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1460591267] SERVICE ALERT: ADDC2;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460591272] SERVICE ALERT: ADDC2;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460591292] SERVICE ALERT: ADDC2;Service: Intersite Messaging;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460591292] SERVICE ALERT: ADDC2;AD Replication;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 30 seconds.
[1460591293] SERVICE ALERT: ADDC2;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460591293] SERVICE ALERT: ADDC2;Service: AD Web Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460591293] SERVICE ALERT: ADDC2;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460591297] HOST ALERT: ADDC2;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1460591314] SERVICE ALERT: ADDC2;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460591314] SERVICE ALERT: ADDC3;AD Replication;WARNING;SOFT;2;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1460591316] SERVICE ALERT: ADDC2;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460591335] SERVICE ALERT: ADDC2;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460591335] SERVICE ALERT: ADDC2;Service: DFS Namespace;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460591338] SERVICE ALERT: ADDC2;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460591340] SERVICE ALERT: ADDC2;Service: Key Distribution Center;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460591344] HOST ALERT: ADDC2;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1460591362] SERVICE ALERT: ADDC2;Service: AD Domain Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460591362] SERVICE ALERT: ADDC2;Service: Remote Management;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460591374] SERVICE ALERT: ADDC2;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460591374] SERVICE ALERT: ADDC2;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460591387] SERVICE ALERT: ADDC2;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460591392] SERVICE ALERT: ADDC2;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460591413] SERVICE ALERT: ADDC2;Service: AD Web Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460591413] SERVICE ALERT: ADDC2;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460591423] SERVICE ALERT: ADDC4;AD Replication;WARNING;SOFT;1;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1460591426] SERVICE ALERT: ADDC2;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460591428] SERVICE ALERT: ADDC2;Service: DNS Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460591434] SERVICE ALERT: ADDC2;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460591434] SERVICE ALERT: ADDC3;AD Replication;WARNING;HARD;3;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1460591436] SERVICE ALERT: ADDC2;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460591455] SERVICE ALERT: ADDC2;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460591455] SERVICE ALERT: ADDC2;Service: DFS Namespace;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460591473] SERVICE ALERT: ADDC2;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460591473] SERVICE ALERT: ADDC2;Service: Dell Backup Agent;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460591482] SERVICE ALERT: ADDC2;Service: AD Domain Service;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460591482] SERVICE ALERT: ADDC2;Service: Remote Management;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460591494] SERVICE ALERT: ADDC2;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460591494] SERVICE ALERT: ADDC2;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460591533] SERVICE ALERT: ADDC2;Service: AD Web Service;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460591533] SERVICE ALERT: ADDC2;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460591542] SERVICE ALERT: ADDC4;AD Replication;WARNING;SOFT;2;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1460591546] SERVICE ALERT: ADDC2;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460591548] SERVICE ALERT: ADDC2;Service: DNS Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460591575] SERVICE ALERT: ADDC2;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460591575] SERVICE ALERT: ADDC2;Service: DFS Namespace;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460591593] SERVICE ALERT: ADDC2;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460591593] SERVICE ALERT: ADDC2;Service: Dell Backup Agent;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460591614] SERVICE ALERT: ADDC2;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460591614] SERVICE ALERT: ADDC2;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460591663] SERVICE ALERT: ADDC4;AD Replication;WARNING;HARD;3;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1460591666] SERVICE ALERT: ADDC2;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460591668] SERVICE ALERT: ADDC2;Service: DNS Service;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460591691] SERVICE ALERT: hallefirst2;Printer Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1460591713] SERVICE ALERT: ADDC2;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460591713] SERVICE ALERT: ADDC2;Service: Dell Backup Agent;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460591751] SERVICE ALERT: hallefirst2;Printer Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1460591811] SERVICE ALERT: hallefirst2;Printer Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1460592362] SERVICE ALERT: hallefirst2;Maintenance Kit HP 110V-C9152A;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Maintenance Kit HP 110V-C9152A, 220V-C9153A'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1460592411] SERVICE ALERT: hallefirst2;Printer Status;OK;HARD;3;Clearing paper path
[1460592422] SERVICE ALERT: hallefirst2;Maintenance Kit HP 110V-C9152A;OK;SOFT;2;Maintenance Kit HP 110V-C9152A, 220V-C9153A is at 89% - OK!
[1460594101] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 93% (5 min average)
[1460594701] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1460595772] SERVICE ALERT: TARGETX-2012;CPU Load;WARNING;SOFT;1;CPU Load 94% (5 min average)
[1460595892] SERVICE ALERT: TARGETX-2012;CPU Load;CRITICAL;SOFT;2;CPU Load 95% (5 min average)
[1460596012] SERVICE ALERT: TARGETX-2012;CPU Load;OK;SOFT;3;CPU Load 73% (5 min average)
[1460596501] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 88% (5 min average)
[1460597101] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460597221] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1460597341] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1460597368] SERVICE ALERT: hallefirst1;Printer Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1460597423] SERVICE ALERT: hallefirst1;Printer Status;OK;SOFT;2;RETRIEVE JOB  NO STORED JOBS
[1460597941] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 94% (5 min average)
[1460598540] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 96% (5 min average)
[1460600844] SERVICE ALERT: ADRODC-DET;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460600867] SERVICE ALERT: ADRODC-DET;Service: DNS Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460600870] SERVICE ALERT: ADRODC-DET;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460600874] HOST ALERT: ADRODC-DET;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460600892] SERVICE ALERT: ADRODC-DET;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460600922] HOST ALERT: ADRODC-DET;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460600950] SERVICE ALERT: VC;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460600964] SERVICE ALERT: ADRODC-DET;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460600987] SERVICE ALERT: ADRODC-DET;Service: DNS Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460600990] SERVICE ALERT: ADRODC-DET;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460600994] HOST ALERT: ADRODC-DET;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460601000] HOST DOWNTIME ALERT: PNCDB-2012;STARTED; Host has entered a period of scheduled downtime
[1460601000] HOST DOWNTIME ALERT: AD-RMAD;STARTED; Host has entered a period of scheduled downtime
[1460601000] HOST DOWNTIME ALERT: DPS-SECURE02;STARTED; Host has entered a period of scheduled downtime
[1460601000] HOST DOWNTIME ALERT: PNCWEB-2012;STARTED; Host has entered a period of scheduled downtime
[1460601000] HOST DOWNTIME ALERT: PROCARE;STARTED; Host has entered a period of scheduled downtime
[1460601000] HOST DOWNTIME ALERT: EMU-VALUE01;STARTED; Host has entered a period of scheduled downtime
[1460601000] HOST DOWNTIME ALERT: R25-ALPHA2012;STARTED; Host has entered a period of scheduled downtime
[1460601000] HOST DOWNTIME ALERT: R25-DELTA2012;STARTED; Host has entered a period of scheduled downtime
[1460601000] HOST DOWNTIME ALERT: ADAUTH-RMAD;STARTED; Host has entered a period of scheduled downtime
[1460601000] HOST DOWNTIME ALERT: EMUSCRIPT;STARTED; Host has entered a period of scheduled downtime
[1460601000] HOST DOWNTIME ALERT: R25-ZETA2012;STARTED; Host has entered a period of scheduled downtime
[1460601000] HOST DOWNTIME ALERT: REFWIN;STARTED; Host has entered a period of scheduled downtime
[1460601000] HOST DOWNTIME ALERT: EMUSQLEXPRESS;STARTED; Host has entered a period of scheduled downtime
[1460601000] HOST DOWNTIME ALERT: SANDBOX-DB;STARTED; Host has entered a period of scheduled downtime
[1460601000] HOST DOWNTIME ALERT: SP2010;STARTED; Host has entered a period of scheduled downtime
[1460601000] HOST DOWNTIME ALERT: ADWSUSCLIENTS;STARTED; Host has entered a period of scheduled downtime
[1460601000] HOST DOWNTIME ALERT: Echo-01;STARTED; Host has entered a period of scheduled downtime
[1460601000] HOST DOWNTIME ALERT: SP2010DB;STARTED; Host has entered a period of scheduled downtime
[1460601000] HOST DOWNTIME ALERT: TABLEAU;STARTED; Host has entered a period of scheduled downtime
[1460601000] HOST DOWNTIME ALERT: FTK;STARTED; Host has entered a period of scheduled downtime
[1460601000] HOST DOWNTIME ALERT: TI22;STARTED; Host has entered a period of scheduled downtime
[1460601000] HOST DOWNTIME ALERT: TI22-COE;STARTED; Host has entered a period of scheduled downtime
[1460601000] HOST DOWNTIME ALERT: ADWU1;STARTED; Host has entered a period of scheduled downtime
[1460601000] HOST DOWNTIME ALERT: HDBACKUPS;STARTED; Host has entered a period of scheduled downtime
[1460601000] HOST DOWNTIME ALERT: TI22-PSYCLINIC;STARTED; Host has entered a period of scheduled downtime
[1460601000] HOST DOWNTIME ALERT: UHSPHARM-2012;STARTED; Host has entered a period of scheduled downtime
[1460601000] HOST DOWNTIME ALERT: IMAGES;STARTED; Host has entered a period of scheduled downtime
[1460601000] HOST DOWNTIME ALERT: WINKMS;STARTED; Host has entered a period of scheduled downtime
[1460601000] HOST DOWNTIME ALERT: XDM;STARTED; Host has entered a period of scheduled downtime
[1460601000] HOST DOWNTIME ALERT: CBORD;STARTED; Host has entered a period of scheduled downtime
[1460601000] HOST DOWNTIME ALERT: LANSWEEPER;STARTED; Host has entered a period of scheduled downtime
[1460601000] HOST DOWNTIME ALERT: MBAM;STARTED; Host has entered a period of scheduled downtime
[1460601000] HOST DOWNTIME ALERT: MBAMDB;STARTED; Host has entered a period of scheduled downtime
[1460601000] HOST DOWNTIME ALERT: CEDARON;STARTED; Host has entered a period of scheduled downtime
[1460601000] HOST DOWNTIME ALERT: NESSUS;STARTED; Host has entered a period of scheduled downtime
[1460601000] HOST DOWNTIME ALERT: NETENG-PA;STARTED; Host has entered a period of scheduled downtime
[1460601000] HOST DOWNTIME ALERT: COE-WEBDAV;STARTED; Host has entered a period of scheduled downtime
[1460601000] HOST DOWNTIME ALERT: OFFICEKMS;STARTED; Host has entered a period of scheduled downtime
[1460601000] HOST DOWNTIME ALERT: OPIE;STARTED; Host has entered a period of scheduled downtime
[1460601000] HOST DOWNTIME ALERT: DIRECTACCESS;STARTED; Host has entered a period of scheduled downtime
[1460601000] HOST DOWNTIME ALERT: PNCAPP-2012;STARTED; Host has entered a period of scheduled downtime
[1460601012] SERVICE ALERT: ADRODC-DET;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460601042] HOST ALERT: ADRODC-DET;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460601050] SERVICE ALERT: VC;Service: RPC;OK;SOFT;2;RpcSs: Started
[1460601084] SERVICE ALERT: ADRODC-DET;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460601090] HOST ALERT: ADRODC-DET;DOWN;SOFT;5;CRITICAL - Time to live exceeded (198.109.113.140)
[1460601096] SERVICE ALERT: ADRODC-DET;Service: DNS Service;OK;SOFT;3;DNS: Started
[1460601099] SERVICE ALERT: ADRODC-DET;Service: Windows Time;OK;SOFT;3;W32Time: Started
[1460601112] SERVICE ALERT: ADRODC-DET;Service: Group Policy;OK;SOFT;3;gpsvc: Started
[1460601116] HOST ALERT: ADRODC-DET;UP;SOFT;6;PING OK - Packet loss = 0%, RTA = 8.74 ms
[1460601256] SERVICE ALERT: ADRODC-DET;Service: Remote Management;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460601262] SERVICE ALERT: ADRODC-DET;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460601283] SERVICE ALERT: ADRODC-DET;Service: Intersite Messaging;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460601286] HOST ALERT: ADRODC-DET;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460601317] SERVICE ALERT: ADRODC-DET;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460601321] SERVICE ALERT: ADRODC-DET;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460601323] SERVICE ALERT: ADRODC-DET;Service: Key Distribution Center;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460601323] SERVICE ALERT: ADRODC-DET;Service: AD Web Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460601331] HOST ALERT: ADRODC-DET;DOWN;SOFT;2;CRITICAL - Time to live exceeded (198.109.113.140)
[1460601341] SERVICE ALERT: ADRODC-DET;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460601345] HOST ALERT: ADRODC-DET;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 10.06 ms
[1460601356] SERVICE ALERT: ADRODC-DET;Service: Remote Management;OK;SOFT;2;WinRM: Started
[1460601362] SERVICE ALERT: ADRODC-DET;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1460601383] SERVICE ALERT: ADRODC-DET;Service: Intersite Messaging;OK;SOFT;2;IsmServ: Started
[1460601409] SERVICE ALERT: ADDC5;GPO Replication;WARNING;SOFT;1;WARNING: [PROCESS] Something wrong happened
[1460601417] SERVICE ALERT: ADRODC-DET;Memory Usage;OK;SOFT;2;Memory usage: total:37586.86 MB - used: 3189.25 MB (8%) - free: 34397.61 MB (92%)
[1460601421] SERVICE ALERT: ADRODC-DET;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1460601423] SERVICE ALERT: ADRODC-DET;Service: Key Distribution Center;OK;SOFT;2;KDC: Started
[1460601423] SERVICE ALERT: ADRODC-DET;Service: AD Web Service;OK;SOFT;2;ADWS: Started
[1460601441] SERVICE ALERT: ADRODC-DET;C:\ Drive Space;OK;SOFT;2;c:\ - total: 1115.16 Gb - used: 267.66 Gb (24%) - free 847.49 Gb (76%)
[1460601554] SERVICE ALERT: ADDC5;GPO Replication;OK;SOFT;2;WARNING: [PROCESS] Something wrong happened
[1460601664] SERVICE ALERT: ADRODC-DET;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1460601976] SERVICE ALERT: ADRODC-DET;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460601976] SERVICE ALERT: ADRODC-DET;Service: Remote Management;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460601979] HOST ALERT: ADRODC-DET;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460601981] SERVICE ALERT: ADRODC-DET;Service: PXE Boot;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460601981] SERVICE ALERT: ADRODC-DET;Service: DFS Namespace;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460601982] SERVICE ALERT: ADRODC-DET;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460601982] SERVICE ALERT: ADRODC-DET;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460601982] SERVICE ALERT: ADRODC-DET;Service: MDT Monitor;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460601982] SERVICE ALERT: ADRODC-DET;Service: Distributed File System Replication;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460601982] SERVICE ALERT: ADRODC-DET;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460601999] SERVICE ALERT: ADRODC-DET;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460602003] SERVICE ALERT: ADRODC-DET;Service: Intersite Messaging;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460602004] SERVICE ALERT: ADRODC-DET;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460602006] SERVICE ALERT: ADRODC-DET;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460602007] SERVICE ALERT: ADRODC-DET;Service: AD Domain Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460602011] HOST ALERT: ADRODC-DET;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460602037] SERVICE ALERT: ADRODC-DET;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460602040] SERVICE ALERT: ADRODC-DET;AD Replication;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 30 seconds.
[1460602041] SERVICE ALERT: ADRODC-DET;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460602043] SERVICE ALERT: ADRODC-DET;Service: Key Distribution Center;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460602043] SERVICE ALERT: ADRODC-DET;Service: AD Web Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460602061] SERVICE ALERT: ADRODC-DET;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460602067] HOST ALERT: ADRODC-DET;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460602096] SERVICE ALERT: ADRODC-DET;Service: Remote Management;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460602096] SERVICE ALERT: ADRODC-DET;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460602101] SERVICE ALERT: ADRODC-DET;Service: PXE Boot;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460602101] SERVICE ALERT: ADRODC-DET;Service: DFS Namespace;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460602102] SERVICE ALERT: ADRODC-DET;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460602102] SERVICE ALERT: ADRODC-DET;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460602102] SERVICE ALERT: ADRODC-DET;Service: MDT Monitor;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460602102] SERVICE ALERT: ADRODC-DET;Service: Distributed File System Replication;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460602102] SERVICE ALERT: ADRODC-DET;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460602119] SERVICE ALERT: ADRODC-DET;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460602123] SERVICE ALERT: ADRODC-DET;Service: Intersite Messaging;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460602124] SERVICE ALERT: ADRODC-DET;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460602126] SERVICE ALERT: ADRODC-DET;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460602126] HOST ALERT: ADRODC-DET;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460602127] SERVICE ALERT: ADRODC-DET;Service: AD Domain Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460602157] HOST ALERT: ADRODC-DET;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460602157] SERVICE ALERT: ADRODC-DET;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460602160] SERVICE ALERT: ADRODC-DET;AD Replication;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 30 seconds.
[1460602161] SERVICE ALERT: ADRODC-DET;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460602163] SERVICE ALERT: ADRODC-DET;Service: Key Distribution Center;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460602163] SERVICE ALERT: ADRODC-DET;Service: AD Web Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460602181] SERVICE ALERT: ADRODC-DET;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460602187] HOST ALERT: ADRODC-DET;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1460602216] SERVICE ALERT: ADRODC-DET;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460602216] SERVICE ALERT: ADRODC-DET;Service: Remote Management;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460602221] SERVICE ALERT: ADRODC-DET;Service: PXE Boot;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460602221] SERVICE ALERT: ADRODC-DET;Service: DFS Namespace;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460602222] SERVICE ALERT: ADRODC-DET;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460602222] SERVICE ALERT: ADRODC-DET;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460602222] SERVICE ALERT: ADRODC-DET;Service: MDT Monitor;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460602222] SERVICE ALERT: ADRODC-DET;Service: Distributed File System Replication;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460602222] SERVICE ALERT: ADRODC-DET;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460602239] SERVICE ALERT: ADRODC-DET;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460602243] SERVICE ALERT: ADRODC-DET;Service: Intersite Messaging;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460602244] SERVICE ALERT: ADRODC-DET;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460602246] SERVICE ALERT: ADRODC-DET;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460602246] HOST ALERT: ADRODC-DET;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1460602247] SERVICE ALERT: ADRODC-DET;Service: AD Domain Service;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460602270] HOST ALERT: ADRODC-DET;DOWN;SOFT;8;CRITICAL - Time to live exceeded (198.109.113.140)
[1460602270] SERVICE ALERT: ADRODC-DET;Service: NetLogin;CRITICAL;HARD;3;connect to address 198.109.113.140 and port 12489: No route to host
[1460602270] HOST ALERT: ADRODC-DET;DOWN;SOFT;9;CRITICAL - Time to live exceeded (198.109.113.140)
[1460602272] SERVICE ALERT: ADRODC-DET;Service: Key Distribution Center;OK;SOFT;3;KDC: Started
[1460602272] SERVICE ALERT: ADRODC-DET;Service: AD Web Service;OK;SOFT;3;ADWS: Started
[1460602277] SERVICE ALERT: ADRODC-DET;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1460602278] SERVICE ALERT: ADRODC-DET;AD Replication;OK;SOFT;3;OK: replication is up and running
[1460602281] SERVICE ALERT: ADRODC-DET;C:\ Drive Space;OK;SOFT;3;c:\ - total: 1115.16 Gb - used: 267.66 Gb (24%) - free 847.49 Gb (76%)
[1460602291] HOST ALERT: ADRODC-DET;UP;SOFT;10;PING OK - Packet loss = 0%, RTA = 8.60 ms
[1460602620] HOST ALERT: ADRODC-DET;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460602710] HOST ALERT: ADRODC-DET;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460602800] HOST ALERT: ADRODC-DET;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460602846] HOST ALERT: ADRODC-DET;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460602877] HOST ALERT: ADRODC-DET;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460602880] SERVICE ALERT: ADRODC-DET;AD Replication;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 30 seconds.
[1460602882] SERVICE ALERT: ADRODC-DET;Service: Key Distribution Center;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460602883] SERVICE ALERT: ADRODC-DET;Service: AD Web Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460602884] SERVICE ALERT: ADRODC-DET;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460602901] SERVICE ALERT: ADRODC-DET;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460602907] HOST ALERT: ADRODC-DET;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1460602907] SERVICE ALERT: ADRODC-DET;Service: DNS Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460602910] SERVICE ALERT: ADRODC-DET;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460602932] SERVICE ALERT: ADRODC-DET;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460602937] HOST ALERT: ADRODC-DET;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1460602956] SERVICE ALERT: CEDARON;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460602979] SERVICE ALERT: SANDBOX-DB;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460603000] SERVICE ALERT: ADRODC-DET;AD Replication;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 30 seconds.
[1460603001] SERVICE ALERT: ADRODC-DET;Service: Key Distribution Center;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460603002] SERVICE ALERT: ADRODC-DET;Service: AD Web Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460603004] SERVICE ALERT: ADRODC-DET;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460603021] SERVICE ALERT: ADRODC-DET;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460603027] SERVICE ALERT: ADRODC-DET;Service: DNS Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460603027] HOST ALERT: ADRODC-DET;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1460603030] SERVICE ALERT: ADRODC-DET;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460603046] SERVICE ALERT: NETENG-PA;CPU Load;WARNING;SOFT;1;CPU Load 91% (5 min average)
[1460603052] SERVICE ALERT: ADRODC-DET;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460603060] HOST ALERT: ADRODC-DET;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1460603063] SERVICE ALERT: PNCDB-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1460603070] SERVICE ALERT: OFFICEKMS;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1460603076] SERVICE ALERT: CEDARON;CPU Load;OK;SOFT;2;CPU Load 82% (5 min average)
[1460603096] SERVICE ALERT: ADRODC-DET;AD Replication;OK;SOFT;3;OK: replication is up and running
[1460603099] SERVICE ALERT: SANDBOX-DB;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1460603100] HOST ALERT: ADRODC-DET;UP;SOFT;10;PING OK - Packet loss = 0%, RTA = 62.76 ms
[1460603102] SERVICE ALERT: ADRODC-DET;Service: Key Distribution Center;OK;SOFT;3;KDC: Started
[1460603103] SERVICE ALERT: ADRODC-DET;Service: AD Web Service;OK;SOFT;3;ADWS: Started
[1460603104] SERVICE ALERT: ADRODC-DET;Service: Firewall;OK;SOFT;3;MpsSvc: Started
[1460603121] SERVICE ALERT: ADRODC-DET;C:\ Drive Space;OK;SOFT;3;c:\ - total: 1115.16 Gb - used: 267.67 Gb (24%) - free 847.49 Gb (76%)
[1460603127] SERVICE ALERT: ADRODC-DET;Service: DNS Service;OK;SOFT;3;DNS: Started
[1460603130] SERVICE ALERT: ADRODC-DET;Service: Windows Time;OK;SOFT;3;W32Time: Started
[1460603152] SERVICE ALERT: ADRODC-DET;Service: Group Policy;OK;SOFT;3;gpsvc: Started
[1460603167] SERVICE ALERT: NETENG-PA;CPU Load;OK;SOFT;2;CPU Load 87% (5 min average)
[1460603183] SERVICE ALERT: PNCDB-2012;CPU Load;OK;SOFT;2;CPU Load 86% (5 min average)
[1460603190] SERVICE ALERT: OFFICEKMS;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1460603219] SERVICE ALERT: SANDBOX-DB;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1460603287] SERVICE ALERT: XDM;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1460603396] SERVICE ALERT: ADRODC-DET;Service: Remote Management;OK;HARD;3;WinRM: Started
[1460603396] SERVICE ALERT: ADRODC-DET;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1460603400] SERVICE ALERT: ADRODC-DET;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1460603401] SERVICE ALERT: ADRODC-DET;Service: PXE Boot;OK;HARD;3;WDSServer: Started
[1460603401] SERVICE ALERT: ADRODC-DET;Service: DFS Namespace;OK;HARD;3;Dfs: Started
[1460603402] SERVICE ALERT: ADRODC-DET;Service: MDT Monitor;OK;HARD;3;MDT_Monitor: Started
[1460603402] SERVICE ALERT: ADRODC-DET;Uptime;OK;HARD;3;System Uptime - 5 day(s) 3 hour(s) 56 minute(s)
[1460603402] SERVICE ALERT: ADRODC-DET;Service: Distributed File System Replication;OK;HARD;3;DFSR: Started
[1460603402] SERVICE ALERT: ADRODC-DET;CPU Load;OK;HARD;3;CPU Load 0% (5 min average)
[1460603407] SERVICE ALERT: XDM;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1460603419] SERVICE ALERT: ADRODC-DET;Service: RPC;OK;HARD;3;RpcSs: Started
[1460603423] SERVICE ALERT: ADRODC-DET;Service: Intersite Messaging;OK;HARD;3;IsmServ: Started
[1460603424] SERVICE ALERT: ADRODC-DET;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1460603426] SERVICE ALERT: ADRODC-DET;Service: WMI;OK;HARD;3;Winmgmt: Started
[1460603427] SERVICE ALERT: ADRODC-DET;Service: AD Domain Service;OK;HARD;3;NTDS: Started
[1460603440] SERVICE ALERT: XDM;Service: Group Policy;CRITICAL;SOFT;1;gpsvc: Stopped
[1460603457] SERVICE ALERT: ADRODC-DET;Memory Usage;OK;HARD;3;Memory usage: total:37586.86 MB - used: 3188.95 MB (8%) - free: 34397.91 MB (92%)
[1460603461] SERVICE ALERT: ADRODC-DET;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1460603470] SERVICE ALERT: XDM;Service: NetLogin;CRITICAL;SOFT;1;connect to address 164.76.2.79 and port 12489: Connection refused
[1460603500] HOST ALERT: XDM;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460603527] SERVICE ALERT: XDM;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1460603531] HOST ALERT: XDM;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1460603559] SERVICE ALERT: XDM;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1460603581] SERVICE ALERT: XDM;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1460603654] SERVICE ALERT: MBAM;Uptime;CRITICAL;SOFT;1;connect to address 164.76.7.98 and port 12489: Connection refused
[1460603657] SERVICE ALERT: MBAMDB;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1460603670] HOST ALERT: MBAM;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460603704] SERVICE ALERT: TI22-COE;Service: Group Policy;CRITICAL;SOFT;1;gpsvc: Stopped
[1460603707] SERVICE ALERT: ADWU1;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.4.133 and port 12489: Connection refused
[1460603708] HOST ALERT: MBAM;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1460603708] SERVICE ALERT: ADWU1;CPU Load;CRITICAL;SOFT;1;connect to address 164.76.4.133 and port 12489: Connection refused
[1460603718] SERVICE ALERT: IMAGES;PXE Boot Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460603726] SERVICE ALERT: OFFICEKMS;VMWare Tools;CRITICAL;SOFT;1;connect to address 164.76.4.134 and port 12489: Connection refused
[1460603737] SERVICE ALERT: OFFICEKMS;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460603748] HOST ALERT: IMAGES;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460603750] SERVICE ALERT: WINKMS;Uptime;CRITICAL;SOFT;1;connect to address 164.76.2.49 and port 12489: Connection refused
[1460603752] SERVICE ALERT: PROCARE;CPU Load;WARNING;SOFT;1;CPU Load 94% (5 min average)
[1460603754] SERVICE ALERT: EMUSQLEXPRESS;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.7.169 and port 12489: Connection refused
[1460603763] SERVICE ALERT: TI22;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.99.227 and port 12489: Connection refused
[1460603764] SERVICE ALERT: WINKMS;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.2.49 and port 12489: Connection refused
[1460603770] SERVICE ALERT: TI22;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.99.227 and port 12489: Connection refused
[1460603774] SERVICE ALERT: MBAMDB;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1460603774] SERVICE ALERT: MBAM;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 2 minute(s)
[1460603781] SERVICE ALERT: HDBACKUPS;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460603782] SERVICE ALERT: CEDARON;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.99.137 and port 12489: Connection refused
[1460603784] SERVICE ALERT: COE-WEBDAV;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.99.167 and port 12489: Connection refused
[1460603806] SERVICE ALERT: FTK;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460603811] HOST ALERT: HDBACKUPS;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460603819] SERVICE ALERT: SANDBOX-DB;CPU Load;OK;HARD;3;CPU Load 2% (5 min average)
[1460603821] HOST ALERT: FTK;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460603824] SERVICE ALERT: TI22-COE;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1460603824] SERVICE ALERT: ADWU1;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1460603828] SERVICE ALERT: ADWU1;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1460603832] SERVICE ALERT: R25-ALPHA2012;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460603833] SERVICE ALERT: PNCWEB-2012;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460603836] SERVICE ALERT: HDBACKUPS;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460603837] SERVICE ALERT: OFFICEKMS;Memory Usage;OK;SOFT;2;Memory usage: total:10943.64 MB - used: 820.21 MB (7%) - free: 10123.43 MB (93%)
[1460603837] SERVICE ALERT: OFFICEKMS;VMWare Tools;OK;SOFT;2;VMTools: Started
[1460603838] SERVICE ALERT: IMAGES;PXE Boot Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460603838] HOST ALERT: IMAGES;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460603839] SERVICE ALERT: HDBACKUPS;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460603840] SERVICE ALERT: UHSPHARM-2012;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460603844] SERVICE ALERT: PNCAPP-2012;Service: Group Policy;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1460603845] SERVICE ALERT: Echo-01;Service: User Profile;CRITICAL;SOFT;1;connect to address 164.76.2.241 and port 12489: Connection refused
[1460603851] SERVICE ALERT: R25-ZETA2012;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.99.206 and port 12489: Connection refused
[1460603854] SERVICE ALERT: TI22-PSYCLINIC;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.99.219 and port 12489: Connection refused
[1460603860] SERVICE ALERT: HDBACKUPS;CPU Load;CRITICAL;SOFT;1;connect to address 164.76.4.109 and port 12489: Connection refused
[1460603861] SERVICE ALERT: FTK;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460603866] HOST ALERT: HDBACKUPS;UP;SOFT;2;PING WARNING - Packet loss = 96%, RTA = 0.29 ms
[1460603870] SERVICE ALERT: WINKMS;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 2 minute(s)
[1460603871] SERVICE ALERT: EMUSQLEXPRESS;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1460603872] SERVICE ALERT: PROCARE;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1460603873] SERVICE ALERT: MBAM;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1460603880] SERVICE ALERT: TI22;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1460603882] SERVICE ALERT: HDBACKUPS;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1460603884] SERVICE ALERT: WINKMS;Memory Usage;OK;SOFT;2;Memory usage: total:20479.73 MB - used: 1154.46 MB (6%) - free: 19325.26 MB (94%)
[1460603890] SERVICE ALERT: TI22;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1460603891] HOST ALERT: FTK;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460603896] SERVICE ALERT: PNCDB-2012;Service: Group Policy;CRITICAL;SOFT;1;gpsvc: Stopped
[1460603902] HOST ALERT: IMAGES;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1460603902] SERVICE ALERT: CEDARON;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1460603904] SERVICE ALERT: COE-WEBDAV;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1460603906] SERVICE ALERT: FTK;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460603926] SERVICE ALERT: FTK;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460603930] SERVICE ALERT: CBORD;Service: WMI;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1460603932] SERVICE ALERT: R25-ALPHA2012;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1460603933] SERVICE ALERT: PNCWEB-2012;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1460603934] SERVICE ALERT: DPS-SECURE02;Service: Group Policy;CRITICAL;SOFT;1;gpsvc: Stopped
[1460603936] SERVICE ALERT: HDBACKUPS;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 2 minute(s)
[1460603936] HOST ALERT: FTK;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460603938] SERVICE ALERT: IMAGES;PXE Boot Service;OK;SOFT;3;WDSServer: Started
[1460603939] SERVICE ALERT: HDBACKUPS;C:\ Drive Space;OK;SOFT;2;c:\ - total: 409.26 Gb - used: 141.55 Gb (35%) - free 267.70 Gb (65%)
[1460603940] SERVICE ALERT: UHSPHARM-2012;C:\ Drive Space;OK;SOFT;2;c:\ - total: 99.66 Gb - used: 49.12 Gb (49%) - free 50.54 Gb (51%)
[1460603941] SERVICE ALERT: PNCDB-2012;Service: NetLogin;CRITICAL;SOFT;1;connect to address 164.76.2.176 and port 12489: Connection refused
[1460603949] SERVICE ALERT: FTK;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460603956] SERVICE ALERT: Echo-01;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1460603961] SERVICE ALERT: FTK;Service: NetLogin;CRITICAL;SOFT;2;connect to address 164.76.7.103 and port 12489: Connection refused
[1460603964] SERVICE ALERT: PNCAPP-2012;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1460603967] SERVICE ALERT: CBORD;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.99.168 and port 12489: Connection refused
[1460603970] SERVICE ALERT: R25-ZETA2012;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1460603974] SERVICE ALERT: TI22-PSYCLINIC;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1460603979] HOST ALERT: FTK;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460603980] SERVICE ALERT: HDBACKUPS;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1460603993] SERVICE ALERT: MBAM;CPU Load;WARNING;SOFT;2;CPU Load 92% (5 min average)
[1460604006] SERVICE ALERT: FTK;Service: RPC;OK;SOFT;2;RpcSs: Started
[1460604010] HOST ALERT: FTK;UP;SOFT;5;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1460604016] SERVICE ALERT: PNCDB-2012;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1460604026] SERVICE ALERT: FTK;Service: Group Policy;OK;SOFT;3;gpsvc: Started
[1460604035] SERVICE ALERT: DPS-SECURE02;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460604049] SERVICE ALERT: FTK;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1460604050] SERVICE ALERT: CBORD;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1460604061] SERVICE ALERT: PNCDB-2012;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1460604061] SERVICE ALERT: TI22;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1460604065] HOST ALERT: DPS-SECURE02;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460604074] SERVICE ALERT: DPS-SECURE02;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460604081] SERVICE ALERT: FTK;Service: NetLogin;OK;SOFT;3;Netlogon: Started
[1460604086] SERVICE ALERT: DPS-SECURE02;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460604087] SERVICE ALERT: CBORD;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1460604104] HOST ALERT: DPS-SECURE02;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460604113] SERVICE ALERT: MBAM;CPU Load;OK;SOFT;3;CPU Load 0% (5 min average)
[1460604121] SERVICE ALERT: TI22-PSYCLINIC;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1460604127] SERVICE ALERT: XDM;CPU Load;OK;HARD;3;CPU Load 3% (5 min average)
[1460604130] SERVICE ALERT: DPS-SECURE02;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.2.46 and port 12489: Connection refused
[1460604135] SERVICE ALERT: DPS-SECURE02;Service: RPC;CRITICAL;SOFT;2;connect to address 164.76.2.46 and port 12489: Connection refused
[1460604160] HOST ALERT: DPS-SECURE02;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460604172] SERVICE ALERT: DPS-SECURE02;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1460604174] SERVICE ALERT: DPS-SECURE02;Service: Group Policy;OK;SOFT;3;gpsvc: Started
[1460604181] HOST ALERT: DPS-SECURE02;UP;SOFT;4;PING OK - Packet loss = 50%, RTA = 0.35 ms
[1460604181] SERVICE ALERT: TI22;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1460604186] SERVICE ALERT: DPS-SECURE02;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1460604189] SERVICE ALERT: EMU-VALUE01;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460604219] HOST ALERT: EMU-VALUE01;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1460604241] SERVICE ALERT: TI22-PSYCLINIC;CPU Load;OK;SOFT;2;CPU Load 1% (5 min average)
[1460604241] SERVICE ALERT: DPS-SECURE02;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1460604255] SERVICE ALERT: DPS-SECURE02;Service: RPC;OK;SOFT;3;RpcSs: Started
[1460604266] SERVICE ALERT: EMU-VALUE01;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460604291] SERVICE ALERT: DPS-SECURE02;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1460604296] HOST ALERT: EMU-VALUE01;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1460604309] SERVICE ALERT: EMU-VALUE01;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1460604321] SERVICE ALERT: EMU-VALUE01;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1460604339] HOST ALERT: EMU-VALUE01;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1460604357] SERVICE ALERT: EMU-VALUE01;CPU Load;CRITICAL;SOFT;1;connect to address 164.76.4.85 and port 12489: Connection refused
[1460604365] SERVICE ALERT: EMU-VALUE01;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.4.85 and port 12489: Connection refused
[1460604366] SERVICE ALERT: EMU-VALUE01;Uptime;CRITICAL;SOFT;2;connect to address 164.76.4.85 and port 12489: Connection refused
[1460604384] SERVICE ALERT: EMU-VALUE01;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.4.85 and port 12489: Connection refused
[1460604387] HOST ALERT: EMU-VALUE01;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1460604409] SERVICE ALERT: EMU-VALUE01;Service: WMI;CRITICAL;HARD;3;connect to address 164.76.4.85 and port 12489: Connection refused
[1460604412] SERVICE ALERT: REFWIN;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.158.184 and port 12489: Connection refused
[1460604412] SERVICE ALERT: REFWIN;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.158.184 and port 12489: Connection refused
[1460604412] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 94% (5 min average)
[1460604421] SERVICE ALERT: EMU-VALUE01;C:\ Drive Space;CRITICAL;SOFT;2;connect to address 164.76.4.85 and port 12489: Connection refused
[1460604434] SERVICE ALERT: EMU-VALUE01;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.4.85 and port 12489: Connection refused
[1460604439] HOST ALERT: EMU-VALUE01;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1460604475] SERVICE ALERT: EMU-VALUE01;CPU Load;OK;SOFT;2;CPU Load 25% (5 min average)
[1460604479] HOST ALERT: EMU-VALUE01;UP;SOFT;6;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1460604485] SERVICE ALERT: EMU-VALUE01;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1460604486] SERVICE ALERT: EMU-VALUE01;Uptime;OK;SOFT;3;System Uptime - 0 day(s) 0 hour(s) 4 minute(s)
[1460604504] SERVICE ALERT: EMU-VALUE01;Memory Usage;OK;SOFT;2;Memory usage: total:9444.43 MB - used: 2077.46 MB (22%) - free: 7366.97 MB (78%)
[1460604532] SERVICE ALERT: REFWIN;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1460604532] SERVICE ALERT: REFWIN;Service: RPC;OK;SOFT;2;RpcSs: Started
[1460604532] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 91% (5 min average)
[1460604541] SERVICE ALERT: EMU-VALUE01;C:\ Drive Space;OK;SOFT;3;c:\ - total: 465.42 Gb - used: 15.38 Gb (3%) - free 450.03 Gb (97%)
[1460604554] SERVICE ALERT: EMU-VALUE01;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1460604652] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 81% (5 min average)
[1460605009] SERVICE ALERT: EMU-VALUE01;Service: WMI;OK;HARD;3;Winmgmt: Started