[1457327136] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 87% (5 min average)
[1457327736] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457327856] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1457327976] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1457332176] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 68% (5 min average)
[1457332776] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457332896] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1457333016] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 72% (5 min average)
[1457333616] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1457333736] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1457333856] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 69% (5 min average)
[1457334456] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457334576] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1457334696] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1457337309] SERVICE ALERT: MBAM;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457337429] SERVICE ALERT: MBAM;CPU Load;OK;SOFT;2;CPU Load 20% (5 min average)
[1457342496] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 56% (5 min average)
[1457343096] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1457343216] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 63% (5 min average)
[1457344416] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1457344536] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1457344656] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1457345256] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 84% (5 min average)
[1457348856] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457348976] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1457349096] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1457352946] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;SOFT;1;WARNING: Replication with ADDC5 failed!
[1457353065] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;SOFT;2;WARNING: Replication with ADDC5 failed!
[1457353186] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;HARD;3;WARNING: Replication with ADDC5 failed!
[1457354298] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;OK;HARD;3;Black Cartridge HP CE260X is at 100% - OK!
[1457356896] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 48% (5 min average)
[1457357471] SERVICE ALERT: IMAGES;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457357591] SERVICE ALERT: IMAGES;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1457359023] SERVICE ALERT: ADDC5;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457359023] SERVICE ALERT: ADDC5;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457359123] SERVICE ALERT: ADDC5;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1457359123] SERVICE ALERT: ADDC5;C:\ Drive Space;OK;SOFT;2;c:\ - total: 79.66 Gb - used: 11.26 Gb (14%) - free 68.40 Gb (86%)
[1457359296] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457359416] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1457359536] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1457360058] SERVICE ALERT: ADRODC-LIV;AD Replication;OK;HARD;3;OK: replication is up and running
[1457360187] SERVICE ALERT: owen101lab3;Device Status;OK;HARD;3;Warming up
[1457360191] HOST ALERT: owen101lab3;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.56 ms
[1457360534] SERVICE ALERT: owen101lab3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1457360551] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 99% - OK!
[1457360736] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 86% (5 min average)
[1457360793] SERVICE ALERT: owen101lab3;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1457360823] HOST ALERT: owen101lab3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1457360853] SERVICE ALERT: owen101lab3;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1457360883] HOST ALERT: owen101lab3;DOWN;SOFT;2;(Host check timed out after 30.00 seconds)
[1457360913] SERVICE ALERT: owen101lab3;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1457360943] HOST ALERT: owen101lab3;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1457361033] HOST ALERT: owen101lab3;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1457361123] HOST ALERT: owen101lab3;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1457361139] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1457361157] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1457361169] HOST ALERT: owen101lab3;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1457361199] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1457361217] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1457361229] HOST ALERT: owen101lab3;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1457361259] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1457361277] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1457361289] HOST ALERT: owen101lab3;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1457361336] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457361379] HOST ALERT: owen101lab3;DOWN;SOFT;9;(Host check timed out after 30.00 seconds)
[1457361456] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1457361469] HOST ALERT: owen101lab3;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1457361576] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 72% (5 min average)
[1457362393] SERVICE ALERT: ADDC5;GPO Replication;WARNING;SOFT;1;Warning! 1 policies are not replicating properly!
[1457362686] SERVICE ALERT: ADDC5;GPO Replication;WARNING;SOFT;2;Warning! 1 policies are not replicating properly!
[1457362988] SERVICE ALERT: ADDC5;GPO Replication;OK;SOFT;3;OK! All GPOs are properly replicated.
[1457363691] SERVICE ALERT: mckennypk;Black Cartridge HP CC364X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.152.26. Is the printer off?  Is the network down?
[1457363745] SERVICE ALERT: mckennypk;Black Cartridge HP CC364X;OK;SOFT;2;Black Cartridge HP CC364X is at 14% - OK!
[1457363976] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457364096] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1457364216] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1457365391] SERVICE ALERT: IMAGES;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457365511] SERVICE ALERT: IMAGES;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1457365855] SERVICE ALERT: SCCM-PRIM01;Uptime;CRITICAL;SOFT;1;connect to address 164.76.102.80 and port 12489: No route to host
[1457365858] HOST ALERT: SCCM-PRIM01;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.80)
[1457365919] HOST ALERT: SCCM-PRIM01;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.80)
[1457365964] SERVICE ALERT: SCCM-PRIM01;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 1 minute(s)
[1457365967] SERVICE ALERT: SCCM-PRIM01;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457365968] HOST ALERT: SCCM-PRIM01;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.54 ms
[1457366047] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1457366067] SERVICE ALERT: SCCM-PRIM01;Memory Usage;OK;SOFT;2;Memory usage: total:32756.36 MB - used: 4790.26 MB (15%) - free: 27966.10 MB (85%)
[1457366816] SERVICE ALERT: hallegroundcolor;Cyan Cartridge HP CE261A;OK;HARD;3;Cyan Cartridge HP CE261A is at 100% - OK!
[1457370899] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 74% (5 min average)
[1457371499] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457371619] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 94% (5 min average)
[1457371736] SERVICE ALERT: DB01-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457371739] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1457371856] SERVICE ALERT: DB01-2012;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1457372339] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 17% (5 min average)
[1457374139] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457374259] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1457374379] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1457375402] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (26.8% change >= 20.0% threshold)
[1457375579] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 92% (5 min average)
[1457376179] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1457376334] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;CRITICAL;SOFT;1;Magenta Cartridge HP CE263A is at 1% - CRITICAL!
[1457376393] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;CRITICAL;SOFT;2;Magenta Cartridge HP CE263A is at 1% - CRITICAL!
[1457376454] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;CRITICAL;HARD;3;Magenta Cartridge HP CE263A is at 1% - CRITICAL!
[1457376779] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 68% (5 min average)
[1457377379] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457377499] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1457377619] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1457380619] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 92% (5 min average)
[1457381219] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1457383105] SERVICE ALERT: SCCM-PRIM01;Service: User Profile;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1457383136] SERVICE ALERT: SCCM-PRIM01;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457383139] SERVICE ALERT: SCCM-PRIM01;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1457383159] HOST ALERT: SCCM-PRIM01;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.80)
[1457383166] SERVICE ALERT: SCCM-PRIM01;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.102.80 and port 12489: No route to host
[1457383169] HOST ALERT: SCCM-PRIM01;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.80)
[1457383228] SERVICE ALERT: SCCM-PRIM01;Service: User Profile;CRITICAL;SOFT;2;connect to address 164.76.102.80 and port 12489: No route to host
[1457383231] HOST ALERT: SCCM-PRIM01;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.102.80)
[1457383239] SERVICE ALERT: SCCM-PRIM01;Service: Windows Time;CRITICAL;SOFT;2;connect to address 164.76.102.80 and port 12489: No route to host
[1457383242] SERVICE ALERT: SCCM-PRIM01;Service: NetLogin;CRITICAL;SOFT;2;connect to address 164.76.102.80 and port 12489: No route to host
[1457383293] SERVICE ALERT: SCCM-PRIM01;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1457383297] HOST ALERT: SCCM-PRIM01;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.56 ms
[1457383345] SERVICE ALERT: SCCM-PRIM01;Service: User Profile;OK;SOFT;3;ProfSvc: Started
[1457383356] SERVICE ALERT: SCCM-PRIM01;Service: Windows Time;OK;SOFT;3;W32Time: Started
[1457383359] SERVICE ALERT: SCCM-PRIM01;Service: NetLogin;OK;SOFT;3;Netlogon: Started
[1457383407] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 81% (5 min average)
[1457384304] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457384424] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1457384544] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1457384581] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (44.9% change >= 20.0% threshold)
[1457384643] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (44.9% change >= 20.0% threshold)
[1457384696] SERVICE ALERT: AD-RMAD;Windows Firewall Status;WARNING;SOFT;1;WARNING: Firewall is Disabled.
[1457384751] SERVICE ALERT: ADDC5;GPO Replication;WARNING;SOFT;1;Warning! 1 policies are not replicating properly!
[1457384814] SERVICE ALERT: AD-RMAD;Windows Firewall Status;WARNING;SOFT;2;WARNING: Firewall is Disabled.
[1457384922] SERVICE ALERT: AD-RMAD;Windows Firewall Status;OK;SOFT;3;OK: Firewall is Enabled
[1457384998] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (44.9% change >= 20.0% threshold)
[1457385048] SERVICE ALERT: ADAUTHDC1;Windows Firewall Status;WARNING;SOFT;1;\\ad\Nagios\Scripts\WinTeamPowershell.ps1 : The term
[1457385071] SERVICE ALERT: VC;Windows Firewall Status;WARNING;SOFT;1;File \\ad\Nagios\Scripts\WinTeamPowershell.ps1 cannot be loaded because the exe
[1457385094] SERVICE ALERT: ADDFS01;Windows Firewall Status;WARNING;SOFT;1;(No output on stdout) stderr: Could not resolve hostname ADDFS01: Name or service not known
[1457385098] SERVICE ALERT: XDM;Windows Firewall Status;WARNING;SOFT;1;File \\ad\Nagios\Scripts\WinTeamPowershell.ps1 cannot be loaded because the exe
[1457385099] SERVICE ALERT: ADAUTHDC2;Windows Firewall Status;WARNING;SOFT;1;\\ad\Nagios\Scripts\WinTeamPowershell.ps1 : The term
[1457385142] SERVICE ALERT: ADRODC-DET;Windows Firewall Status;UNKNOWN;SOFT;1;Unknown command(s): winteampowershell
[1457385143] SERVICE ALERT: EVENTSTEST-2012;Windows Firewall Status;WARNING;SOFT;1;(No output on stdout) stderr: Could not resolve hostname EVENTSTEST-2012: Name or service not known
[1457385144] SERVICE ALERT: cava2;Windows Firewall Status;UNKNOWN;SOFT;1;Unknown command(s): winteampowershell
[1457385146] SERVICE ALERT: ADAUTHDC3;Windows Firewall Status;WARNING;SOFT;1;\\ad\Nagios\Scripts\WinTeamPowershell.ps1 : The term
[1457385168] SERVICE ALERT: ADAUTHDC1;Windows Firewall Status;WARNING;SOFT;2;\\ad\Nagios\Scripts\WinTeamPowershell.ps1 : The term
[1457385168] SERVICE ALERT: VC;Windows Firewall Status;WARNING;SOFT;2;File \\ad\Nagios\Scripts\WinTeamPowershell.ps1 cannot be loaded because the exe
[1457385189] SERVICE ALERT: ADCA1;Windows Firewall Status;WARNING;SOFT;1;(No output on stdout) stderr: Could not resolve hostname ADCA1: Name or service not known
[1457385190] SERVICE ALERT: ADRODC-LIV;Windows Firewall Status;UNKNOWN;SOFT;1;Unknown command(s): winteampowershell
[1457385214] SERVICE ALERT: ADDFS01;Windows Firewall Status;WARNING;SOFT;2;(No output on stdout) stderr: Could not resolve hostname ADDFS01: Name or service not known
[1457385215] SERVICE ALERT: ADAUTHDC2;Windows Firewall Status;WARNING;SOFT;2;\\ad\Nagios\Scripts\WinTeamPowershell.ps1 : The term
[1457385216] SERVICE ALERT: XDM;Windows Firewall Status;WARNING;SOFT;2;File \\ad\Nagios\Scripts\WinTeamPowershell.ps1 cannot be loaded because the exe
[1457385221] SERVICE ALERT: ADDC5;GPO Replication;WARNING;SOFT;2;Warning! 1 policies are not replicating properly!
[1457385237] SERVICE ALERT: ADCA2;Windows Firewall Status;WARNING;SOFT;1;(No output on stdout) stderr: Could not resolve hostname ADCA2: Name or service not known
[1457385238] SERVICE ALERT: ADWSUSCLIENTS;Windows Firewall Status;UNKNOWN;SOFT;1;No handler for command: winteampowershell
[1457385245] SERVICE ALERT: IMAGES;Windows Firewall Status;WARNING;SOFT;1;\\ad\Nagios\Scripts\WinTeamPowershell.ps1 : File
[1457385262] SERVICE ALERT: ADRODC-DET;Windows Firewall Status;UNKNOWN;SOFT;2;Unknown command(s): winteampowershell
[1457385263] SERVICE ALERT: ADAUTHDC3;Windows Firewall Status;WARNING;SOFT;2;\\ad\Nagios\Scripts\WinTeamPowershell.ps1 : The term
[1457385263] SERVICE ALERT: EVENTSTEST-2012;Windows Firewall Status;WARNING;SOFT;2;(No output on stdout) stderr: Could not resolve hostname EVENTSTEST-2012: Name or service not known
[1457385264] SERVICE ALERT: cava2;Windows Firewall Status;UNKNOWN;SOFT;2;Unknown command(s): winteampowershell
[1457385286] SERVICE ALERT: ADWSUSSERVERS;Windows Firewall Status;UNKNOWN;SOFT;1;No handler for command: winteampowershell
[1457385288] SERVICE ALERT: ADAUTHDC1;Windows Firewall Status;WARNING;HARD;3;\\ad\Nagios\Scripts\WinTeamPowershell.ps1 : The term
[1457385288] SERVICE ALERT: VC;Windows Firewall Status;WARNING;HARD;3;File \\ad\Nagios\Scripts\WinTeamPowershell.ps1 cannot be loaded because the exe
[1457385309] SERVICE ALERT: ADCA1;Windows Firewall Status;WARNING;SOFT;2;(No output on stdout) stderr: Could not resolve hostname ADCA1: Name or service not known
[1457385310] SERVICE ALERT: ADRODC-LIV;Windows Firewall Status;UNKNOWN;SOFT;2;Unknown command(s): winteampowershell
[1457385334] SERVICE ALERT: ADDFS01;Windows Firewall Status;WARNING;HARD;3;(No output on stdout) stderr: Could not resolve hostname ADDFS01: Name or service not known
[1457385336] SERVICE ALERT: ADAUTHDC2;Windows Firewall Status;WARNING;HARD;3;\\ad\Nagios\Scripts\WinTeamPowershell.ps1 : The term
[1457385336] SERVICE ALERT: XDM;Windows Firewall Status;WARNING;HARD;3;File \\ad\Nagios\Scripts\WinTeamPowershell.ps1 cannot be loaded because the exe
[1457385357] SERVICE ALERT: ADCA2;Windows Firewall Status;WARNING;SOFT;2;(No output on stdout) stderr: Could not resolve hostname ADCA2: Name or service not known
[1457385357] SERVICE ALERT: EMUPRINTPS;Windows Firewall Status;WARNING;SOFT;1;(No output on stdout) stderr: Could not resolve hostname EMUPRINTPS: Name or service not known
[1457385358] SERVICE ALERT: ADWSUSCLIENTS;Windows Firewall Status;UNKNOWN;SOFT;2;No handler for command: winteampowershell
[1457385359] SERVICE ALERT: MBAM;Windows Firewall Status;WARNING;SOFT;1;\\ad\Nagios\Scripts\WinTeamPowershell.ps1 : File
[1457385359] SERVICE ALERT: IMAGES;Windows Firewall Status;WARNING;SOFT;2;\\ad\Nagios\Scripts\WinTeamPowershell.ps1 : File
[1457385382] SERVICE ALERT: ADRODC-DET;Windows Firewall Status;UNKNOWN;HARD;3;Unknown command(s): winteampowershell
[1457385383] SERVICE ALERT: ADAUTHDC3;Windows Firewall Status;WARNING;HARD;3;\\ad\Nagios\Scripts\WinTeamPowershell.ps1 : The term
[1457385383] SERVICE ALERT: EVENTSTEST-2012;Windows Firewall Status;WARNING;HARD;3;(No output on stdout) stderr: Could not resolve hostname EVENTSTEST-2012: Name or service not known
[1457385384] SERVICE ALERT: cava2;Windows Firewall Status;UNKNOWN;HARD;3;Unknown command(s): winteampowershell
[1457385386] SERVICE ALERT: MBAMDB;Windows Firewall Status;WARNING;SOFT;1;\\ad\Nagios\Scripts\WinTeamPowershell.ps1 : File
[1457385406] SERVICE ALERT: ADWSUSSERVERS;Windows Firewall Status;UNKNOWN;SOFT;2;No handler for command: winteampowershell
[1457385429] SERVICE ALERT: ADCA1;Windows Firewall Status;WARNING;HARD;3;(No output on stdout) stderr: Could not resolve hostname ADCA1: Name or service not known
[1457385430] SERVICE ALERT: ADRODC-LIV;Windows Firewall Status;UNKNOWN;HARD;3;Unknown command(s): winteampowershell
[1457385434] SERVICE ALERT: IMAGES;Windows Firewall Status;WARNING;HARD;3;\\ad\Nagios\Scripts\WinTeamPowershell.ps1 : File
[1457385436] SERVICE ALERT: SCCM-PRIM01;Windows Firewall Status;WARNING;SOFT;1;(No output on stdout) stderr: Could not resolve hostname SCCM-PRIM01: Name or service not known
[1457385455] SERVICE ALERT: MBAM;Windows Firewall Status;WARNING;SOFT;2;\\ad\Nagios\Scripts\WinTeamPowershell.ps1 : File
[1457385477] SERVICE ALERT: ADCA2;Windows Firewall Status;WARNING;HARD;3;(No output on stdout) stderr: Could not resolve hostname ADCA2: Name or service not known
[1457385477] SERVICE ALERT: EMUPRINTPS;Windows Firewall Status;WARNING;SOFT;2;(No output on stdout) stderr: Could not resolve hostname EMUPRINTPS: Name or service not known
[1457385478] SERVICE ALERT: ADWSUSCLIENTS;Windows Firewall Status;UNKNOWN;HARD;3;No handler for command: winteampowershell
[1457385483] SERVICE ALERT: CSCFILESERV;Windows Firewall Status;WARNING;SOFT;1;File \\ad\Nagios\Scripts\WinTeamPowershell.ps1 cannot be loaded because the exe
[1457385503] SERVICE ALERT: MBAMDB;Windows Firewall Status;WARNING;SOFT;2;\\ad\Nagios\Scripts\WinTeamPowershell.ps1 : File
[1457385516] SERVICE ALERT: ADDC5;GPO Replication;OK;SOFT;3;OK! All GPOs are properly replicated.
[1457385526] SERVICE ALERT: ADWSUSSERVERS;Windows Firewall Status;UNKNOWN;HARD;3;No handler for command: winteampowershell
[1457385526] SERVICE ALERT: DB01-2012;Windows Firewall Status;WARNING;SOFT;1;(No output on stdout) stderr: Could not resolve hostname DB01-2012: Name or service not known
[1457385551] SERVICE ALERT: REFWIN;Windows Firewall Status;WARNING;SOFT;1;File \\ad\Nagios\Scripts\WinTeamPowershell.ps1 cannot be loaded because the exe
[1457385556] SERVICE ALERT: SCCM-PRIM01;Windows Firewall Status;WARNING;SOFT;2;(No output on stdout) stderr: Could not resolve hostname SCCM-PRIM01: Name or service not known
[1457385574] SERVICE ALERT: EMU-VALUE01;Windows Firewall Status;WARNING;SOFT;1;(No output on stdout) stderr: Could not resolve hostname EMU-VALUE01: Name or service not known
[1457385575] SERVICE ALERT: SANDBOX-DB;Windows Firewall Status;WARNING;SOFT;1;(No output on stdout) stderr: Could not resolve hostname SANDBOX-DB: Name or service not known
[1457385576] SERVICE ALERT: MBAM;Windows Firewall Status;WARNING;HARD;3;\\ad\Nagios\Scripts\WinTeamPowershell.ps1 : File
[1457385597] SERVICE ALERT: EMUPRINTPS;Windows Firewall Status;WARNING;HARD;3;(No output on stdout) stderr: Could not resolve hostname EMUPRINTPS: Name or service not known
[1457385598] SERVICE ALERT: CSCFILESERV;Windows Firewall Status;WARNING;SOFT;2;File \\ad\Nagios\Scripts\WinTeamPowershell.ps1 cannot be loaded because the exe
[1457385621] SERVICE ALERT: NETIDAD;Windows Firewall Status;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1457385624] SERVICE ALERT: MBAMDB;Windows Firewall Status;WARNING;HARD;3;\\ad\Nagios\Scripts\WinTeamPowershell.ps1 : File
[1457385646] SERVICE ALERT: DB01-2012;Windows Firewall Status;WARNING;SOFT;2;(No output on stdout) stderr: Could not resolve hostname DB01-2012: Name or service not known
[1457385647] SERVICE ALERT: REFWIN;Windows Firewall Status;WARNING;SOFT;2;File \\ad\Nagios\Scripts\WinTeamPowershell.ps1 cannot be loaded because the exe
[1457385648] SERVICE ALERT: EMUSCRIPT;Windows Firewall Status;WARNING;SOFT;1;At \\ad\Nagios\Scripts\WinTeamPowershell.ps1:585 char:35
[1457385676] SERVICE ALERT: SCCM-PRIM01;Windows Firewall Status;WARNING;HARD;3;(No output on stdout) stderr: Could not resolve hostname SCCM-PRIM01: Name or service not known
[1457385694] SERVICE ALERT: EMU-VALUE01;Windows Firewall Status;WARNING;SOFT;2;(No output on stdout) stderr: Could not resolve hostname EMU-VALUE01: Name or service not known
[1457385695] SERVICE ALERT: EMUSQLEXPRESS;Windows Firewall Status;WARNING;SOFT;1;At \\ad\Nagios\Scripts\WinTeamPowershell.ps1:585 char:35
[1457385695] SERVICE ALERT: SANDBOX-DB;Windows Firewall Status;WARNING;SOFT;2;(No output on stdout) stderr: Could not resolve hostname SANDBOX-DB: Name or service not known
[1457385715] SERVICE ALERT: EMUSCRIPT;Windows Firewall Status;OK;SOFT;2;OK: Firewall is Enabled
[1457385719] SERVICE ALERT: CSCFILESERV;Windows Firewall Status;WARNING;HARD;3;File \\ad\Nagios\Scripts\WinTeamPowershell.ps1 cannot be loaded because the exe
[1457385760] SERVICE ALERT: EMUSQLEXPRESS;Windows Firewall Status;OK;SOFT;2;OK: Firewall is Enabled
[1457385766] SERVICE ALERT: DB01-2012;Windows Firewall Status;WARNING;HARD;3;(No output on stdout) stderr: Could not resolve hostname DB01-2012: Name or service not known
[1457385768] SERVICE ALERT: REFWIN;Windows Firewall Status;WARNING;HARD;3;File \\ad\Nagios\Scripts\WinTeamPowershell.ps1 cannot be loaded because the exe
[1457385810] SERVICE ALERT: NETIDAD;Windows Firewall Status;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1457385814] SERVICE ALERT: EMU-VALUE01;Windows Firewall Status;WARNING;HARD;3;(No output on stdout) stderr: Could not resolve hostname EMU-VALUE01: Name or service not known
[1457385815] SERVICE ALERT: SANDBOX-DB;Windows Firewall Status;WARNING;HARD;3;(No output on stdout) stderr: Could not resolve hostname SANDBOX-DB: Name or service not known
[1457385952] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (42.6% change >= 20.0% threshold)
[1457386202] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (42.6% change >= 20.0% threshold)
[1457389376] SERVICE ALERT: owen101lab2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1457389436] SERVICE ALERT: owen101lab2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1457389496] SERVICE ALERT: owen101lab2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1457390696] SERVICE ALERT: owen101lab2;Device Status;OK;HARD;3;TRAY 3 EMPTY PLAIN LETTER   Ready
[1457394744] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1457395344] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1457396051] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;SOFT;1;WARNING: Replication with ADDC5 failed!
[1457396171] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;SOFT;2;WARNING: Replication with ADDC5 failed!
[1457396291] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;HARD;3;WARNING: Replication with ADDC5 failed!
[1457397744] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 87% (5 min average)
[1457398344] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457398464] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1457398584] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 89% (5 min average)
[1457399184] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457399304] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1457399424] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 97% (5 min average)
[1457404256] SERVICE ALERT: DB01-2012;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1457404376] SERVICE ALERT: DB01-2012;CPU Load;WARNING;SOFT;2;CPU Load 94% (5 min average)
[1457404496] SERVICE ALERT: DB01-2012;CPU Load;OK;SOFT;3;CPU Load 85% (5 min average)
[1457407095] SERVICE ALERT: ADRODC-LIV;AD Replication;OK;HARD;3;OK: replication is up and running
[1457407224] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 90% (5 min average)
[1457407824] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1457409025] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 12072;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1457409025] SERVICE ALERT: ADWSUSSERVERS;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.102.133 and port 12489: No route to host
[1457409028] HOST ALERT: ADWSUSSERVERS;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.133)
[1457409084] HOST ALERT: ADWSUSSERVERS;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1457409135] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 12072;OK;SOFT;2;Eventlog check ok
[1457409139] SERVICE ALERT: ADWSUSSERVERS;Service: DNS Client;OK;SOFT;2;Dnscache: Started