[1463545001] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463545601] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1463546201] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1463546321] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1463546441] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463547041] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463548241] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463548841] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463549441] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463551306] SERVICE ALERT: BOEAPPTEST;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1463551425] SERVICE ALERT: BOEAPPTEST;CPU Load;OK;SOFT;2;CPU Load 20% (15 min average)
[1463554841] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 72% (15 min average)
[1463555441] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463555561] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 82% (15 min average)
[1463556161] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463556281] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463556401] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1463557001] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463557243] SERVICE ALERT: AXIOM-TEST-2012;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1463557363] SERVICE ALERT: AXIOM-TEST-2012;CPU Load;OK;SOFT;2;CPU Load 72% (15 min average)
[1463557601] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463557963] SERVICE ALERT: AXIOM-TEST-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1463558083] SERVICE ALERT: AXIOM-TEST-2012;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463558203] SERVICE ALERT: AXIOM-TEST-2012;CPU Load;OK;SOFT;3;CPU Load 67% (15 min average)
[1463559401] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 67% (15 min average)
[1463561201] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463561321] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1463561441] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1463562041] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 77% (15 min average)
[1463563841] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463563961] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463564081] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 74% (15 min average)
[1463564681] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463564801] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 65% (15 min average)
[1463566001] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463566121] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1463566241] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463567441] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1463567997] SERVICE ALERT: BBTS;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463568033] SERVICE ALERT: BBTS;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463568065] SERVICE ALERT: BBTS;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463568083] SERVICE ALERT: BBTS;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463568117] SERVICE ALERT: BBTS;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463568118] SERVICE ALERT: BBTS_DB;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463568120] SERVICE ALERT: BBTS;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463568123] SERVICE ALERT: BBTS;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463568123] SERVICE ALERT: BBTS_DB;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463568133] SERVICE ALERT: BBTS_DB;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463568143] SERVICE ALERT: BBTS_DB;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463568153] SERVICE ALERT: BBTS;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463568171] SERVICE ALERT: BBTS_DB;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463568175] SERVICE ALERT: BBTS;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463568181] SERVICE ALERT: BBTS_DB;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463568185] SERVICE ALERT: BBTS;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463568185] SERVICE ALERT: BBTS_DB;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463568185] SERVICE ALERT: BBTS_DB;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463568203] SERVICE ALERT: BBTS;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463568203] SERVICE ALERT: BBTS;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463568205] SERVICE ALERT: BBTS;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463568237] SERVICE ALERT: BBTS;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1463568237] SERVICE ALERT: BBTS_DB;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463568238] SERVICE ALERT: BBTS;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463568240] SERVICE ALERT: BBTS;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463568243] SERVICE ALERT: BBTS;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463568243] SERVICE ALERT: BBTS_DB;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463568253] SERVICE ALERT: BBTS_DB;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463568263] SERVICE ALERT: BBTS_DB;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463568271] SERVICE ALERT: BBTS;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1463568285] SERVICE ALERT: BBTS_DB;Service: WMI;CRITICAL;SOFT;2;Winmgmt: Stopped
[1463568285] SERVICE ALERT: BBTS_DB;Service: RPC;OK;SOFT;2;RpcSs: Started
[1463568290] SERVICE ALERT: BBTS_DB;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1463568291] SERVICE ALERT: BBTS_DB;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463568295] SERVICE ALERT: BBTS;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463568305] SERVICE ALERT: BBTS;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1463568323] SERVICE ALERT: BBTS;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463568323] SERVICE ALERT: BBTS;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463568323] SERVICE ALERT: BBTS;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1463568325] SERVICE ALERT: BBTS;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463568337] SERVICE ALERT: BBTS_DB;NSClient++ Version;OK;SOFT;3;NSClient++ 0.4.4.15 2015-11-25
[1463568341] SERVICE ALERT: BBTS;Uptime;CRITICAL;SOFT;2;connect to address 164.76.227.3 and port 12489: Connection refused
[1463568343] SERVICE ALERT: BBTS;Service: NetLogin;CRITICAL;HARD;3;connect to address 164.76.227.3 and port 12489: Connection refused
[1463568343] SERVICE ALERT: BBTS_DB;Service: Group Policy;OK;SOFT;3;gpsvc: Started
[1463568343] SERVICE ALERT: BBTS;Service: WMI;CRITICAL;HARD;3;connect to address 164.76.227.3 and port 12489: Connection refused
[1463568353] SERVICE ALERT: BBTS_DB;Service: User Profile;OK;SOFT;3;ProfSvc: Started
[1463568363] SERVICE ALERT: BBTS_DB;Service: DNS Client;OK;SOFT;3;Dnscache: Started
[1463568391] SERVICE ALERT: BBTS_DB;Service: Firewall;OK;SOFT;3;MpsSvc: Started
[1463568395] SERVICE ALERT: BBTS;Service: Windows Time;OK;SOFT;3;W32Time: Started
[1463568405] SERVICE ALERT: BBTS_DB;Service: WMI;OK;SOFT;3;Winmgmt: Started
[1463568423] SERVICE ALERT: BBTS;CPU Load;OK;SOFT;2;CPU Load 33% (15 min average)
[1463568423] SERVICE ALERT: BBTS;Service: Group Policy;OK;SOFT;3;gpsvc: Started
[1463568425] SERVICE ALERT: BBTS;Service: Firewall;OK;SOFT;3;MpsSvc: Started
[1463568458] SERVICE ALERT: BBTS;Uptime;OK;SOFT;3;System Uptime - 0 day(s) 0 hour(s) 2 minute(s)
[1463568663] SERVICE ALERT: BBTS_DB;VMWare Tools;UNKNOWN;HARD;3;Failed to open service VMTools: 1060: The specified service does not exist as an installed service.
[1463568851] SERVICE ALERT: BBTS;Memory Usage;OK;HARD;3;Memory usage: total:28669.54 MB - used: 846.77 MB (3%) - free: 27822.77 MB (97%)
[1463568885] SERVICE ALERT: BBTS;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1463568903] SERVICE ALERT: BBTS;Service: RPC;OK;HARD;3;RpcSs: Started
[1463568940] SERVICE ALERT: BBTS;Service: WMI;OK;HARD;3;Winmgmt: Started
[1463568943] SERVICE ALERT: BBTS;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1463569263] SERVICE ALERT: BBTS_DB;VMWare Tools;OK;HARD;3;VMTools: Started
[1463569417] SERVICE ALERT: BBTS;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1463571041] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463571641] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1463571837] SERVICE ALERT: COE-WEBDAV;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463571839] HOST ALERT: COE-WEBDAV;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1463571848] SERVICE ALERT: COE-WEBDAV;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463571868] SERVICE ALERT: COE-WEBDAV;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463571878] HOST ALERT: COE-WEBDAV;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1463571924] SERVICE ALERT: COE-WEBDAV;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463571954] HOST ALERT: COE-WEBDAV;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1463571957] SERVICE ALERT: COE-WEBDAV;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463571968] SERVICE ALERT: COE-WEBDAV;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463571975] SERVICE ALERT: COE-WEBDAV;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463571987] HOST ALERT: COE-WEBDAV;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1463571988] SERVICE ALERT: COE-WEBDAV;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463571989] SERVICE ALERT: COE-WEBDAV;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463572018] HOST ALERT: COE-WEBDAV;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1463572043] SERVICE ALERT: COE-WEBDAV;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463572044] SERVICE ALERT: COE-WEBDAV;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463572055] SERVICE ALERT: COE-WEBDAV;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463572073] HOST ALERT: COE-WEBDAV;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1463572077] SERVICE ALERT: COE-WEBDAV;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1463572088] SERVICE ALERT: COE-WEBDAV;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1463572095] SERVICE ALERT: COE-WEBDAV;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463572107] HOST ALERT: COE-WEBDAV;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1463572108] SERVICE ALERT: COE-WEBDAV;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1463572109] SERVICE ALERT: COE-WEBDAV;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463572138] HOST ALERT: COE-WEBDAV;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1463572143] SERVICE ALERT: COE-WEBDAV;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463572152] SERVICE ALERT: COE-WEBDAV;Service: RPC;OK;SOFT;2;RpcSs: Started
[1463572153] SERVICE ALERT: COE-WEBDAV;Service: Group Policy;OK;SOFT;3;gpsvc: Started
[1463572155] SERVICE ALERT: COE-WEBDAV;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1463572161] HOST ALERT: COE-WEBDAV;UP;SOFT;9;PING OK - Packet loss = 73%, RTA = 0.31 ms
[1463572195] SERVICE ALERT: COE-WEBDAV;Service: NetLogin;OK;SOFT;3;Netlogon: Started
[1463572208] SERVICE ALERT: COE-WEBDAV;Service: Firewall;UNKNOWN;HARD;3;Failed to open service manager: 1115: A system shutdown is in progress.
[1463572243] SERVICE ALERT: COE-WEBDAV;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1463572657] SERVICE ALERT: COE-WEBDAV;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1463572668] SERVICE ALERT: COE-WEBDAV;CPU Load;OK;HARD;3;CPU Load 2% (15 min average)
[1463572688] SERVICE ALERT: COE-WEBDAV;Service: Windows Time;OK;HARD;3;W32Time: Started
[1463572808] SERVICE ALERT: COE-WEBDAV;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1463572841] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 60% (15 min average)
[1463573754] SERVICE ALERT: PAPERCUT;Service: Group Policy;CRITICAL;SOFT;1;gpsvc: Stopped
[1463573874] SERVICE ALERT: PAPERCUT;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1463574303] SERVICE ALERT: hallefirstcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1463574364] SERVICE ALERT: hallefirstcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1463574423] SERVICE ALERT: hallefirstcolor;Device Status;OK;SOFT;3;Open tray indicated.
[1463574641] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463574761] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 80% (15 min average)
[1463575023] SERVICE ALERT: hallefirstcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1463575031] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 209 found since last run. oneLevelSearchOps OK. searchOps is too high! 210 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463575083] SERVICE ALERT: hallefirstcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1463575143] SERVICE ALERT: hallefirstcolor;Device Status;OK;SOFT;3;Processing job from tray 1
[1463575151] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 231 found since last run. oneLevelSearchOps OK. searchOps is too high! 232 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463575271] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463575871] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 213 found since last run. oneLevelSearchOps OK. searchOps is too high! 214 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463575991] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463576425] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1463576505] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1463576545] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 81% (15 min average)
[1463576565] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1463576625] SERVICE ALERT: mckennypk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1463577191] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 454 found since last run. oneLevelSearchOps OK. searchOps is too high! 455 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463577310] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463577761] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463577881] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1463577910] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 281 found since last run. oneLevelSearchOps OK. searchOps is too high! 285 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463578001] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463578030] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463578425] SERVICE ALERT: mckennypk;Device Status;OK;HARD;3;Warming up
[1463578579] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1463578601] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463578609] HOST ALERT: mckennypk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1463578639] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1463578643] HOST ALERT: mckennypk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.61 ms
[1463578699] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1463578729] HOST ALERT: mckennypk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1463578793] HOST ALERT: mckennypk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.60 ms
[1463578793] HOST FLAPPING ALERT: mckennypk;STARTED; Host appears to have started flapping (23.8% change > 20.0% threshold)
[1463579293] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 97% - OK!
[1463579944] HOST DOWNTIME ALERT: epprogsql2;STOPPED; Host has exited from a period of scheduled downtime
[1463580574] SERVICE ALERT: sillpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1463580633] SERVICE ALERT: sillpk;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! All Trays Still Full
[1463581093] HOST FLAPPING ALERT: mckennypk;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1463581479] HOST DOWNTIME ALERT: GOVGIS;STOPPED; Host has exited from a period of scheduled downtime
[1463581601] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (15 min average)
[1463581987] HOST DOWNTIME ALERT: ADAUTHDC1;STOPPED; Host has exited from a period of scheduled downtime
[1463582201] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463582321] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1463582441] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463582654] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1463582714] SERVICE ALERT: studentcenterpk;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! All Trays Still Full
[1463583237] SERVICE ALERT: GOVGIS;C:\ Drive Space;OK;HARD;3;c:\ - total: 174.66 Gb - used: 35.20 Gb (20%) - free 139.45 Gb (80%)
[1463583237] SERVICE ALERT: GOVGIS;VMWare Tools;OK;HARD;3;VMTools: Started
[1463583241] HOST ALERT: GOVGIS;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1463583365] SERVICE ALERT: GOVGIS;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1463583401] SERVICE ALERT: GOVGIS;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1463583423] SERVICE ALERT: GOVGIS;Uptime;OK;HARD;3;System Uptime - 4 day(s) 21 hour(s) 5 minute(s)
[1463583451] SERVICE ALERT: GOVGIS;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1463583451] SERVICE ALERT: GOVGIS;Service: Windows Time;OK;HARD;3;W32Time: Started
[1463583453] SERVICE ALERT: GOVGIS;CPU Load;OK;HARD;3;CPU Load 0% (15 min average)
[1463583536] SERVICE ALERT: GOVGIS;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1463583545] SERVICE ALERT: GOVGIS;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1463583609] SERVICE ALERT: GOVGIS;Service: RPC;OK;HARD;3;RpcSs: Started
[1463583639] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1463583663] SERVICE ALERT: GOVGIS;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1463583699] SERVICE ALERT: GOVGIS;Service: WMI;OK;HARD;3;Winmgmt: Started
[1463583723] SERVICE ALERT: GOVGIS;Memory Usage;OK;HARD;3;Memory usage: total:10239.64 MB - used: 5291.48 MB (52%) - free: 4948.16 MB (48%)
[1463584240] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463584359] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463584477] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463586903] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1463586963] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1463587023] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1463587030] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 494 found since last run. oneLevelSearchOps OK. searchOps is too high! 495 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463587150] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463587390] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1463587450] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1463587459] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1463587510] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1463587519] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1463587579] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1463587579] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1463587639] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1463587699] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1463588033] HOST ALERT: hallefirst2;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1463588064] HOST ALERT: backup;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1463588077] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463588110] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1463588113] HOST ALERT: hallefirst2;UP;SOFT;2;PING OK - Packet loss = 76%, RTA = 0.52 ms
[1463588127] HOST ALERT: backup;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.39 ms
[1463588144] HOST ALERT: redbox;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1463588173] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Warming up
[1463588208] HOST ALERT: redbox;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1463588293] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 41% - OK!
[1463588677] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463588704] SERVICE ALERT: owen101lab1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1463588950] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 991 found since last run. oneLevelSearchOps OK. searchOps is too high! 992 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 960 found since last run.
[1463589070] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 217 found since last run. oneLevelSearchOps OK. searchOps is too high! 218 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463589190] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463590991] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 139872 found since last run. oneLevelSearchOps OK. searchOps is too high! 139873 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463591077] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463591110] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 28566 found since last run. oneLevelSearchOps OK. searchOps is too high! 28567 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463591233] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 29919 found since last run. oneLevelSearchOps OK. searchOps is too high! 29920 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463592277] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463592877] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463593030] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463593401] SERVICE ALERT: idm1.emich.edu;Loopback Driver;CRITICAL;SOFT;1;CRITICAL: cn=Generic Loopback,cn=driverset1,o=system is stopped!
[1463593477] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463593533] SERVICE ALERT: idm1.emich.edu;Loopback Driver;CRITICAL;SOFT;2;CRITICAL: cn=Generic Loopback,cn=driverset1,o=system is stopped!
[1463593639] SERVICE ALERT: idm1.emich.edu;Loopback Driver;CRITICAL;HARD;3;CRITICAL: cn=Generic Loopback,cn=driverset1,o=system is stopped!
[1463595277] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 66% (15 min average)
[1463595539] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1463595599] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1463595659] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1463595733] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1463595743] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1463595793] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1463595849] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1463595853] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1463595863] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 71% (15 min average)
[1463595877] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463595909] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1463595969] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1463595998] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463596117] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463596717] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1463596854] SERVICE ALERT: hallefirst1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1463597047] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;Ready
[1463597163] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 49% - OK!
[1463597317] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1463597437] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463597557] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463598157] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463598757] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463598863] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1463598983] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1463599103] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 87% (15 min average)
[1463599357] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 62% (15 min average)
[1463599957] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1463600077] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463600197] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463600797] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463601135] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 1 Trays Empty
[1463601194] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 1 Trays Empty
[1463601254] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 1 Trays Empty
[1463601397] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463603023] SERVICE ALERT: VC;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1463603143] SERVICE ALERT: VC;CPU Load;OK;SOFT;2;CPU Load 2% (15 min average)
[1463609195] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1463609725] HOST ALERT: hallesecondsouth;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1463609795] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463609815] HOST ALERT: hallesecondsouth;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1463609869] SERVICE ALERT: hallesecondsouth;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1463609899] HOST ALERT: hallesecondsouth;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1463609915] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463609919] SERVICE ALERT: hallesecondsouth;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1463609929] SERVICE ALERT: hallesecondsouth;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1463609949] HOST ALERT: hallesecondsouth;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1463609979] SERVICE ALERT: hallesecondsouth;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1463609989] SERVICE ALERT: hallesecondsouth;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1463610009] HOST ALERT: hallesecondsouth;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1463610035] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463610039] SERVICE ALERT: hallesecondsouth;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1463610069] HOST ALERT: hallesecondsouth;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1463610156] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1463610159] HOST ALERT: hallesecondsouth;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1463610216] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1463610246] HOST ALERT: hallesecondsouth;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1463610276] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1463610306] HOST ALERT: hallesecondsouth;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1463610396] HOST ALERT: hallesecondsouth;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1463610635] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463611235] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463611835] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 82% (15 min average)
[1463612435] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1463612555] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463612675] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463613275] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463615075] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1463615654] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1463615675] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463615795] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463615915] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1463622514] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 76% (15 min average)
[1463623114] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1463623234] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1463623354] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463623954] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463624183] SERVICE ALERT: TI22;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463624283] SERVICE ALERT: TI22;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1463624553] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1463624626] SERVICE ALERT: VC;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463624726] SERVICE ALERT: VC;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1463625000] HOST DOWNTIME ALERT: ADAUTH-RMAD;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: MBAMDB;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: PNCAPP-2012;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: PNCDB-2012;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: ADDC33;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: ADAUTHDC1-OLD;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: EMU-VALUE01;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: UHSPHARM-2012;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: HDBACKUPS;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: R25-ALPHA2012;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: ADCA2;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: ADDC22;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: ADDC3;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: ADDC11;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: TABLEAU;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: ADCA1;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: OPIE;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: SP2010;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: EMUSCRIPT;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: R25-ZETA2012;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: R25-WEBSERVICES;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: DIRECTACCESS;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: R25-DELTA2012;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: WINKMS;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: TI22-PSYCLINIC;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: CBORD;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: NESSUS;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: REFWIN;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: ADWSUSCLIENTS;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: NETENG-PA;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: AD-RMAD;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: PROCARE;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: OFFICEKMS;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: XDM;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: TI22;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: FTK;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: R25-WEBVIEWER;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: ADAUTHDC1;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: ADDC55;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: LANSWEEPER;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: IMAGES;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: DPS-SECURE02;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: EMUSQLEXPRESS;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: TI22-COE;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: ADDC2;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: CEDARON;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: PNCWEB-2012;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: SCCM-PRIM01;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: MBAM;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: SANDBOX-DB;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: COE-WEBDAV;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: R25-INTERFACE;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: ADDC4;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: SP2010DB;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: Echo-01;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: ADWU1;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: ADARC;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: ADDC5;STARTED; Host has entered a period of scheduled downtime
[1463625000] HOST DOWNTIME ALERT: ADDC1;STARTED; Host has entered a period of scheduled downtime
[1463625153] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463625273] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1463625393] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)