[1458101944] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1458102544] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 83% (5 min average)
[1458103144] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1458103264] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1458103384] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 71% (5 min average)
[1458103984] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1458104104] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1458104224] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1458104344] SERVICE ALERT: SP2010;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458104464] SERVICE ALERT: SP2010;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1458107186] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1458107306] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 85% (5 min average)
[1458108387] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 84.90 Gb - used: 78.70 Gb (93%) - free 6.20 Gb (7%)
[1458108506] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 91% (5 min average)
[1458108507] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 84.90 Gb - used: 78.70 Gb (93%) - free 6.20 Gb (7%)
[1458108625] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 64% (5 min average)
[1458108627] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 84.90 Gb - used: 78.70 Gb (93%) - free 6.20 Gb (7%)
[1458109825] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1458109945] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 83% (5 min average)
[1458111881] SERVICE ALERT: halleground1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1458111885] HOST ALERT: halleground1;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.54 ms
[1458111991] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 44% - OK!
[1458112024] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (5 min average)
[1458112295] SERVICE ALERT: halleground1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1458112624] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1458112744] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1458112864] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1458112864] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1458115945] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1458116065] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 53% (5 min average)
[1458116464] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 59% (5 min average)
[1458117064] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458117184] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 97% (5 min average)
[1458117304] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 92% (5 min average)
[1458117904] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1458119104] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 53% (5 min average)
[1458119704] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458119824] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1458119944] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1458121144] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 94% (5 min average)
[1458121744] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 70% (5 min average)
[1458122111] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;1;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1458122231] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;2;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1458122249] HOST ALERT: hallefirst1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1458122317] HOST ALERT: hallefirst1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1458122351] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;HARD;3;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1458122407] HOST ALERT: hallefirst1;DOWN;SOFT;3;(Host check timed out after 30.00 seconds)
[1458122497] HOST ALERT: hallefirst1;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1458122587] HOST ALERT: hallefirst1;DOWN;SOFT;5;(Host check timed out after 30.00 seconds)
[1458122676] HOST ALERT: hallefirst1;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1458122757] HOST ALERT: hallefirst1;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1458122847] HOST ALERT: hallefirst1;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1458122917] HOST ALERT: hallefirst1;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1458122944] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1458122981] HOST ALERT: hallefirst1;UP;SOFT;10;PING OK - Packet loss = 0%, RTA = 0.55 ms
[1458123064] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1458123184] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1458123311] HOST ALERT: hallefirst1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1458123357] HOST ALERT: hallefirst1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1458123447] HOST ALERT: hallefirst1;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1458123517] HOST ALERT: hallefirst1;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1458123607] HOST ALERT: hallefirst1;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1458123697] HOST ALERT: hallefirst1;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1458123786] HOST ALERT: hallefirst1;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1458123875] HOST ALERT: hallefirst1;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1458123957] HOST ALERT: hallefirst1;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1458124047] HOST ALERT: hallefirst1;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1458127384] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 52% (5 min average)
[1458128338] SERVICE ALERT: hallefirst3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1458128386] SERVICE ALERT: hallefirst3;Device Status;OK;HARD;3;TRAY 2 EMPTY PLAIN LETTER   Processing job from tray 4
[1458128584] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458128704] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 95% (5 min average)
[1458128824] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1458128923] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 36% - OK!
[1458129413] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;Warming up
[1458129417] HOST ALERT: hallefirst1;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.56 ms
[1458129551] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;OK;HARD;3;OK: No entries found
[1458129820] HOST ALERT: hallefirst1;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1458129910] HOST ALERT: hallefirst1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1458129949] HOST ALERT: hallefirst1;UP;SOFT;3;PING OK - Packet loss = 77%, RTA = 0.66 ms
[1458130019] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1458130079] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1458130139] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1458130791] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;MENUS  RETRIEVE JOB  INFORMATION  PAPER HANDLING
[1458131224] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 92% (5 min average)
[1458131285] HOST DOWNTIME ALERT: DIRECTACCESS;STARTED; Host has entered a period of scheduled downtime
[1458131318] HOST DOWNTIME ALERT: MAGICAPPTEST;STARTED; Host has entered a period of scheduled downtime
[1458131339] HOST DOWNTIME ALERT: MAGICDBPROD;STARTED; Host has entered a period of scheduled downtime
[1458131391] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1458131451] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1458131512] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1458131824] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 78% (5 min average)
[1458132424] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458132516] HOST ALERT: hallefirst1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1458132544] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 91% (5 min average)
[1458132544] HOST ALERT: hallefirst1;UP;SOFT;2;PING OK - Packet loss = 16%, RTA = 0.71 ms
[1458132664] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1458132921] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 33% - OK!
[1458133080] SERVICE ALERT: hallefirst1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1458133134] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;Ready
[1458133264] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 74% (5 min average)
[1458133571] SERVICE ALERT: strongpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1458133864] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458133885] SERVICE ALERT: ADAUTH-RMAD;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1458133913] SERVICE ALERT: ADWSUSSERVERS;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458133984] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1458134005] SERVICE ALERT: ADAUTH-RMAD;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1458134013] SERVICE ALERT: ADWSUSSERVERS;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1458134104] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 96% (5 min average)
[1458135025] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;OK;HARD;3;c:\ - total: 84.90 Gb - used: 75.50 Gb (89%) - free 9.40 Gb (11%)
[1458135904] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 75% (5 min average)
[1458136802] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (42.5% change >= 20.0% threshold)
[1458136802] HOST DOWNTIME ALERT: DIRECTACCESS;STARTED; Host has entered a period of scheduled downtime
[1458136802] HOST DOWNTIME ALERT: MAGICAPPTEST;STARTED; Host has entered a period of scheduled downtime
[1458136802] HOST DOWNTIME ALERT: MAGICDBPROD;STARTED; Host has entered a period of scheduled downtime
[1458137704] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458137824] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 97% (5 min average)
[1458137944] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1458138112] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Processing job from tray 4
[1458138372] SERVICE ALERT: strongpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1458138431] SERVICE ALERT: strongpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1458138492] SERVICE ALERT: strongpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1458139691] SERVICE ALERT: strongpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1458141544] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 78% (5 min average)
[1458142144] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1458142264] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 89% (5 min average)
[1458142864] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1458142984] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1458143104] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 94% (5 min average)
[1458143704] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1458145724] SERVICE ALERT: emu-fndn-fe7;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458145754] HOST ALERT: emu-fndn-fe7;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1458145771] SERVICE ALERT: emu-fndn-fe7;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458145797] SERVICE ALERT: emu-fndn-fe7;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458145801] HOST ALERT: emu-fndn-fe7;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1458145816] SERVICE ALERT: emu-fndn-fe7;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458145844] SERVICE ALERT: emu-fndn-fe7;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458145846] HOST ALERT: emu-fndn-fe7;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1458145891] SERVICE ALERT: emu-fndn-fe7;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458145917] SERVICE ALERT: emu-fndn-fe7;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458145921] HOST ALERT: emu-fndn-fe7;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1458145936] SERVICE ALERT: emu-fndn-fe7;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458145951] SERVICE ALERT: emu-fndn-fe7;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458145964] SERVICE ALERT: emu-fndn-fe7;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458145966] HOST ALERT: emu-fndn-fe7;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1458145993] SERVICE ALERT: emu-fndn-fe7;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458145995] SERVICE ALERT: emu-fndn-fe7;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458146011] SERVICE ALERT: emu-fndn-fe7;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458146023] HOST ALERT: emu-fndn-fe7;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1458146037] SERVICE ALERT: emu-fndn-fe7;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458146041] SERVICE ALERT: emu-fndn-fe7;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458146056] SERVICE ALERT: emu-fndn-fe7;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458146067] HOST ALERT: emu-fndn-fe7;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1458146071] SERVICE ALERT: emu-fndn-fe7;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458146089] SERVICE ALERT: emu-fndn-fe7;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458146101] HOST ALERT: emu-fndn-fe7;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1458146113] SERVICE ALERT: emu-fndn-fe7;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458146115] SERVICE ALERT: emu-fndn-fe7;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458146140] SERVICE ALERT: emu-fndn-fe7;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458146143] HOST ALERT: emu-fndn-fe7;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1458146161] SERVICE ALERT: emu-fndn-fe7;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458146182] SERVICE ALERT: emu-fndn-fe7;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458146191] HOST ALERT: emu-fndn-fe7;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1458146191] SERVICE ALERT: emu-fndn-fe7;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458146209] SERVICE ALERT: emu-fndn-fe7;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458146231] SERVICE ALERT: emu-fndn-fe7;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458146233] SERVICE ALERT: emu-fndn-fe7;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458146235] SERVICE ALERT: emu-fndn-fe7;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458146260] SERVICE ALERT: emu-fndn-fe7;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458146274] SERVICE ALERT: emu-fndn-fe7;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458146281] SERVICE ALERT: emu-fndn-fe7;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458146302] SERVICE ALERT: emu-fndn-fe7;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458146329] SERVICE ALERT: emu-fndn-fe7;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458146350] SERVICE ALERT: emu-fndn-fe7;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458146380] SERVICE ALERT: emu-fndn-fe7;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458146394] SERVICE ALERT: emu-fndn-fe7;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458146422] SERVICE ALERT: emu-fndn-fe7;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458146470] SERVICE ALERT: emu-fndn-fe7;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458146514] SERVICE ALERT: emu-fndn-fe7;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458146518] HOST ALERT: emu-fndn-fe7;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.49 ms
[1458146544] SERVICE ALERT: emu-fndn-fe7;Service: RPC;OK;HARD;3;RpcSs: Started
[1458146591] SERVICE ALERT: emu-fndn-fe7;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1458146635] SERVICE ALERT: emu-fndn-fe7;Service: WMI;OK;HARD;3;Winmgmt: Started
[1458146771] SERVICE ALERT: emu-fndn-fe7;C:\ Drive Space;OK;HARD;3;c:\ - total: 698.53 Gb - used: 42.90 Gb (6%) - free 655.63 Gb (94%)
[1458146813] SERVICE ALERT: emu-fndn-fe7;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 5 minute(s)
[1458146815] SERVICE ALERT: emu-fndn-fe7;CPU Load;OK;HARD;3;CPU Load 18% (5 min average)
[1458146860] SERVICE ALERT: emu-fndn-fe7;Memory Usage;OK;HARD;3;Memory usage: total:18781.43 MB - used: 2000.05 MB (11%) - free: 16781.38 MB (89%)
[1458146909] SERVICE ALERT: emu-fndn-fe7;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1458146960] SERVICE ALERT: emu-fndn-fe7;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1458147002] SERVICE ALERT: emu-fndn-fe7;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1458147050] SERVICE ALERT: emu-fndn-fe7;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1458147094] SERVICE ALERT: emu-fndn-fe7;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1458147217] SERVICE ALERT: emu-fndn-fe7;Service: Windows Time;OK;HARD;3;W32Time: Started
[1458152464] SERVICE ALERT: owen101lab3;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1458152524] SERVICE ALERT: owen101lab3;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1458152584] SERVICE ALERT: owen101lab3;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1458153110] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1458153170] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1458153230] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1458153518] HOST ALERT: owen101lab3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1458153593] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1458153608] HOST ALERT: owen101lab3;DOWN;SOFT;2;(Host check timed out after 30.00 seconds)
[1458153653] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1458153683] HOST ALERT: owen101lab3;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1458153713] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1458153743] HOST ALERT: owen101lab3;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1458153791] SERVICE ALERT: owen101lab3;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1458153821] HOST ALERT: owen101lab3;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1458153831] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Processing job from tray 4
[1458153910] HOST ALERT: owen101lab3;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1458153946] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1458153976] HOST ALERT: owen101lab3;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1458154006] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1458154036] HOST ALERT: owen101lab3;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1458154066] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1458154096] HOST ALERT: owen101lab3;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1458154186] HOST ALERT: owen101lab3;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1458154430] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1458154490] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1458154550] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1458155103] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1458155150] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Ready
[1458155150] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1458156951] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1458157010] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1458157070] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1458157670] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Warming up
[1458158196] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1458158226] HOST ALERT: halleground1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1458158251] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 44% - OK!
[1458158255] HOST ALERT: halleground1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.81 ms
[1458158270] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1458158330] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1458158391] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1458158703] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 51% (5 min average)
[1458158990] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Ready
[1458160463] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1458160503] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458160583] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 92% (5 min average)
[1458160623] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1458160650] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1458160703] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1458160710] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1458160743] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 85% (5 min average)
[1458160770] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1458160894] SERVICE ALERT: halleground1;Device Status;OK;HARD;3;Clearing paper path
[1458161303] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 63% (5 min average)
[1458161343] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458161370] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 44% - OK!
[1458161463] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 93% (5 min average)
[1458161583] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 89% (5 min average)
[1458161594] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1458161714] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1458161834] SERVICE ALERT: BOEAPPTEST;CPU Load;OK;SOFT;3;CPU Load 3% (5 min average)
[1458161876] SERVICE ALERT: DB01-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1458161996] SERVICE ALERT: DB01-2012;CPU Load;OK;SOFT;2;CPU Load 43% (5 min average)
[1458162183] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1458162303] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 95% (5 min average)
[1458162423] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 97% (5 min average)
[1458162591] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1458162650] SERVICE ALERT: hallefirst2;Device Status;OK;SOFT;2;Processing job from tray 4
[1458163023] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 82% (5 min average)
[1458163250] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1458163310] SERVICE ALERT: hallefirst2;Device Status;OK;SOFT;2;Clearing paper path
[1458163623] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1458163743] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1458163863] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1458163863] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1458163910] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1458163970] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1458164030] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1458164933] HOST ALERT: hallefirstcolor;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1458164955] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1458164985] HOST ALERT: hallefirstcolor;UP;SOFT;2;PING WARNING - Packet loss = 96%, RTA = 0.62 ms
[1458164987] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Cyan Cartridge HP CE261A'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1458164989] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Cartridge HP CE260X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1458165010] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;OK;SOFT;2;Magenta Cartridge HP CE263A is at 40% - OK!
[1458165047] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;OK;SOFT;2;Cyan Cartridge HP CE261A is at 76% - OK!
[1458165049] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;OK;SOFT;2;Black Cartridge HP CE260X is at 48% - OK!
[1458166429] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Ready
[1458168229] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1458168290] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1458168349] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1458171348] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Ready
[1458171665] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1458171690] HOST ALERT: hallefirst2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1458171725] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1458171755] HOST ALERT: hallefirst2;UP;SOFT;2;PING WARNING - Packet loss = 93%, RTA = 0.55 ms
[1458171779] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;SOFT;3;Black Print Cartridge HP C8543X is at 29% - OK!
[1458175263] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1458175863] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 81% (5 min average)
[1458176463] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1458176583] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1458176703] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 88% (5 min average)
[1458176814] SERVICE ALERT: INTLTESTDB;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1458176934] SERVICE ALERT: INTLTESTDB;CPU Load;OK;SOFT;2;CPU Load 28% (5 min average)
[1458177303] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458177423] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1458177493] SERVICE ALERT: SANDBOX-DB;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1458177543] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1458177613] SERVICE ALERT: SANDBOX-DB;CPU Load;WARNING;SOFT;2;CPU Load 94% (5 min average)
[1458177733] SERVICE ALERT: SANDBOX-DB;CPU Load;OK;SOFT;3;CPU Load 0% (5 min average)
[1458180348] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1458181182] SERVICE ALERT: ADWSUSSERVERS;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458181185] HOST ALERT: ADWSUSSERVERS;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.133)
[1458181190] SERVICE ALERT: ADWSUSSERVERS;E:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.102.133 and port 12489: No route to host
[1458181190] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10032;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1458181190] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 12072;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1458181211] SERVICE ALERT: ADWSUSSERVERS;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.102.133 and port 12489: No route to host
[1458181211] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 364;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1458181214] HOST ALERT: ADWSUSSERVERS;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.133)
[1458181278] HOST ALERT: ADWSUSSERVERS;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.39 ms
[1458181282] SERVICE ALERT: ADWSUSSERVERS;NSClient++ Version;OK;SOFT;2;NSClient++ 0,4,1,73 2012-12-17
[1458181287] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 12072;OK;SOFT;2;Eventlog check ok
[1458181307] SERVICE ALERT: ADWSUSSERVERS;E:\ Drive Space;OK;SOFT;2;e:\ - total: 99.87 Gb - used: 59.98 Gb (60%) - free 39.89 Gb (40%)
[1458181307] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10032;OK;SOFT;2;Eventlog check ok
[1458181328] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 364;OK;SOFT;2;Eventlog check ok
[1458181329] SERVICE ALERT: ADWSUSSERVERS;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1458181800] HOST DOWNTIME ALERT: SANDBOX-DB;STARTED; Host has entered a period of scheduled downtime
[1458181800] HOST DOWNTIME ALERT: CBORD;STARTED; Host has entered a period of scheduled downtime
[1458181800] HOST DOWNTIME ALERT: Echo-01;STARTED; Host has entered a period of scheduled downtime
[1458181800] HOST DOWNTIME ALERT: TABLEAU;STARTED; Host has entered a period of scheduled downtime
[1458181800] HOST DOWNTIME ALERT: TI22;STARTED; Host has entered a period of scheduled downtime
[1458181800] HOST DOWNTIME ALERT: FTK;STARTED; Host has entered a period of scheduled downtime
[1458181800] HOST DOWNTIME ALERT: TI22-PSYCLINIC;STARTED; Host has entered a period of scheduled downtime
[1458181800] HOST DOWNTIME ALERT: UHSPHARM-2012;STARTED; Host has entered a period of scheduled downtime
[1458181800] HOST DOWNTIME ALERT: AD-RMAD;STARTED; Host has entered a period of scheduled downtime
[1458181800] HOST DOWNTIME ALERT: HDBACKUPS;STARTED; Host has entered a period of scheduled downtime
[1458181800] HOST DOWNTIME ALERT: WINKMS;STARTED; Host has entered a period of scheduled downtime
[1458181800] HOST DOWNTIME ALERT: XDM;STARTED; Host has entered a period of scheduled downtime
[1458181800] HOST DOWNTIME ALERT: CEDARON;STARTED; Host has entered a period of scheduled downtime
[1458181800] HOST DOWNTIME ALERT: IMAGES;STARTED; Host has entered a period of scheduled downtime
[1458181800] HOST DOWNTIME ALERT: LANSWEEPER;STARTED; Host has entered a period of scheduled downtime
[1458181800] HOST DOWNTIME ALERT: ADAUTH-RMAD;STARTED; Host has entered a period of scheduled downtime
[1458181800] HOST DOWNTIME ALERT: COE-WEBDAV;STARTED; Host has entered a period of scheduled downtime
[1458181800] HOST DOWNTIME ALERT: MBAM;STARTED; Host has entered a period of scheduled downtime
[1458181800] HOST DOWNTIME ALERT: MBAMDB;STARTED; Host has entered a period of scheduled downtime
[1458181800] HOST DOWNTIME ALERT: NESSUS;STARTED; Host has entered a period of scheduled downtime
[1458181800] HOST DOWNTIME ALERT: NETENG-PA;STARTED; Host has entered a period of scheduled downtime
[1458181800] HOST DOWNTIME ALERT: ADWSUSCLIENTS;STARTED; Host has entered a period of scheduled downtime
[1458181800] HOST DOWNTIME ALERT: DPS-SECURE02;STARTED; Host has entered a period of scheduled downtime
[1458181800] HOST DOWNTIME ALERT: OFFICEKMS;STARTED; Host has entered a period of scheduled downtime
[1458181800] HOST DOWNTIME ALERT: OPIE;STARTED; Host has entered a period of scheduled downtime
[1458181800] HOST DOWNTIME ALERT: EMU-VALUE01;STARTED; Host has entered a period of scheduled downtime
[1458181800] HOST DOWNTIME ALERT: PNCAPP-2012;STARTED; Host has entered a period of scheduled downtime
[1458181800] HOST DOWNTIME ALERT: PNCDB-2012;STARTED; Host has entered a period of scheduled downtime
[1458181800] HOST DOWNTIME ALERT: ADWU1;STARTED; Host has entered a period of scheduled downtime
[1458181800] HOST DOWNTIME ALERT: EMUSCRIPT;STARTED; Host has entered a period of scheduled downtime
[1458181800] HOST DOWNTIME ALERT: PNCWEB-2012;STARTED; Host has entered a period of scheduled downtime
[1458181800] HOST DOWNTIME ALERT: PROCARE;STARTED; Host has entered a period of scheduled downtime
[1458181800] HOST DOWNTIME ALERT: EMUSQLEXPRESS;STARTED; Host has entered a period of scheduled downtime
[1458181800] HOST DOWNTIME ALERT: REFWIN;STARTED; Host has entered a period of scheduled downtime
[1458183502] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 91% (5 min average)
[1458183622] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 91% (5 min average)
[1458183711] SERVICE ALERT: MBAM;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1458183733] SERVICE ALERT: SANDBOX-DB;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1458183742] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 95% (5 min average)
[1458183765] SERVICE ALERT: REFWIN;Service: User Profile;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1458183771] SERVICE ALERT: XDM;Service: DNS Client;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1458183830] SERVICE ALERT: MBAM;CPU Load;OK;SOFT;2;CPU Load 14% (5 min average)
[1458183834] SERVICE ALERT: MBAMDB;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458183853] SERVICE ALERT: SANDBOX-DB;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1458183856] SERVICE ALERT: FTK;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.102.240 and port 12489: No route to host
[1458183859] HOST ALERT: FTK;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.240)
[1458183885] SERVICE ALERT: REFWIN;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1458183891] SERVICE ALERT: XDM;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1458183895] SERVICE ALERT: FTK;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.102.240 and port 12489: No route to host
[1458183898] HOST ALERT: FTK;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.240)
[1458183905] SERVICE ALERT: FTK;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.102.240 and port 12489: No route to host
[1458183934] SERVICE ALERT: MBAMDB;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 1 minute(s)
[1458183942] SERVICE ALERT: FTK;Uptime;CRITICAL;SOFT;1;connect to address 164.76.102.240 and port 12489: No route to host
[1458183945] HOST ALERT: FTK;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.102.240)
[1458183967] SERVICE ALERT: FTK;Service: WMI;CRITICAL;SOFT;2;connect to address 164.76.102.240 and port 12489: No route to host
[1458183970] HOST ALERT: FTK;DOWN;SOFT;4;CRITICAL - Host Unreachable (164.76.102.240)
[1458183987] SERVICE ALERT: FTK;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.102.240 and port 12489: No route to host
[1458183990] HOST ALERT: FTK;DOWN;SOFT;5;CRITICAL - Host Unreachable (164.76.102.240)
[1458184015] SERVICE ALERT: FTK;Service: Windows Time;CRITICAL;SOFT;2;connect to address 164.76.102.240 and port 12489: No route to host
[1458184018] HOST ALERT: FTK;DOWN;SOFT;6;CRITICAL - Host Unreachable (164.76.102.240)
[1458184023] SERVICE ALERT: FTK;Memory Usage;OK;SOFT;2;Memory usage: total:65524.64 MB - used: 19127.96 MB (29%) - free: 46396.69 MB (71%)
[1458184059] SERVICE ALERT: FTK;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 1 minute(s)
[1458184063] HOST ALERT: FTK;UP;SOFT;7;PING OK - Packet loss = 0%, RTA = 0.24 ms
[1458184084] SERVICE ALERT: FTK;Service: WMI;OK;SOFT;3;Winmgmt: Started
[1458184104] SERVICE ALERT: FTK;C:\ Drive Space;OK;SOFT;2;c:\ - total: 408.28 Gb - used: 138.79 Gb (34%) - free 269.49 Gb (66%)
[1458184109] SERVICE ALERT: XDM;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1458184131] SERVICE ALERT: FTK;Service: Windows Time;OK;SOFT;3;W32Time: Started
[1458184229] SERVICE ALERT: XDM;CPU Load;OK;SOFT;2;CPU Load 15% (5 min average)
[1458184319] SERVICE ALERT: SANDBOX-DB;Service: Group Policy;CRITICAL;SOFT;1;gpsvc: Stopped
[1458184342] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 85% (5 min average)
[1458184350] SERVICE ALERT: Echo-01;Service: WMI;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1458184362] SERVICE ALERT: PNCAPP-2012;Service: Group Policy;CRITICAL;SOFT;1;gpsvc: Stopped
[1458184362] SERVICE ALERT: PNCWEB-2012;Service: DNS Client;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1458184368] SERVICE ALERT: EMUSQLEXPRESS;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458184368] SERVICE ALERT: EMUSQLEXPRESS;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458184375] HOST ALERT: DIRECTACCESS;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.41)
[1458184376] SERVICE ALERT: Echo-01;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458184382] SERVICE ALERT: SANDBOX-DB;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458184382] SERVICE ALERT: OPIE;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458184388] SERVICE ALERT: UHSPHARM-2012;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458184392] SERVICE ALERT: WINKMS;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458184392] SERVICE ALERT: DPS-SECURE02;Service: Group Policy;CRITICAL;SOFT;1;gpsvc: Stopped
[1458184408] SERVICE ALERT: CBORD;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458184411] SERVICE ALERT: LANSWEEPER;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.2.33 and port 12489: Connection refused
[1458184432] SERVICE ALERT: COE-WEBDAV;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458184433] HOST ALERT: DIRECTACCESS;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.39 ms
[1458184439] SERVICE ALERT: SANDBOX-DB;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1458184454] SERVICE ALERT: SANDBOX-DB;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1458184468] SERVICE ALERT: EMUSQLEXPRESS;VMWare Tools;OK;SOFT;2;VMTools: Started
[1458184468] SERVICE ALERT: EMUSQLEXPRESS;C:\ Drive Space;OK;SOFT;2;c:\ - total: 79.66 Gb - used: 25.35 Gb (32%) - free 54.30 Gb (68%)
[1458184470] SERVICE ALERT: Echo-01;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1458184476] SERVICE ALERT: Echo-01;Service: RPC;OK;SOFT;2;RpcSs: Started
[1458184482] SERVICE ALERT: SANDBOX-DB;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1458184482] SERVICE ALERT: OPIE;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1458184482] SERVICE ALERT: PNCWEB-2012;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1458184482] SERVICE ALERT: PNCAPP-2012;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1458184488] SERVICE ALERT: UHSPHARM-2012;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1458184492] SERVICE ALERT: WINKMS;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1458184508] SERVICE ALERT: CBORD;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1458184522] SERVICE ALERT: LANSWEEPER;Service: RPC;OK;SOFT;2;RpcSs: Started
[1458184532] SERVICE ALERT: COE-WEBDAV;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1458184532] SERVICE ALERT: DPS-SECURE02;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458184552] SERVICE ALERT: DPS-SECURE02;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458184562] HOST ALERT: DPS-SECURE02;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1458184573] SERVICE ALERT: SANDBOX-DB;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1458184602] SERVICE ALERT: DPS-SECURE02;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458184632] HOST ALERT: DPS-SECURE02;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1458184633] SERVICE ALERT: DPS-SECURE02;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458184648] SERVICE ALERT: DPS-SECURE02;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458184649] SERVICE ALERT: EMU-VALUE01;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458184652] SERVICE ALERT: DPS-SECURE02;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1458184663] HOST ALERT: DPS-SECURE02;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1458184672] SERVICE ALERT: DPS-SECURE02;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458184676] HOST ALERT: DPS-SECURE02;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1458184678] SERVICE ALERT: EMU-VALUE01;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458184679] HOST ALERT: EMU-VALUE01;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1458184696] SERVICE ALERT: EMU-VALUE01;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458184702] SERVICE ALERT: DPS-SECURE02;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1458184726] HOST ALERT: EMU-VALUE01;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1458184733] SERVICE ALERT: DPS-SECURE02;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1458184743] SERVICE ALERT: EMU-VALUE01;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458184748] SERVICE ALERT: DPS-SECURE02;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1458184769] SERVICE ALERT: EMU-VALUE01;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458184772] SERVICE ALERT: DPS-SECURE02;Service: User Profile;OK;SOFT;3;ProfSvc: Started
[1458184773] HOST ALERT: EMU-VALUE01;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1458184790] SERVICE ALERT: EMU-VALUE01;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458184798] SERVICE ALERT: EMU-VALUE01;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458184816] SERVICE ALERT: EMU-VALUE01;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458184820] HOST ALERT: EMU-VALUE01;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1458184836] SERVICE ALERT: EMU-VALUE01;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1458184863] SERVICE ALERT: EMU-VALUE01;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1458184866] HOST ALERT: EMU-VALUE01;DOWN;SOFT;5;(Host check timed out after 30.01 seconds)
[1458184879] SERVICE ALERT: EMU-VALUE01;Service: WMI;OK;SOFT;3;Winmgmt: Started
[1458184890] SERVICE ALERT: EMU-VALUE01;C:\ Drive Space;OK;SOFT;2;c:\ - total: 465.42 Gb - used: 15.16 Gb (3%) - free 450.26 Gb (97%)
[1458184898] SERVICE ALERT: EMU-VALUE01;Memory Usage;OK;SOFT;3;Memory usage: total:9444.43 MB - used: 2023.42 MB (21%) - free: 7421.01 MB (79%)
[1458184902] HOST ALERT: EMU-VALUE01;UP;SOFT;6;PING OK - Packet loss = 78%, RTA = 0.48 ms
[1458184916] SERVICE ALERT: EMU-VALUE01;Service: Windows Time;OK;SOFT;3;W32Time: Started
[1458184936] SERVICE ALERT: EMU-VALUE01;CPU Load;OK;SOFT;2;CPU Load 17% (5 min average)
[1458184963] SERVICE ALERT: EMU-VALUE01;Uptime;OK;SOFT;3;System Uptime - 0 day(s) 0 hour(s) 5 minute(s)
[1458185232] SERVICE ALERT: DPS-SECURE02;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1458185343] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 94% (5 min average)
[1458185717] HOST ALERT: hallethirdsouth;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1458185943] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1458185943] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1458186543] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 80% (5 min average)
[1458187143] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)