[1461124825] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461124945] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1461125065] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461128065] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461128665] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461131785] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1461131905] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1461132025] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 53% (15 min average)
[1461132625] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1461132745] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1461132865] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 82% (15 min average)
[1461132865] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461133465] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1461133465] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461133584] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1461133704] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461134304] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1461135865] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461136104] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1461136224] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1461136465] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461136824] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1461136943] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1461137062] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1461137065] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1461137662] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1461137662] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1461137665] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1461137785] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461137905] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461140305] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 83% (15 min average)
[1461140905] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461141025] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1461141145] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 94% (15 min average)
[1461141745] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461141865] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1461141985] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461146185] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 83% (15 min average)
[1461146785] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1461146905] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1461147025] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461148225] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461148825] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461149062] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1461151225] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461151825] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461151876] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Warming up
[1461152476] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461152537] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461152596] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461153025] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1461153196] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Waiting for tray 4 to lift   Processing job from tray 3
[1461154225] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461154345] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461154465] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461155065] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 90% (15 min average)
[1461155403] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (48.0% change >= 20.0% threshold)
[1461155403] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1461155403] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1461155403] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1461155403] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1461155637] SERVICE ALERT: TI22-COE;D:\ Drive Space;CRITICAL;SOFT;1;d:\ - total: 2.96 Gb - used: 2.96 Gb (100%) - free 0.00 Gb (0%)
[1461155665] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1461155757] SERVICE ALERT: TI22-COE;D:\ Drive Space;CRITICAL;SOFT;2;d:\ - total: 2.96 Gb - used: 2.96 Gb (100%) - free 0.00 Gb (0%)
[1461155785] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 91% (15 min average)
[1461155877] SERVICE ALERT: TI22-COE;D:\ Drive Space;CRITICAL;HARD;3;d:\ - total: 2.96 Gb - used: 2.96 Gb (100%) - free 0.00 Gb (0%)
[1461156985] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1461157105] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461157225] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461157252] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors is too high! 1140 found since last run. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1461157372] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1461157666] SERVICE ALERT: TI22-COE;D:\ Drive Space;UNKNOWN;HARD;3;Free disk space : Invalid drive
[1461157825] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1461158425] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461158545] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461158665] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461158853] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1461158912] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1461158973] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1461159002] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (60.4% change >= 20.0% threshold)
[1461159002] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1461159002] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1461159002] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1461159002] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1461160372] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors is too high! 1726 found since last run. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1461160493] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;2;errors is too high! 15950 found since last run. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1461160612] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;HARD;3;errors is too high! 33582 found since last run. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1461161212] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1461161814] SERVICE ALERT: sillpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.126.102. Is the printer off?  Is the network down?
[1461161874] SERVICE ALERT: sillpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.126.102. Is the printer off?  Is the network down?
[1461161934] SERVICE ALERT: sillpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.126.102. Is the printer off?  Is the network down?
[1461162118] SERVICE ALERT: sillpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.126.102. Is the printer off?  Is the network down?
[1461162178] SERVICE ALERT: sillpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.126.102. Is the printer off?  Is the network down?
[1461162238] SERVICE ALERT: sillpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.126.102. Is the printer off?  Is the network down?
[1461162273] SERVICE ALERT: sillpk;Black Cartridge HP CC364X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.126.102. Is the printer off?  Is the network down?
[1461162333] SERVICE ALERT: sillpk;Black Cartridge HP CC364X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.126.102. Is the printer off?  Is the network down?
[1461162393] SERVICE ALERT: sillpk;Black Cartridge HP CC364X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.126.102. Is the printer off?  Is the network down?
[1461163937] SERVICE ALERT: owen101lab1;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1461163997] SERVICE ALERT: owen101lab1;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1461164057] SERVICE ALERT: owen101lab1;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1461165558] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1461165588] HOST ALERT: prayharroldpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1461165618] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1461165648] HOST ALERT: prayharroldpk;UP;SOFT;2;PING WARNING - Packet loss = 93%, RTA = 0.73 ms
[1461165678] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1461165769] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1461165829] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1461165889] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1461165970] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1461166030] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1461166090] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1461166258] SERVICE ALERT: mckennypk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1461166288] HOST ALERT: mckennypk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1461166352] HOST ALERT: mckennypk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.57 ms
[1461166457] SERVICE ALERT: owen101lab1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1461166852] SERVICE ALERT: mckennypk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461166908] HOST ALERT: prayharroldpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1461166960] SERVICE ALERT: halleground1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461166972] HOST ALERT: prayharroldpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1461167019] SERVICE ALERT: halleground1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461167079] SERVICE ALERT: halleground1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461167194] SERVICE ALERT: halleground1;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1461167254] SERVICE ALERT: halleground1;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1461167314] SERVICE ALERT: halleground1;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1461168195] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461168255] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461168315] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461168652] SERVICE ALERT: mckennypk;Device Status;OK;HARD;3;Ready
[1461168915] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Clearing paper path
[1461168945] SERVICE ALERT: hallefirstcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461169004] SERVICE ALERT: hallefirstcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461169064] SERVICE ALERT: hallefirstcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461169515] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461169575] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461169636] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461169664] SERVICE ALERT: hallefirstcolor;Device Status;OK;HARD;3;TONER COLLECTION UNIT FULL 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1461169713] SERVICE ALERT: halleground1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1461170064] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461170079] SERVICE ALERT: halleground1;Device Status;OK;HARD;3;Warming up
[1461170235] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Ready
[1461170235] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1461170664] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461171076] SERVICE ALERT: prayharroldpk;Device Status;OK;HARD;3;Warming up
[1461171283] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 9% - OK!
[1461171485] SERVICE ALERT: prayharroldpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1461171864] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461172224] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461172284] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461172344] SERVICE ALERT: hallefirst1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461172464] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461172944] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;Ready
[1461173403] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1461173403] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1461173403] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1461173403] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1461173403] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1461173403] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1461173544] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461173604] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461173664] SERVICE ALERT: hallefirst1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461174264] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;Ready
[1461174264] SERVICE FLAPPING ALERT: hallefirst1;Device Status;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1461175202] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1461175202] SERVICE FLAPPING ALERT: hallefirst1;Device Status;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1461175202] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1461175203] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1461175203] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1461175203] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1461175203] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1461175467] HOST ALERT: sillpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1461175557] HOST ALERT: sillpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1461175619] HOST ALERT: sillpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1461175709] HOST ALERT: sillpk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1461175764] HOST ALERT: sillpk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1461175780] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1461175840] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1461175854] HOST ALERT: sillpk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1461175900] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1461175944] HOST ALERT: sillpk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1461176034] HOST ALERT: sillpk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1461176067] HOST ALERT: sillpk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1461176157] HOST ALERT: sillpk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1461177100] SERVICE ALERT: studentcenterpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1461177832] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;REPLACE MAINTENANCE KIT 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1461177836] HOST ALERT: sillpk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.41 ms
[1461177862] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1461178019] HOST ALERT: sillpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1461178035] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461178095] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461178109] HOST ALERT: sillpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1461178155] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461178164] HOST ALERT: sillpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1461178254] HOST ALERT: sillpk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1461178343] HOST ALERT: sillpk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1461178433] HOST ALERT: sillpk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1461178437] SERVICE ALERT: sillpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.126.102. Is the printer off?  Is the network down?
[1461178462] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1461178467] HOST ALERT: sillpk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1461178497] SERVICE ALERT: sillpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.126.102. Is the printer off?  Is the network down?
[1461178527] HOST ALERT: sillpk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1461178557] SERVICE ALERT: sillpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.126.102. Is the printer off?  Is the network down?
[1461178582] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461178587] HOST ALERT: sillpk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1461178619] HOST ALERT: sillpk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1461178702] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461179302] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461179356] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Ready
[1461179902] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461181252] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461181312] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461181372] SERVICE ALERT: mckennypk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461181583] SERVICE ALERT: idm1.emich.edu;NSDS Errors;WARNING;SOFT;1;NRPE: Unable to read output
[1461181703] SERVICE ALERT: idm1.emich.edu;NSDS Errors;WARNING;SOFT;2;NRPE: Unable to read output
[1461181770] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1461181823] SERVICE ALERT: idm1.emich.edu;NSDS Errors;OK;SOFT;3;OK: 0 Errors Reported
[1461182302] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461182902] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461184102] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461184202] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (42.4% change >= 20.0% threshold)
[1461184202] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1461184202] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1461184202] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1461184202] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1461184702] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461184770] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1461184831] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1461184890] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1461185302] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461185572] SERVICE ALERT: mckennypk;Device Status;OK;HARD;3;Ready
[1461185663] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461185723] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461185784] SERVICE ALERT: hallefirst1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461185902] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1461186502] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461187102] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1461187155] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461187214] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461187274] SERVICE ALERT: hallefirst2;Device Status;OK;SOFT;3;Clearing paper path
[1461187583] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;Clearing paper path
[1461188183] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461188243] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461188302] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 91% (15 min average)
[1461188303] SERVICE ALERT: hallefirst1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461188902] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1461188909] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1461188909] SERVICE FLAPPING ALERT: hallefirst1;Device Status;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1461188932] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1461188939] HOST ALERT: hallefirst1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1461188987] SERVICE ALERT: hallefirst1;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! All Trays Still Full
[1461188991] HOST ALERT: hallefirst1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.53 ms
[1461189022] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1461189142] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461189503] SERVICE ALERT: hallefirst1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461190942] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461191542] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461192680] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1461192740] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1461192800] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1461193148] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1461193185] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1461193208] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1461193245] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1461193268] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1461193305] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1461194303] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;MENUS  RETRIEVE JOB  INFORMATION  PAPER HANDLING
[1461194542] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461194903] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461194934] HOST ALERT: hallefirst2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1461194963] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461194998] HOST ALERT: hallefirst2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.79 ms
[1461195024] SERVICE ALERT: hallefirst1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461195062] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1461195100] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 26% - OK!
[1461195142] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461195194] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Ready
[1461196222] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;Canceling job
[1461198089] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1461198690] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1461198749] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1461198810] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1461199222] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461199282] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461199343] SERVICE ALERT: hallefirst1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461199942] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;Ready
[1461200542] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1461201142] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461201262] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1461201382] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461202582] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461202942] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1461203002] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1461203062] SERVICE ALERT: hallefirst1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1461203663] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;Ready
[1461204382] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 85% (15 min average)
[1461204982] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1461205102] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1461205222] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461205800] HOST DOWNTIME ALERT: LANSWEEPER;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: ADDC11;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: ADDC55;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: MBAMDB;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: ADDC66;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: NESSUS;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: NETENG-PA;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: ADDC2;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: ADWSUSCLIENTS;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: OFFICEKMS;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: OPIE;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: ADWU1;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: PNCAPP-2012;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: PNCDB-2012;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: ADDC22;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: CBORD;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: PNCWEB-2012;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: AD-RMAD;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: PROCARE;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: CEDARON;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: R25-ALPHA2012;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: R25-DELTA2012;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: ADDC3;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: COE-WEBDAV;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: R25-ZETA2012;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: REFWIN;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: DIRECTACCESS;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: SANDBOX-DB;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: SP2010;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: ADDC33;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: DPS-SECURE02;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: SP2010DB;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: ADARC;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: TABLEAU;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: EMU-VALUE01;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: TI22;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: TI22-COE;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: ADDC4;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: EMUSCRIPT;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: TI22-PSYCLINIC;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: UHSPHARM-2012;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: EMUSQLEXPRESS;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: WINKMS;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: XDM;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: ADDC44;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: MBAM;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: Echo-01;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: FTK;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: ADAUTH-RMAD;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: HDBACKUPS;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: ADDC5;STARTED; Host has entered a period of scheduled downtime
[1461205800] HOST DOWNTIME ALERT: IMAGES;STARTED; Host has entered a period of scheduled downtime
[1461205822] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1461206422] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1461206542] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1461206662] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461207262] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461207728] SERVICE ALERT: OFFICEKMS;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1461207808] SERVICE ALERT: XDM;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Starting
[1461207848] SERVICE ALERT: OFFICEKMS;CPU Load;OK;SOFT;2;CPU Load 2% (15 min average)
[1461207854] SERVICE ALERT: ADCA2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1461207928] SERVICE ALERT: XDM;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1461207942] HOST ALERT: FTK;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1461207963] HOST ALERT: FTK;UP;SOFT;2;PING OK - Packet loss = 58%, RTA = 0.34 ms
[1461207974] SERVICE ALERT: ADCA2;CPU Load;OK;SOFT;2;CPU Load 64% (15 min average)
[1461208239] SERVICE ALERT: ADDC11;C:\ Drive Space;UNKNOWN;SOFT;1;NSClient - ERROR: Invalid return from command: check_drivesize
[1461208272] SERVICE ALERT: ADCA1;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.7.150 and port 12489: Connection refused
[1461208272] SERVICE ALERT: ADCA2;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.7.151 and port 12489: Connection refused
[1461208299] SERVICE ALERT: HDBACKUPS;Service: Group Policy;CRITICAL;SOFT;1;gpsvc: Stopped
[1461208302] SERVICE ALERT: R25-ALPHA2012;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.99.204 and port 12489: Connection refused
[1461208309] SERVICE ALERT: ADDC66;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461208312] SERVICE ALERT: ADWU1;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461208313] SERVICE ALERT: R25-DELTA2012;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461208314] SERVICE ALERT: EMUSCRIPT;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461208320] SERVICE ALERT: OPIE;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461208320] SERVICE ALERT: PNCWEB-2012;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461208322] SERVICE ALERT: OPIE;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461208322] SERVICE ALERT: PNCDB-2012;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461208322] SERVICE ALERT: OFFICEKMS;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461208324] SERVICE ALERT: R25-ZETA2012;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461208346] SERVICE ALERT: DPS-SECURE02;Service: Group Policy;CRITICAL;SOFT;1;gpsvc: Stopped
[1461208355] SERVICE ALERT: CBORD;VMWare Tools;CRITICAL;SOFT;1;VMTools: Stopped
[1461208358] HOST ALERT: CBORD;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1461208359] SERVICE ALERT: ADDC11;C:\ Drive Space;OK;SOFT;2;c:\ - total: 79.66 Gb - used: 33.35 Gb (42%) - free 46.31 Gb (58%)
[1461208372] SERVICE ALERT: HDBACKUPS;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461208392] SERVICE ALERT: ADCA1;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1461208392] SERVICE ALERT: ADCA2;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1461208402] HOST ALERT: HDBACKUPS;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1461208409] SERVICE ALERT: ADDC66;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1461208411] SERVICE ALERT: ADDC4;Service: AD Domain Service;CRITICAL;SOFT;1;connect to address 164.76.99.238 and port 12489: Connection refused
[1461208411] SERVICE ALERT: ADDC4;Service: Remote Management;CRITICAL;SOFT;1;connect to address 164.76.99.238 and port 12489: Connection refused
[1461208412] SERVICE ALERT: HDBACKUPS;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461208412] SERVICE ALERT: ADWU1;Memory Usage;OK;SOFT;2;Memory usage: total:22471.68 MB - used: 1415.18 MB (6%) - free: 21056.50 MB (94%)
[1461208413] SERVICE ALERT: R25-DELTA2012;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1461208414] SERVICE ALERT: EMUSCRIPT;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1461208420] SERVICE ALERT: OPIE;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 2 minute(s)
[1461208420] SERVICE ALERT: PNCWEB-2012;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1461208421] HOST ALERT: CBORD;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1461208422] SERVICE ALERT: OPIE;Memory Usage;OK;SOFT;2;Memory usage: total:10943.64 MB - used: 1096.37 MB (10%) - free: 9847.28 MB (90%)
[1461208422] SERVICE ALERT: R25-ALPHA2012;Service: RPC;OK;SOFT;2;RpcSs: Started
[1461208422] SERVICE ALERT: OFFICEKMS;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1461208422] SERVICE ALERT: PNCDB-2012;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1461208424] SERVICE ALERT: R25-ZETA2012;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1461208430] SERVICE ALERT: DPS-SECURE02;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461208439] SERVICE ALERT: ADDC2;Service: Group Policy;CRITICAL;SOFT;1;gpsvc: Stopped
[1461208439] SERVICE ALERT: HDBACKUPS;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1461208440] SERVICE ALERT: DPS-SECURE02;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461208442] HOST ALERT: HDBACKUPS;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1461208452] SERVICE ALERT: DPS-SECURE02;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461208460] HOST ALERT: DPS-SECURE02;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1461208462] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1461208469] SERVICE ALERT: HDBACKUPS;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461208475] SERVICE ALERT: CBORD;VMWare Tools;OK;SOFT;2;VMTools: Started
[1461208479] SERVICE ALERT: ADDC2;AD Replication;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1461208481] SERVICE ALERT: HDBACKUPS;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1461208482] HOST ALERT: HDBACKUPS;UP;SOFT;3;PING OK - Packet loss = 64%, RTA = 0.26 ms
[1461208483] SERVICE ALERT: DPS-SECURE02;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461208484] SERVICE ALERT: ADDC3;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461208486] SERVICE ALERT: DPS-SECURE02;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1461208486] SERVICE ALERT: DPS-SECURE02;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461208486] SERVICE ALERT: ADDC3;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461208503] HOST ALERT: EMU-VALUE01;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1461208512] SERVICE ALERT: HDBACKUPS;Service: RPC;OK;SOFT;2;RpcSs: Started
[1461208513] HOST ALERT: DPS-SECURE02;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1461208520] SERVICE ALERT: NETENG-PA;Service: Windows Time;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1461208522] SERVICE ALERT: ADDC4;Service: AD Domain Service;OK;SOFT;2;NTDS: Started
[1461208522] SERVICE ALERT: ADDC4;Service: Remote Management;OK;SOFT;2;WinRM: Started
[1461208532] SERVICE ALERT: DPS-SECURE02;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461208539] SERVICE ALERT: HDBACKUPS;Service: Group Policy;OK;SOFT;3;gpsvc: Started
[1461208540] SERVICE ALERT: ADDC4;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1461208550] SERVICE ALERT: DPS-SECURE02;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1461208559] SERVICE ALERT: ADDC2;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1461208560] HOST ALERT: DPS-SECURE02;UP;SOFT;3;PING WARNING - Packet loss = 82%, RTA = 0.34 ms
[1461208560] SERVICE ALERT: DPS-SECURE02;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1461208560] SERVICE ALERT: ADAUTHDC1;Service: AD Web Service;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1461208560] SERVICE ALERT: NETENG-PA;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461208561] SERVICE ALERT: DPS-SECURE02;Service: RPC;OK;SOFT;2;RpcSs: Started
[1461208569] SERVICE ALERT: HDBACKUPS;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1461208582] SERVICE ALERT: ADDC5;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1461208582] SERVICE ALERT: ADDC5;Service: Dell Backup Agent;CRITICAL;SOFT;1;ErdAgent: Starting
[1461208583] SERVICE ALERT: DPS-SECURE02;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1461208584] SERVICE ALERT: ADDC3;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1461208585] SERVICE ALERT: ADAUTHDC1;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.102.171 and port 12489: Connection refused
[1461208586] SERVICE ALERT: DPS-SECURE02;Service: Group Policy;OK;SOFT;3;gpsvc: Started
[1461208586] SERVICE ALERT: ADDC2;AD Replication;OK;SOFT;2;OK: replication is up and running
[1461208586] SERVICE ALERT: DPS-SECURE02;CPU Load;OK;SOFT;2;CPU Load 4% (15 min average)
[1461208586] SERVICE ALERT: ADDC3;Service: RPC;OK;SOFT;2;RpcSs: Started
[1461208593] HOST ALERT: EMU-VALUE01;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1461208599] SERVICE ALERT: EMU-VALUE01;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461208629] HOST ALERT: EMU-VALUE01;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1461208630] SERVICE ALERT: ADAUTHDC2;Service: DFS Namespace;CRITICAL;SOFT;1;connect to address 164.76.102.172 and port 12489: Connection refused
[1461208630] SERVICE ALERT: ADAUTHDC2;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.102.172 and port 12489: Connection refused
[1461208632] SERVICE ALERT: DPS-SECURE02;Memory Usage;OK;SOFT;2;Memory usage: total:4790.25 MB - used: 942.00 MB (20%) - free: 3848.25 MB (80%)
[1461208640] SERVICE ALERT: NETENG-PA;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1461208650] SERVICE ALERT: DPS-SECURE02;Service: WMI;OK;SOFT;3;Winmgmt: Started
[1461208650] SERVICE ALERT: ADAUTHDC2;Service: DNS Service;CRITICAL;SOFT;1;DNS: Starting
[1461208660] SERVICE ALERT: DPS-SECURE02;C:\ Drive Space;OK;SOFT;3;c:\ - total: 278.53 Gb - used: 13.97 Gb (5%) - free 264.56 Gb (95%)
[1461208660] SERVICE ALERT: ADDC4;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1461208660] SERVICE ALERT: NETENG-PA;CPU Load;OK;SOFT;2;CPU Load 3% (15 min average)
[1461208664] SERVICE ALERT: EMU-VALUE01;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461208673] SERVICE ALERT: EMU-VALUE01;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461208674] SERVICE ALERT: ADAUTHDC2;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1461208680] SERVICE ALERT: ADAUTHDC1;Service: AD Web Service;OK;SOFT;2;ADWS: Started
[1461208680] SERVICE ALERT: EMU-VALUE01;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461208683] SERVICE ALERT: EMU-VALUE01;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461208694] HOST ALERT: EMU-VALUE01;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1461208699] SERVICE ALERT: EMU-VALUE01;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1461208702] SERVICE ALERT: ADDC5;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1461208702] SERVICE ALERT: ADAUTHDC1;Service: RPC;OK;SOFT;2;RpcSs: Started
[1461208702] SERVICE ALERT: ADDC5;Service: Dell Backup Agent;OK;SOFT;2;ErdAgent: Started
[1461208703] HOST ALERT: EMU-VALUE01;UP;SOFT;5;PING OK - Packet loss = 0%, RTA = 0.36 ms
[1461208729] SERVICE ALERT: ADAUTHDC1;Service: Dell Backup Agent;CRITICAL;SOFT;1;ErdAgent: Starting
[1461208750] SERVICE ALERT: ADAUTHDC2;Service: DFS Namespace;OK;SOFT;2;Dfs: Started
[1461208750] SERVICE ALERT: ADAUTHDC2;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1461208762] SERVICE ALERT: ADAUTHDC2;Service: Dell Backup Agent;CRITICAL;SOFT;1;ErdAgent: Starting
[1461208764] SERVICE ALERT: EMU-VALUE01;CPU Load;OK;SOFT;2;CPU Load 25% (15 min average)
[1461208770] SERVICE ALERT: ADAUTHDC2;Service: DNS Service;OK;SOFT;2;DNS: Started
[1461208773] SERVICE ALERT: EMU-VALUE01;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1461208780] SERVICE ALERT: EMU-VALUE01;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1461208783] SERVICE ALERT: EMU-VALUE01;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 5 minute(s)
[1461208794] SERVICE ALERT: ADAUTHDC2;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1461208849] SERVICE ALERT: ADAUTHDC1;Service: Dell Backup Agent;OK;SOFT;2;ErdAgent: Started
[1461208882] SERVICE ALERT: ADAUTHDC2;Service: Dell Backup Agent;OK;SOFT;2;ErdAgent: Started
[1461209034] SERVICE ALERT: hallesecondsouth;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1461209043] SERVICE ALERT: ADDC4;AD Replication;WARNING;SOFT;1;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1461209060] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1461209089] SERVICE ALERT: hallesecondsouth;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! All Trays Still Full
[1461209097] SERVICE ALERT: SCCM-PRIM01;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461209100] HOST ALERT: SCCM-PRIM01;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1461209101] SERVICE ALERT: ADAUTHDC1;AD Replication;WARNING;SOFT;1;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1461209140] SERVICE ALERT: SCCM-PRIM01;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461209152] SERVICE ALERT: SCCM-PRIM01;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461209155] SERVICE ALERT: SCCM-PRIM01;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1461209162] SERVICE ALERT: SCCM-PRIM01;Service: NetLogin;CRITICAL;SOFT;1;connect to address 164.76.4.64 and port 12489: Connection refused
[1461209164] SERVICE ALERT: ADDC4;AD Replication;WARNING;SOFT;2;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1461209170] HOST ALERT: SCCM-PRIM01;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1461209180] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1461209192] HOST ALERT: SCCM-PRIM01;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1461209197] SERVICE ALERT: SCCM-PRIM01;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1461209220] SERVICE ALERT: ADAUTHDC1;AD Replication;WARNING;SOFT;2;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1461209240] SERVICE ALERT: SCCM-PRIM01;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 2 minute(s)
[1461209252] SERVICE ALERT: SP2010DB;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.4.123 and port 12489: Connection refused
[1461209252] SERVICE ALERT: SCCM-PRIM01;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1461209255] SERVICE ALERT: SCCM-PRIM01;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1461209273] SERVICE ALERT: SCCM-PRIM01;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1461209283] SERVICE ALERT: ADDC4;AD Replication;WARNING;HARD;3;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1461209300] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 94% (15 min average)
[1461209340] SERVICE ALERT: ADAUTHDC1;AD Replication;WARNING;HARD;3;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1461209372] SERVICE ALERT: SP2010DB;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1461209500] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1461209559] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1461209620] SERVICE ALERT: studentcenterpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1461209670] SERVICE ALERT: SP2010;Service: Firewall;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1461209790] SERVICE ALERT: SP2010;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1461209900] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1461210020] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1461210140] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 92% (15 min average)
[1461210262] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1461210544] SERVICE ALERT: ADAUTHDC1;AD Replication;OK;HARD;3;OK: replication is up and running
[1461210740] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1461210860] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1461210981] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)