[1462336078] SERVICE ALERT: ADWU1;Windows Update Services Event 10022;CRITICAL;SOFT;1;Windows Server Update Services, The last catalog synchronization attempt was unsuccessful.
[1462336198] SERVICE ALERT: ADWU1;Windows Update Services Event 10022;CRITICAL;SOFT;2;Windows Server Update Services, The last catalog synchronization attempt was unsuccessful.
[1462336318] SERVICE ALERT: ADWU1;Windows Update Services Event 10022;CRITICAL;HARD;3;Windows Server Update Services, The last catalog synchronization attempt was unsuccessful.
[1462337973] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 80% (15 min average)
[1462338573] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462338693] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1462338813] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 67% (15 min average)
[1462339020] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10022;CRITICAL;SOFT;1;Windows Server Update Services(error, 10022, success)[The last catalog synchronization attempt was unsuccessful., ], eventlog: 1 > critical
[1462339140] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10022;CRITICAL;SOFT;2;Windows Server Update Services(error, 10022, success)[The last catalog synchronization attempt was unsuccessful., ], eventlog: 1 > critical
[1462339261] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10022;CRITICAL;HARD;3;Windows Server Update Services, eventlog: 1 > critical
[1462339413] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462339533] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 74% (15 min average)
[1462340133] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462340253] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462340373] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462340963] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462341083] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1462341203] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462341933] SERVICE ALERT: ADCA2;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462342053] SERVICE ALERT: ADCA2;CPU Load;OK;SOFT;2;CPU Load 1% (15 min average)
[1462342918] SERVICE ALERT: ADWU1;Windows Update Services Event 10022;OK;HARD;3;OK: No entries found
[1462344803] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462345318] SERVICE ALERT: ADWU1;Windows Update Services Event 10022;CRITICAL;SOFT;1;Windows Server Update Services, The last catalog synchronization attempt was unsuccessful.
[1462345403] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1462345438] SERVICE ALERT: ADWU1;Windows Update Services Event 10022;CRITICAL;SOFT;2;Windows Server Update Services, The last catalog synchronization attempt was unsuccessful.
[1462345558] SERVICE ALERT: ADWU1;Windows Update Services Event 10022;CRITICAL;HARD;3;Windows Server Update Services, The last catalog synchronization attempt was unsuccessful.
[1462346003] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462346003] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1462347203] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1462347804] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462349003] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 90% (15 min average)
[1462349603] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462349723] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 84% (15 min average)
[1462349973] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1462350323] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462350443] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462350563] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1462350573] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462350693] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1462350813] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 36% (15 min average)
[1462351259] SERVICE ALERT: ADWSUSCLIENTS;Windows Update Services Event 10022;OK;HARD;3;Eventlog check ok
[1462352363] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 89% (15 min average)
[1462352613] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462352733] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462352758] SERVICE ALERT: ADWU1;Windows Update Services Event 10022;OK;HARD;3;OK: No entries found
[1462352853] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462352963] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1462353083] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1462353453] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 76% (15 min average)
[1462353683] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462353803] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1462353923] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1462354653] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1462354773] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 52% (15 min average)
[1462355123] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1462355723] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1462355843] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1462355963] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 89% (15 min average)
[1462356563] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462356573] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1462356683] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462356693] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1462356804] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462356813] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462357403] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1462358003] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1462358123] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 87% (15 min average)
[1462361311] SERVICE ALERT: ADARC;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1462361429] SERVICE ALERT: ADAUTHDC1;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1462361431] SERVICE ALERT: ADARC;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1462361483] SERVICE ALERT: ADAUTHDC1;Service: AD Web Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462361495] SERVICE ALERT: ADAUTHDC1-old;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462361495] SERVICE ALERT: ADAUTHDC1-old;Service: DFS Namespace;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462361496] SERVICE ALERT: ADAUTHDC1-old;Service: Key Distribution Center;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462361503] SERVICE ALERT: ADAUTHDC1-old;AD Replication;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 30 seconds.
[1462361509] SERVICE ALERT: ADAUTHDC1-old;Service: Intersite Messaging;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462361511] HOST ALERT: ADAUTH-RMAD;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462361513] HOST ALERT: ADAUTHDC1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462361516] SERVICE ALERT: ADAUTHDC1;Service: DNS Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462361516] SERVICE ALERT: ADAUTHDC1;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462361517] SERVICE ALERT: ADAUTHDC1;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462361520] SERVICE ALERT: ADAUTHDC1;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462361535] SERVICE ALERT: ADAUTH-RMAD;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462361538] SERVICE ALERT: ADAUTHDC1;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462361538] SERVICE ALERT: ADAUTHDC1;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462361546] HOST ALERT: ADAUTHDC1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1462361565] HOST ALERT: ADAUTH-RMAD;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1462361569] SERVICE ALERT: ADAUTHDC1;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462361585] SERVICE ALERT: ADAUTH-RMAD;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462361593] SERVICE ALERT: ADAUTHDC1;Service: Distributed File System Replication;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462361595] SERVICE ALERT: ADAUTHDC1-old;VMWare Tools;OK;SOFT;2;VMTools: Started
[1462361595] SERVICE ALERT: ADAUTHDC1-old;Service: DFS Namespace;OK;SOFT;2;Dfs: Started
[1462361596] SERVICE ALERT: ADAUTHDC1-old;AD Replication;WARNING;SOFT;2;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1462361596] SERVICE ALERT: ADAUTHDC1-old;Service: Key Distribution Center;OK;SOFT;2;KDC: Started
[1462361599] HOST ALERT: ADAUTHDC1;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1462361603] SERVICE ALERT: ADAUTHDC1;Service: AD Web Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462361609] SERVICE ALERT: ADAUTHDC1-old;Service: Intersite Messaging;OK;SOFT;2;IsmServ: Started
[1462361615] HOST ALERT: ADAUTH-RMAD;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1462361633] HOST ALERT: ADAUTHDC1;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1462361636] SERVICE ALERT: ADAUTHDC1;Service: DNS Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462361636] SERVICE ALERT: ADAUTHDC1;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462361637] SERVICE ALERT: ADAUTHDC1;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462361639] SERVICE ALERT: ADAUTHDC1;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462361655] SERVICE ALERT: ADAUTH-RMAD;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462361658] SERVICE ALERT: ADAUTHDC1;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462361658] SERVICE ALERT: ADAUTHDC1;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462361663] SERVICE ALERT: ADAUTHDC1;Service: Dell Backup Agent;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462361663] SERVICE ALERT: ADAUTH-RMAD;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462361663] SERVICE ALERT: ADAUTH-RMAD;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462361666] HOST ALERT: ADAUTHDC1;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1462361685] HOST ALERT: ADAUTH-RMAD;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1462361689] SERVICE ALERT: ADAUTHDC1;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462361690] SERVICE ALERT: ADAUTHDC1;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462361691] SERVICE ALERT: ADAUTHDC1;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462361691] SERVICE ALERT: ADAUTH-RMAD;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462361693] SERVICE ALERT: ADAUTH-RMAD;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462361705] SERVICE ALERT: ADAUTH-RMAD;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462361713] SERVICE ALERT: ADAUTHDC1-old;AD Replication;WARNING;HARD;3;Send-MailMessage : The remote name could not be resolved: 'meritsmpt.emich.edu'
[1462361713] SERVICE ALERT: ADAUTHDC1;Service: Distributed File System Replication;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462361717] SERVICE ALERT: ADAUTH-RMAD;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462361718] SERVICE ALERT: ADAUTH-RMAD;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462361719] HOST ALERT: ADAUTHDC1;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1462361721] HOST ALERT: ADAUTH-RMAD;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1462361723] SERVICE ALERT: ADAUTHDC1;Service: AD Web Service;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462361723] SERVICE ALERT: ADAUTH-RMAD;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462361729] SERVICE ALERT: ADAUTHDC1;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462361729] SERVICE ALERT: ADAUTHDC1;Service: Intersite Messaging;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462361753] HOST ALERT: ADAUTHDC1;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1462361753] HOST ALERT: ADAUTH-RMAD;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1462361756] SERVICE ALERT: ADAUTHDC1;Service: DNS Service;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462361756] SERVICE ALERT: ADAUTHDC1;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462361757] SERVICE ALERT: ADAUTHDC1;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462361758] SERVICE ALERT: ADAUTH-RMAD;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462361759] SERVICE ALERT: ADAUTHDC1;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462361760] SERVICE ALERT: ADAUTH-RMAD;E:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462361771] SERVICE ALERT: ADAUTHDC1;Service: Key Distribution Center;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462361775] SERVICE ALERT: ADAUTH-RMAD;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462361778] SERVICE ALERT: ADAUTHDC1;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462361778] SERVICE ALERT: ADAUTHDC1;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462361783] SERVICE ALERT: ADAUTHDC1;Service: Dell Backup Agent;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462361783] SERVICE ALERT: ADAUTH-RMAD;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462361783] SERVICE ALERT: ADAUTH-RMAD;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462361786] HOST ALERT: ADAUTHDC1;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1462361788] HOST ALERT: ADAUTH-RMAD;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1462361805] SERVICE ALERT: ADAUTH-RMAD;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462361810] SERVICE ALERT: ADAUTHDC1;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462361811] SERVICE ALERT: ADAUTHDC1;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462361811] SERVICE ALERT: ADAUTH-RMAD;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462361813] SERVICE ALERT: ADAUTH-RMAD;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462361818] SERVICE ALERT: ADAUTHDC1;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462361819] SERVICE ALERT: ADAUTHDC1;AD Replication;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 30 seconds.
[1462361825] SERVICE ALERT: ADAUTH-RMAD;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462361833] SERVICE ALERT: ADAUTHDC1;Service: Distributed File System Replication;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462361835] HOST ALERT: ADAUTH-RMAD;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1462361837] SERVICE ALERT: ADAUTH-RMAD;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462361838] SERVICE ALERT: ADAUTH-RMAD;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462361840] HOST ALERT: ADAUTHDC1;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1462361843] SERVICE ALERT: ADAUTH-RMAD;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462361849] SERVICE ALERT: ADAUTHDC1;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462361849] SERVICE ALERT: ADAUTHDC1;Service: Intersite Messaging;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462361851] SERVICE ALERT: ADAUTH-RMAD;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462361851] SERVICE ALERT: ADAUTH-RMAD;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462361867] HOST ALERT: ADAUTH-RMAD;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1462361877] SERVICE ALERT: ADAUTH-RMAD;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462361878] SERVICE ALERT: ADAUTH-RMAD;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462361879] HOST ALERT: ADAUTHDC1;DOWN;HARD;10;(Host check timed out after 30.01 seconds)
[1462361880] SERVICE ALERT: ADAUTH-RMAD;E:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462361891] SERVICE ALERT: ADAUTHDC1;Service: Key Distribution Center;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462361903] SERVICE ALERT: ADAUTHDC1;Service: Dell Backup Agent;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462361903] SERVICE ALERT: ADAUTH-RMAD;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462361903] SERVICE ALERT: ADAUTH-RMAD;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462361907] HOST ALERT: ADAUTH-RMAD;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1462361925] SERVICE ALERT: ADAUTH-RMAD;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462361930] SERVICE ALERT: ADAUTHDC1;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462361931] SERVICE ALERT: ADAUTHDC1;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462361931] SERVICE ALERT: ADAUTH-RMAD;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462361931] SERVICE ALERT: ADAUTHDC1;Service: Remote Management;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462361932] SERVICE ALERT: ADAUTHDC1;Service: AD Domain Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462361933] SERVICE ALERT: ADAUTH-RMAD;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462361938] SERVICE ALERT: ADAUTHDC1;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462361939] SERVICE ALERT: ADAUTHDC1;AD Replication;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 30 seconds.
[1462361945] SERVICE ALERT: ADAUTHDC3;AD Replication;WARNING;SOFT;1;Send-MailMessage : The remote name could not be resolved: 'meritsmtp.emich.edu'
[1462361957] SERVICE ALERT: ADAUTHDC1;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462361957] SERVICE ALERT: ADAUTH-RMAD;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462361957] SERVICE ALERT: ADAUTHDC1;Service: DFS Namespace;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462361957] SERVICE ALERT: ADAUTHDC1;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462361958] SERVICE ALERT: ADAUTH-RMAD;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462361959] SERVICE ALERT: ADAUTHDC1;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462361963] SERVICE ALERT: ADAUTH-RMAD;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462361969] SERVICE ALERT: ADAUTHDC1;Service: Intersite Messaging;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462361969] SERVICE ALERT: ADAUTHDC1;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462361971] SERVICE ALERT: ADAUTH-RMAD;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462361971] SERVICE ALERT: ADAUTH-RMAD;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462361997] SERVICE ALERT: ADAUTH-RMAD;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462361998] SERVICE ALERT: ADAUTH-RMAD;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462362000] SERVICE ALERT: ADAUTH-RMAD;E:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462362011] SERVICE ALERT: ADAUTHDC1;Service: Key Distribution Center;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462362044] SERVICE ALERT: ADAUTHDC1;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462362045] SERVICE ALERT: ADAUTH-RMAD;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462362051] SERVICE ALERT: ADAUTHDC1;Service: Remote Management;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462362052] SERVICE ALERT: ADAUTHDC1;Service: AD Domain Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462362058] SERVICE ALERT: ADAUTHDC1;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462362059] SERVICE ALERT: ADAUTHDC1;AD Replication;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 30 seconds.
[1462362065] SERVICE ALERT: ADAUTHDC3;AD Replication;WARNING;SOFT;2;Send-MailMessage : The remote name could not be resolved: 'meritsmtp.emich.edu'
[1462362076] SERVICE ALERT: ADAUTHDC1;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462362077] SERVICE ALERT: ADAUTHDC1;Service: DFS Namespace;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462362077] SERVICE ALERT: ADAUTHDC1;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462362079] SERVICE ALERT: ADAUTHDC1;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462362091] SERVICE ALERT: ADAUTH-RMAD;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462362091] SERVICE ALERT: ADAUTH-RMAD;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462362117] SERVICE ALERT: ADAUTH-RMAD;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462362164] SERVICE ALERT: ADAUTHDC1;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462362171] SERVICE ALERT: ADAUTHDC1;Service: Remote Management;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462362172] SERVICE ALERT: ADAUTHDC1;Service: AD Domain Service;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462362185] SERVICE ALERT: ADAUTHDC3;AD Replication;WARNING;HARD;3;Send-MailMessage : The remote name could not be resolved: 'meritsmtp.emich.edu'
[1462362195] SERVICE ALERT: ADAUTHDC1;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462362197] SERVICE ALERT: ADAUTHDC1;Service: DFS Namespace;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462362197] SERVICE ALERT: ADAUTHDC1;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462362198] SERVICE ALERT: ADAUTHDC1;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462362284] SERVICE ALERT: ADAUTHDC1;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462362923] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1462363043] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1462363063] SERVICE ALERT: idm1.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 6
[1462363123] SERVICE ALERT: idm4.emich.edu;LDAP Stats;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1462363153] HOST ALERT: idm4.emich.edu;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462363183] SERVICE ALERT: idm1.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 10
[1462363243] HOST ALERT: idm4.emich.edu;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1462363243] SERVICE ALERT: idm4.emich.edu;LDAP Stats;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1462363273] HOST ALERT: idm4.emich.edu;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1462363303] SERVICE ALERT: idm1.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 12
[1462363316] SERVICE ALERT: idm2.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 12
[1462363363] HOST ALERT: idm4.emich.edu;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1462363363] SERVICE ALERT: idm4.emich.edu;LDAP Stats;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1462363381] SERVICE ALERT: idm4.emich.edu;eDirectory Service;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1462363385] SERVICE ALERT: idm3.emich.edu;eDirectory Service;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1462363389] SERVICE ALERT: idm3.emich.edu;NSDS Errors;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1462363393] HOST ALERT: idm4.emich.edu;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1462363408] SERVICE ALERT: idm4.emich.edu;NSDS Errors;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 60 seconds.
[1462363413] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462363436] SERVICE ALERT: idm2.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 12
[1462363438] HOST ALERT: idm4.emich.edu;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1462363462] SERVICE ALERT: idm4.emich.edu;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462363492] HOST ALERT: idm4.emich.edu;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1462363496] SERVICE ALERT: idm3.emich.edu;eDirectory Service;OK;SOFT;2;OK: ndsd status is good.
[1462363501] SERVICE ALERT: idm4.emich.edu;eDirectory Service;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1462363528] SERVICE ALERT: idm4.emich.edu;NSDS Errors;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 60 seconds.
[1462363531] HOST ALERT: idm4.emich.edu;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1462363548] SERVICE ALERT: idm3.emich.edu;NSDS Errors;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 60 seconds.
[1462363556] SERVICE ALERT: idm2.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 12
[1462363582] SERVICE ALERT: idm4.emich.edu;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1462363612] HOST ALERT: idm4.emich.edu;DOWN;SOFT;9;(Host check timed out after 30.00 seconds)
[1462363615] SERVICE ALERT: idm3.emich.edu;LDAP Stats;UNKNOWN;SOFT;1;Difference is negative. LDAP recently rebooted.
[1462363621] SERVICE ALERT: idm4.emich.edu;eDirectory Service;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1462363647] SERVICE ALERT: idm4.emich.edu;NSDS Errors;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 60 seconds.
[1462363651] HOST ALERT: idm4.emich.edu;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1462363668] SERVICE ALERT: idm3.emich.edu;NSDS Errors;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 60 seconds.
[1462363695] SERVICE ALERT: idm4.emich.edu;PING;OK;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1462363699] HOST ALERT: idm4.emich.edu;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1462363735] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1462363766] HOST DOWNTIME ALERT: idm1.emich.edu;STARTED; Host has entered a period of scheduled downtime
[1462363766] HOST DOWNTIME ALERT: idm2.emich.edu;STARTED; Host has entered a period of scheduled downtime
[1462363766] HOST DOWNTIME ALERT: idm3.emich.edu;STARTED; Host has entered a period of scheduled downtime
[1462363766] HOST DOWNTIME ALERT: idm4.emich.edu;STARTED; Host has entered a period of scheduled downtime
[1462363766] HOST DOWNTIME ALERT: idmrbs.emich.edu;STARTED; Host has entered a period of scheduled downtime
[1462363766] HOST DOWNTIME ALERT: idmrbstest.emich.edu;STARTED; Host has entered a period of scheduled downtime
[1462363862] SERVICE ALERT: ADDC5;AD Replication;WARNING;SOFT;1;WARNING: Replication with ADDC11 failed!
[1462363893] SERVICE ALERT: ADDC11;Service: NetLogin;CRITICAL;SOFT;1;Netlogon: Stopped
[1462363953] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1462363982] SERVICE ALERT: ADDC5;AD Replication;WARNING;SOFT;2;WARNING: Replication with ADDC11 failed!
[1462364002] HOST ALERT: idm3.emich.edu;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462364013] SERVICE ALERT: ADDC11;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1462364013] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462364013] SERVICE ALERT: idm3.emich.edu;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462364022] HOST ALERT: idm3.emich.edu;UP;SOFT;2;PING WARNING - Packet loss = 0%, RTA = 4711.90 ms
[1462364102] SERVICE ALERT: ADDC5;AD Replication;WARNING;HARD;3;WARNING: Replication with ADDC11 failed!
[1462364103] SERVICE ALERT: ADDC11;Service: Intersite Messaging;CRITICAL;SOFT;1;IsmServ: Stopped
[1462364127] SERVICE ALERT: idm3.emich.edu;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1462364175] SERVICE ALERT: ADDC11;Service: Key Distribution Center;CRITICAL;SOFT;1;KDC: Stopped
[1462364188] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 8
[1462364208] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 8
[1462364211] SERVICE ALERT: idm4.emich.edu;eDirectory Service;OK;HARD;3;OK: ndsd status is good.
[1462364223] SERVICE ALERT: ADDC11;Service: Intersite Messaging;UNKNOWN;SOFT;2;Failed to open service manager: 1115: A system shutdown is in progress.
[1462364271] SERVICE ALERT: ADDC11;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462364271] SERVICE ALERT: ADDC11;Service: Remote Management;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462364296] SERVICE ALERT: ADDC11;Service: AD Web Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462364301] HOST ALERT: ADDC11;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462364315] SERVICE ALERT: ADDC11;Service: Key Distribution Center;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462364345] HOST ALERT: ADDC11;UP;SOFT;2;PING WARNING - Packet loss = 96%, RTA = 0.42 ms
[1462364349] SERVICE ALERT: ADDC11;Service: AD Domain Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462364363] SERVICE ALERT: ADDC11;Service: Intersite Messaging;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462364391] SERVICE ALERT: ADDC11;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462364391] SERVICE ALERT: ADDC11;Service: Remote Management;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462364416] SERVICE ALERT: ADDC11;Service: AD Web Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462364435] SERVICE ALERT: ADDC11;Service: Key Distribution Center;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462364469] SERVICE ALERT: ADDC11;Service: AD Domain Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462364511] SERVICE ALERT: ADDC11;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462364511] SERVICE ALERT: ADDC11;Service: Remote Management;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462364536] SERVICE ALERT: ADDC11;Service: AD Web Service;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462364564] SERVICE ALERT: ADDC11;Service: DNS Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462364589] SERVICE ALERT: ADDC11;Service: AD Domain Service;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462364613] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 84% (15 min average)
[1462364626] SERVICE ALERT: ADDC11;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462364630] SERVICE ALERT: ADDC11;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462364631] SERVICE ALERT: ADDC11;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462364631] SERVICE ALERT: ADDC11;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462364631] SERVICE ALERT: ADDC11;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462364632] SERVICE ALERT: ADDC11;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462364632] SERVICE ALERT: ADDC11;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462364632] SERVICE ALERT: ADDC11;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462364633] SERVICE ALERT: ADDC11;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462364633] SERVICE ALERT: ADDC11;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462364633] SERVICE ALERT: ADDC11;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462364633] SERVICE ALERT: ADDC11;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462364633] SERVICE ALERT: ADDC11;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462364645] SERVICE ALERT: ADDC11;Service: Distributed File System Replication;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462364684] SERVICE ALERT: ADDC11;Service: DNS Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462364707] SERVICE ALERT: ADDC5;AD Replication;OK;HARD;3;OK: replication is up and running
[1462364733] SERVICE ALERT: ADDC11;Service: DFS Namespace;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462364743] SERVICE ALERT: ADDC11;AD Replication;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 30 seconds.
[1462364746] SERVICE ALERT: ADDC11;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462364750] SERVICE ALERT: ADDC11;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462364751] SERVICE ALERT: ADDC11;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462364751] SERVICE ALERT: ADDC11;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462364751] SERVICE ALERT: ADDC11;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462364752] SERVICE ALERT: ADDC11;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462364752] SERVICE ALERT: ADDC11;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462364752] SERVICE ALERT: ADDC11;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462364753] SERVICE ALERT: ADDC11;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462364753] SERVICE ALERT: ADDC11;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462364753] SERVICE ALERT: ADDC11;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462364753] SERVICE ALERT: ADDC11;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462364753] SERVICE ALERT: ADDC11;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462364764] SERVICE ALERT: ADDC11;Service: Distributed File System Replication;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462364767] SERVICE ALERT: ADDC11;Service: Dell Backup Agent;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462364804] SERVICE ALERT: ADDC11;Service: DNS Service;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462364843] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1462364853] SERVICE ALERT: ADDC11;Service: DFS Namespace;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462364866] SERVICE ALERT: ADDC11;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462364870] SERVICE ALERT: ADDC11;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462364871] SERVICE ALERT: ADDC11;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462364871] SERVICE ALERT: ADDC11;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462364871] SERVICE ALERT: ADDC11;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462364872] SERVICE ALERT: ADDC11;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462364872] SERVICE ALERT: ADDC11;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462364872] SERVICE ALERT: ADDC11;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462364873] SERVICE ALERT: ADDC11;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462364873] SERVICE ALERT: ADDC11;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462364873] SERVICE ALERT: ADDC11;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462364873] SERVICE ALERT: ADDC11;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462364873] SERVICE ALERT: ADDC11;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462364884] SERVICE ALERT: ADDC11;Service: Distributed File System Replication;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462364887] SERVICE ALERT: ADDC11;Service: Dell Backup Agent;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462364963] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 83% (15 min average)
[1462364973] SERVICE ALERT: ADDC11;Service: DFS Namespace;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462365007] SERVICE ALERT: ADDC11;Service: Dell Backup Agent;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462365035] HOST DOWNTIME ALERT: ADDC11;STARTED; Host has entered a period of scheduled downtime
[1462365036] SERVICE ALERT: idmrbs.emich.edu;LDAP Stats;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1462365048] HOST ALERT: idmrbs.emich.edu;DOWN;SOFT;1;PING CRITICAL - Packet loss = 0%, RTA = 6550.76 ms
[1462365112] HOST ALERT: idmrbs.emich.edu;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1462365146] SERVICE ALERT: idmrbs.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1462365356] SERVICE ALERT: idm2.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1462365387] SERVICE ALERT: idm4.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1462365408] SERVICE ALERT: idm3.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1462365563] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1462365683] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 92% (15 min average)
[1462365703] SERVICE ALERT: idm1.emich.edu;NSDS Errors;OK;HARD;3;OK: 0 Errors Reported
[1462365813] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462365832] HOST ALERT: idmwarehouse;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462365896] HOST ALERT: idmwarehouse;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.12 ms
[1462365933] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 53% (15 min average)
[1462366327] SERVICE ALERT: idmrbs.emich.edu;eDirectory Service;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1462366353] HOST ALERT: idmrbs.emich.edu;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462366356] SERVICE ALERT: idmrbs.emich.edu;LDAP Stats;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1462366386] HOST ALERT: idmrbs.emich.edu;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1462366447] SERVICE ALERT: idmrbs.emich.edu;eDirectory Service;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1462366475] SERVICE ALERT: idmrbs.emich.edu;LDAP Stats;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 10 seconds.
[1462366476] HOST ALERT: idmrbs.emich.edu;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1462366507] SERVICE ALERT: idmrbs.emich.edu;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462366537] HOST ALERT: idmrbs.emich.edu;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1462366538] SERVICE ALERT: idmrbs.emich.edu;Banner Driver;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 60 seconds.
[1462366566] SERVICE ALERT: idmrbs.emich.edu;eDirectory Service;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1462366568] HOST ALERT: idmrbs.emich.edu;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1462366595] SERVICE ALERT: idmrbs.emich.edu;LDAP Stats;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 10 seconds.
[1462366625] HOST ALERT: idmrbs.emich.edu;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1462366627] SERVICE ALERT: idmrbs.emich.edu;PING;CRITICAL;SOFT;2;PING CRITICAL - Packet loss = 100%
[1462366657] HOST ALERT: idmrbs.emich.edu;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1462366658] SERVICE ALERT: idmrbs.emich.edu;Banner Driver;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 60 seconds.
[1462366688] HOST ALERT: idmrbs.emich.edu;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1462366691] SERVICE ALERT: idmrbs.emich.edu;NSDS Errors;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 60 seconds.
[1462366721] HOST ALERT: idmrbs.emich.edu;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1462366729] SERVICE ALERT: idmrbs.emich.edu;Banner Driver;CRITICAL;HARD;3;Unknown: Error reported?
[1462366733] HOST ALERT: idmrbs.emich.edu;UP;SOFT;10;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1462366741] SERVICE ALERT: idmrbs.emich.edu;PING;OK;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.38 ms
[1462366751] SERVICE ALERT: idmrbs.emich.edu;NSDS Errors;OK;SOFT;2;OK:
[1462366837] HOST ALERT: ADDC11;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462366883] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1462366927] HOST ALERT: ADDC11;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1462366953] SERVICE ALERT: idm4.emich.edu;LDAP Stats;UNKNOWN;SOFT;1;Difference is negative. LDAP recently rebooted.
[1462366965] HOST ALERT: ADDC11;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1462367003] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 75% (15 min average)
[1462367019] HOST ALERT: ADDC11;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1462367073] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1462367109] HOST ALERT: ADDC11;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1462367157] SERVICE ALERT: idmrbs.emich.edu;eDirectory Service;OK;HARD;3;OK: ndsd status is good.
[1462367157] SERVICE ALERT: ADWU1;Windows Update Services Event 10022;CRITICAL;SOFT;1;Windows Server Update Services, The last catalog synchronization attempt was unsuccessful.
[1462367173] HOST ALERT: ADDC11;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1462367186] SERVICE ALERT: idmrbs.emich.edu;LDAP Stats;UNKNOWN;HARD;3;Difference is negative. LDAP recently rebooted.
[1462367234] HOST ALERT: ADDC11;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1462367277] SERVICE ALERT: ADWU1;Windows Update Services Event 10022;CRITICAL;SOFT;2;Windows Server Update Services, The last catalog synchronization attempt was unsuccessful.
[1462367296] HOST ALERT: ADDC11;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1462367320] SERVICE ALERT: idmrbs.emich.edu;Banner Driver;OK;HARD;3;0
[1462367345] SERVICE ALERT: idm3.emich.edu;LDAP Stats;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 10 seconds.
[1462367375] HOST ALERT: idm3.emich.edu;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462367386] HOST ALERT: ADDC11;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1462367397] SERVICE ALERT: ADWU1;Windows Update Services Event 10022;CRITICAL;HARD;3;Windows Server Update Services, The last catalog synchronization attempt was unsuccessful.
[1462367437] HOST ALERT: ADDC11;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1462367439] HOST ALERT: idm3.emich.edu;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1462367455] SERVICE ALERT: idm3.emich.edu;LDAP Stats;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1462367575] SERVICE ALERT: idm3.emich.edu;LDAP Stats;CRITICAL;HARD;3;CRITICAL: Failed to connect to LDAP!
[1462367785] SERVICE ALERT: idmrbs.emich.edu;LDAP Stats;OK;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1462368103] SERVICE ALERT: idm1.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 4
[1462368175] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1462368203] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1462368223] SERVICE ALERT: idm1.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 5
[1462368323] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1462368333] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462368343] SERVICE ALERT: idm1.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 7
[1462368377] SERVICE ALERT: idm2.emich.edu;NSDS Errors;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1462368387] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 7
[1462368398] SERVICE ALERT: idm2.emich.edu;eDirectory Service;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1462368408] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 7
[1462368443] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 94% (15 min average)
[1462368453] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1462368507] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 7
[1462368518] SERVICE ALERT: idm2.emich.edu;eDirectory Service;OK;SOFT;2;OK: ndsd status is good.
[1462368528] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 7
[1462368534] SERVICE ALERT: idm2.emich.edu;NSDS Errors;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 60 seconds.
[1462368573] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1462368598] SERVICE ALERT: idm2.emich.edu;LDAP Stats;UNKNOWN;SOFT;1;Difference is negative. LDAP recently rebooted.
[1462368627] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 2
[1462368627] SERVICE FLAPPING ALERT: idm4.emich.edu;NSDS Errors;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1462368648] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 2
[1462368648] SERVICE FLAPPING ALERT: idm3.emich.edu;NSDS Errors;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1462368654] SERVICE ALERT: idm2.emich.edu;NSDS Errors;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 60 seconds.
[1462368718] SERVICE ALERT: idm2.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1462368944] SERVICE ALERT: idm1.emich.edu;NSDS Errors;OK;HARD;3;OK: 0 Errors Reported
[1462368944] SERVICE FLAPPING ALERT: idm1.emich.edu;NSDS Errors;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1462369207] SERVICE ALERT: idm2.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 4
[1462369207] SERVICE FLAPPING ALERT: idm2.emich.edu;NSDS Errors;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1462369275] HOST ALERT: ADAUTHDC1;UP;HARD;10;PING OK - Packet loss = 72%, RTA = 0.38 ms
[1462369285] HOST ALERT: idm1.emich.edu;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462369351] SERVICE ALERT: ADAUTHDC1;Service: Remote Management;OK;HARD;3;WinRM: Started
[1462369352] SERVICE ALERT: ADAUTHDC1;Service: AD Domain Service;OK;HARD;3;NTDS: Started
[1462369353] HOST ALERT: idm1.emich.edu;UP;SOFT;2;PING OK - Packet loss = 44%, RTA = 0.31 ms
[1462369375] SERVICE ALERT: ADAUTHDC1;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1462369377] SERVICE ALERT: ADAUTHDC1;Service: DFS Namespace;OK;HARD;3;Dfs: Started
[1462369377] SERVICE ALERT: ADAUTHDC1;Memory Usage;OK;HARD;3;Memory usage: total:10239.64 MB - used: 679.44 MB (7%) - free: 9560.20 MB (93%)
[1462369378] SERVICE ALERT: ADAUTHDC1;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1462369434] SERVICE ALERT: idm1.emich.edu;LDAP Stats;CRITICAL;SOFT;1;CRITICAL: Failed to connect to LDAP!
[1462369464] SERVICE ALERT: ADAUTHDC1;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1462369468] SERVICE ALERT: ADAUTHDC1;Service: Windows Time;OK;HARD;3;W32Time: Started
[1462369503] SERVICE ALERT: ADAUTHDC1;Service: AD Web Service;OK;HARD;3;ADWS: Started
[1462369504] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;CRITICAL;SOFT;1;CRITICAL: cn=ADAUTH Driver,cn=driverset1,o=system is stopped!
[1462369536] SERVICE ALERT: ADAUTHDC1;Service: RPC;OK;HARD;3;RpcSs: Started
[1462369536] SERVICE ALERT: ADAUTHDC1;Service: DNS Service;OK;HARD;3;DNS: Started
[1462369537] SERVICE ALERT: ADAUTHDC1;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 4 minute(s)
[1462369539] SERVICE ALERT: ADAUTHDC1;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1462369554] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1462369558] SERVICE ALERT: ADAUTHDC1;VMWare Tools;OK;HARD;3;VMTools: Started
[1462369558] SERVICE ALERT: ADAUTHDC1;Service: WMI;OK;HARD;3;Winmgmt: Started
[1462369566] SERVICE ALERT: idm1.emich.edu;Loopback Driver;CRITICAL;SOFT;1;CRITICAL: cn=Generic Loopback,cn=driverset1,o=system is stopped!
[1462369603] SERVICE ALERT: idm1.emich.edu;NSDS Errors;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 60 seconds.
[1462369613] SERVICE ALERT: ADAUTHDC1;Service: Distributed File System Replication;OK;HARD;3;DFSR: Started
[1462369623] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;CRITICAL;SOFT;2;CRITICAL: cn=ADAUTH Driver,cn=driverset1,o=system is stopped!
[1462369643] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1462369663] SERVICE ALERT: idm1.emich.edu;NSDS Errors;OK;SOFT;2;OK: 0 Errors Reported
[1462369681] SERVICE ALERT: ADAUTHDC1;Service: Dell Backup Agent;OK;HARD;3;ErdAgent: Started
[1462369687] SERVICE ALERT: idm1.emich.edu;Loopback Driver;OK;SOFT;2;0
[1462369709] SERVICE ALERT: ADAUTHDC1;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1462369711] SERVICE ALERT: ADAUTHDC1;C:\ Drive Space;OK;HARD;3;c:\ - total: 79.66 Gb - used: 15.45 Gb (19%) - free 64.21 Gb (81%)
[1462369744] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;OK;SOFT;3;0
[1462369749] SERVICE ALERT: ADAUTHDC1;Service: Intersite Messaging;OK;HARD;3;IsmServ: Started
[1462369749] SERVICE ALERT: ADAUTHDC1;CPU Load;OK;HARD;3;CPU Load 0% (15 min average)
[1462369752] SERVICE ALERT: ADAUTH-RMAD;Service: RPC;OK;HARD;3;RpcSs: Started
[1462369759] HOST ALERT: ADAUTH-RMAD;UP;HARD;10;PING OK - Packet loss = 37%, RTA = 0.35 ms
[1462369773] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (15 min average)
[1462369778] SERVICE ALERT: ADAUTH-RMAD;VMWare Tools;OK;HARD;3;VMTools: Started
[1462369780] SERVICE ALERT: ADAUTH-RMAD;E:\ Drive Space;OK;HARD;3;e:\ - total: 199.87 Gb - used: 102.93 Gb (52%) - free 96.94 Gb (48%)
[1462369791] SERVICE ALERT: ADAUTHDC1;Service: Key Distribution Center;OK;HARD;3;KDC: Started
[1462369794] SERVICE ALERT: idm2.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1462369825] SERVICE ALERT: ADAUTH-RMAD;Memory Usage;OK;HARD;3;Memory usage: total:10239.64 MB - used: 1472.01 MB (14%) - free: 8767.63 MB (86%)
[1462369827] SERVICE ALERT: idm4.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1462369835] SERVICE ALERT: ADAUTHDC1;AD Replication;OK;HARD;3;OK: replication is up and running
[1462369838] SERVICE ALERT: ADAUTHDC1;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1462369849] SERVICE ALERT: idm3.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1462369871] SERVICE ALERT: ADAUTH-RMAD;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1462369871] SERVICE ALERT: ADAUTH-RMAD;Service: Windows Time;OK;HARD;3;W32Time: Started
[1462369897] SERVICE ALERT: ADAUTH-RMAD;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1462369988] SERVICE ALERT: ADAUTHDC3;AD Replication;OK;HARD;3;OK: replication is up and running
[1462370118] SERVICE ALERT: ADAUTHDC1-old;AD Replication;OK;HARD;3;OK: replication is up and running
[1462370155] SERVICE ALERT: ADAUTH-RMAD;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1462370205] SERVICE ALERT: ADAUTH-RMAD;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1462370243] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1462370281] SERVICE ALERT: ADAUTH-RMAD;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 9 minute(s)
[1462370281] SERVICE ALERT: ADAUTH-RMAD;C:\ Drive Space;OK;HARD;3;c:\ - total: 79.66 Gb - used: 24.78 Gb (31%) - free 54.87 Gb (69%)
[1462370311] SERVICE ALERT: ADAUTH-RMAD;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1462370311] SERVICE ALERT: ADAUTH-RMAD;Service: WMI;OK;HARD;3;Winmgmt: Started
[1462370337] SERVICE ALERT: ADAUTH-RMAD;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1462370337] SERVICE ALERT: ADAUTH-RMAD;CPU Load;OK;HARD;3;CPU Load 0% (15 min average)
[1462370363] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 68% (15 min average)
[1462370373] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1462370493] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 88% (15 min average)
[1462370959] HOST DOWNTIME ALERT: idm3.emich.edu;STOPPED; Host has exited from a period of scheduled downtime
[1462370959] HOST DOWNTIME ALERT: idmrbstest.emich.edu;STOPPED; Host has exited from a period of scheduled downtime
[1462370959] HOST DOWNTIME ALERT: idm2.emich.edu;STOPPED; Host has exited from a period of scheduled downtime
[1462370959] HOST DOWNTIME ALERT: idm1.emich.edu;STOPPED; Host has exited from a period of scheduled downtime
[1462370959] HOST DOWNTIME ALERT: idmrbs.emich.edu;STOPPED; Host has exited from a period of scheduled downtime
[1462370959] HOST DOWNTIME ALERT: idm4.emich.edu;STOPPED; Host has exited from a period of scheduled downtime
[1462370963] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1462371083] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 92% (15 min average)
[1462371595] HOST ALERT: emunix;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462371685] HOST ALERT: emunix;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1462371775] HOST ALERT: emunix;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1462371865] HOST ALERT: emunix;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1462371955] HOST ALERT: emunix;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1462372045] HOST ALERT: emunix;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1462372135] HOST ALERT: emunix;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1462372193] HOST ALERT: ADDC11;UP;HARD;10;PING WARNING - Packet loss = 93%, RTA = 3.10 ms
[1462372225] HOST ALERT: emunix;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1462372229] HOST DOWNTIME ALERT: ADDC11;STOPPED; Host has exited from a period of scheduled downtime
[1462372293] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462372315] HOST ALERT: emunix;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1462372405] HOST ALERT: emunix;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1462372413] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462372533] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462373733] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462374002] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (35.7% change >= 20.0% threshold)
[1462374002] SERVICE FLAPPING ALERT: idm2.emich.edu;NSDS Errors;STARTED; Service appears to have started flapping (24.0% change >= 20.0% threshold)
[1462374002] SERVICE FLAPPING ALERT: idm3.emich.edu;NSDS Errors;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1462374002] SERVICE FLAPPING ALERT: idm4.emich.edu;NSDS Errors;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1462374002] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1462374002] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1462374002] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1462374002] SERVICE DOWNTIME ALERT: NETIDADPROXY;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1462374002] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1462374002] HOST DOWNTIME ALERT: ADAUTHDC1;STARTED; Host has entered a period of scheduled downtime
[1462374002] HOST DOWNTIME ALERT: ADAUTHDC3;STARTED; Host has entered a period of scheduled downtime
[1462374029] HOST ALERT: emunix;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1462374333] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1462374597] SERVICE ALERT: ADWU1;Windows Update Services Event 10022;OK;HARD;3;OK: No entries found
[1462374671] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;UNKNOWN;SOFT;1;CHECK_NRPE: Error receiving data from daemon.
[1462374791] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;OK;SOFT;2;OK: No entries found
[1462375199] HOST DOWNTIME ALERT: ADAUTHDC1;STOPPED; Host has exited from a period of scheduled downtime
[1462375242] HOST DOWNTIME ALERT: ADAUTHDC3;STOPPED; Host has exited from a period of scheduled downtime
[1462375973] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462376033] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462376093] SERVICE ALERT: hallefirst2;Device Status;OK;SOFT;3;Processing job from tray 4
[1462376133] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 82% (15 min average)
[1462376733] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1462376739] SERVICE ALERT: AX-QUAL;Service: Group Policy;CRITICAL;SOFT;1;gpsvc: Stopped
[1462376807] SERVICE ALERT: AX-QUAL;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462376853] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462376859] SERVICE ALERT: AX-QUAL;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1462376907] SERVICE ALERT: AX-QUAL;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1462376973] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462377293] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462377354] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462377413] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462377603] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (37.4% change >= 20.0% threshold)
[1462377603] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1462377603] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1462377603] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1462377603] SERVICE DOWNTIME ALERT: NETIDADPROXY;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1462377603] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1462378059] SERVICE ALERT: AX-QUAL;Service: Group Policy;CRITICAL;SOFT;1;gpsvc: Stopped
[1462378061] SERVICE ALERT: AX-QUAL;Service: RPC;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1462378096] SERVICE ALERT: AX-QUAL;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462378179] SERVICE ALERT: AX-QUAL;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1462378181] SERVICE ALERT: AX-QUAL;Service: RPC;OK;SOFT;2;RpcSs: Started
[1462378196] SERVICE ALERT: AX-QUAL;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1462378614] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Ready
[1462379373] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 86% (15 min average)
[1462379402] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (40.9% change >= 20.0% threshold)
[1462379402] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1462379402] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1462379402] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1462379402] SERVICE DOWNTIME ALERT: NETIDADPROXY;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1462379402] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1462379561] SERVICE ALERT: markjeffersonpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1462379590] SERVICE ALERT: markjeffersonpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1462379591] HOST ALERT: markjeffersonpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462379621] SERVICE ALERT: markjeffersonpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1462379644] HOST ALERT: markjeffersonpk;UP;SOFT;2;PING OK - Packet loss = 79%, RTA = 0.62 ms
[1462379649] SERVICE ALERT: markjeffersonpk;Device Status;OK;SOFT;2;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1462379675] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1462379676] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;SOFT;3;Printer Paper Is OK! All Trays Still Full
[1462379736] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 72% - OK!
[1462379973] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1462380093] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462380213] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1462381283] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1462381403] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 77% (15 min average)
[1462382003] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462382123] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 80% (15 min average)
[1462384413] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 82% (15 min average)
[1462385013] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1462385133] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462385253] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462386453] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462387053] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462388853] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 78% (15 min average)
[1462390053] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1462390173] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462390293] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462390877] SERVICE ALERT: R25-GAMMA2012;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462390907] HOST ALERT: R25-GAMMA2012;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462390913] SERVICE ALERT: R25-GAMMA2012;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462390937] SERVICE ALERT: R25-GAMMA2012;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462390943] HOST ALERT: R25-GAMMA2012;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1462390962] HOST DOWNTIME ALERT: R25-GAMMA2012;STARTED; Host has entered a period of scheduled downtime
[1462390975] SERVICE ALERT: R25-BETA2012;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462390997] SERVICE ALERT: R25-GAMMA2012;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462390998] HOST DOWNTIME ALERT: R25-EPSILON2012;STARTED; Host has entered a period of scheduled downtime
[1462391003] SERVICE ALERT: R25-BETA2012;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462391004] SERVICE ALERT: R25-GAMMA2012;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462391005] HOST ALERT: R25-BETA2012;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462391027] HOST ALERT: R25-GAMMA2012;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1462391033] SERVICE ALERT: R25-GAMMA2012;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462391035] HOST DOWNTIME ALERT: R25-BETA2012;STARTED; Host has entered a period of scheduled downtime
[1462391035] SERVICE ALERT: R25-EPSILON2012;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462391040] SERVICE ALERT: R25-EPSILON2012;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462391041] SERVICE ALERT: R25-BETA2012;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462391057] SERVICE ALERT: R25-GAMMA2012;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462391063] HOST ALERT: R25-GAMMA2012;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1462391065] HOST ALERT: R25-EPSILON2012;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1462391071] HOST ALERT: R25-BETA2012;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1462391075] SERVICE ALERT: R25-EPSILON2012;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462391091] SERVICE ALERT: R25-EPSILON2012;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462391093] SERVICE ALERT: R25-GAMMA2012;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462391095] SERVICE ALERT: R25-BETA2012;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462391105] HOST ALERT: R25-EPSILON2012;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1462391116] SERVICE ALERT: R25-GAMMA2012;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462391117] SERVICE ALERT: R25-GAMMA2012;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462391122] SERVICE ALERT: R25-BETA2012;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462391123] HOST ALERT: R25-GAMMA2012;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1462391124] SERVICE ALERT: R25-GAMMA2012;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462391125] HOST ALERT: R25-BETA2012;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1462391127] SERVICE ALERT: R25-EPSILON2012;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462391149] SERVICE ALERT: R25-GAMMA2012;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462391150] SERVICE ALERT: R25-BETA2012;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462391150] SERVICE ALERT: R25-BETA2012;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462391153] SERVICE ALERT: R25-GAMMA2012;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462391154] HOST ALERT: R25-GAMMA2012;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1462391155] SERVICE ALERT: R25-EPSILON2012;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462391155] SERVICE ALERT: R25-EPSILON2012;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462391157] SERVICE ALERT: R25-BETA2012;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462391157] HOST ALERT: R25-EPSILON2012;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1462391160] SERVICE ALERT: R25-EPSILON2012;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462391161] SERVICE ALERT: R25-BETA2012;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462391177] SERVICE ALERT: R25-GAMMA2012;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462391180] HOST ALERT: R25-BETA2012;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1462391184] SERVICE ALERT: R25-BETA2012;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462391190] HOST ALERT: R25-EPSILON2012;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1462391195] SERVICE ALERT: R25-EPSILON2012;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462391203] SERVICE ALERT: R25-EPSILON2012;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462391203] SERVICE ALERT: R25-GAMMA2012;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462391207] HOST ALERT: R25-GAMMA2012;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1462391211] SERVICE ALERT: R25-EPSILON2012;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462391213] SERVICE ALERT: R25-EPSILON2012;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462391213] SERVICE ALERT: R25-GAMMA2012;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462391214] HOST ALERT: R25-BETA2012;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1462391215] SERVICE ALERT: R25-BETA2012;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462391225] HOST ALERT: R25-EPSILON2012;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1462391236] SERVICE ALERT: R25-EPSILON2012;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462391236] SERVICE ALERT: R25-GAMMA2012;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462391237] SERVICE ALERT: R25-EPSILON2012;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462391238] SERVICE ALERT: R25-GAMMA2012;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462391242] SERVICE ALERT: R25-BETA2012;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462391243] SERVICE ALERT: R25-BETA2012;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462391243] HOST ALERT: R25-GAMMA2012;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1462391244] SERVICE ALERT: R25-GAMMA2012;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462391245] HOST ALERT: R25-BETA2012;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1462391247] SERVICE ALERT: R25-EPSILON2012;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462391266] HOST ALERT: R25-EPSILON2012;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1462391269] SERVICE ALERT: R25-GAMMA2012;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462391270] SERVICE ALERT: R25-BETA2012;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462391270] SERVICE ALERT: R25-BETA2012;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462391274] HOST ALERT: R25-GAMMA2012;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1462391275] SERVICE ALERT: R25-EPSILON2012;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462391275] SERVICE ALERT: R25-EPSILON2012;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462391276] SERVICE ALERT: R25-BETA2012;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462391276] SERVICE ALERT: R25-GAMMA2012;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462391277] SERVICE ALERT: R25-BETA2012;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462391280] SERVICE ALERT: R25-EPSILON2012;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462391281] SERVICE ALERT: R25-BETA2012;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462391294] SERVICE ALERT: R25-EPSILON2012;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462391295] SERVICE ALERT: R25-GAMMA2012;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462391296] SERVICE ALERT: R25-BETA2012;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462391296] SERVICE ALERT: R25-EPSILON2012;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462391296] SERVICE ALERT: R25-GAMMA2012;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462391300] HOST ALERT: R25-BETA2012;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1462391304] SERVICE ALERT: R25-BETA2012;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462391305] HOST ALERT: R25-EPSILON2012;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1462391306] HOST ALERT: R25-GAMMA2012;DOWN;HARD;10;(Host check timed out after 30.01 seconds)
[1462391315] SERVICE ALERT: R25-EPSILON2012;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462391323] SERVICE ALERT: R25-EPSILON2012;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462391323] SERVICE ALERT: R25-GAMMA2012;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462391331] SERVICE ALERT: R25-EPSILON2012;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462391332] SERVICE ALERT: R25-EPSILON2012;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462391333] SERVICE ALERT: R25-GAMMA2012;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462391334] HOST ALERT: R25-BETA2012;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1462391345] HOST ALERT: R25-EPSILON2012;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1462391356] SERVICE ALERT: R25-EPSILON2012;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462391356] SERVICE ALERT: R25-GAMMA2012;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462391357] SERVICE ALERT: R25-EPSILON2012;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462391358] SERVICE ALERT: R25-GAMMA2012;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462391363] SERVICE ALERT: R25-BETA2012;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462391367] SERVICE ALERT: R25-EPSILON2012;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462391371] SERVICE ALERT: R25-GAMMA2012;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462391386] HOST ALERT: R25-EPSILON2012;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1462391389] SERVICE ALERT: R25-GAMMA2012;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462391390] SERVICE ALERT: R25-BETA2012;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462391390] SERVICE ALERT: R25-BETA2012;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462391393] HOST ALERT: R25-BETA2012;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1462391395] SERVICE ALERT: R25-EPSILON2012;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462391396] SERVICE ALERT: R25-BETA2012;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462391396] SERVICE ALERT: R25-GAMMA2012;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462391397] SERVICE ALERT: R25-BETA2012;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462391414] SERVICE ALERT: R25-EPSILON2012;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462391415] SERVICE ALERT: R25-GAMMA2012;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462391416] SERVICE ALERT: R25-BETA2012;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462391416] SERVICE ALERT: R25-EPSILON2012;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462391416] SERVICE ALERT: R25-GAMMA2012;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462391424] SERVICE ALERT: R25-BETA2012;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462391425] HOST ALERT: R25-EPSILON2012;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1462391426] HOST ALERT: R25-BETA2012;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1462391433] SERVICE ALERT: R25-GAMMA2012;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462391433] SERVICE ALERT: R25-EPSILON2012;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462391443] SERVICE ALERT: R25-EPSILON2012;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462391443] SERVICE ALERT: R25-GAMMA2012;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462391452] SERVICE ALERT: R25-EPSILON2012;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462391475] SERVICE ALERT: R25-BETA2012;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462391476] SERVICE ALERT: R25-EPSILON2012;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462391476] SERVICE ALERT: R25-EPSILON2012;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462391476] SERVICE ALERT: R25-BETA2012;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462391477] SERVICE ALERT: R25-EPSILON2012;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462391478] SERVICE ALERT: R25-GAMMA2012;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462391483] SERVICE ALERT: R25-BETA2012;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462391491] SERVICE ALERT: R25-GAMMA2012;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462391493] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462391516] SERVICE ALERT: R25-BETA2012;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462391516] SERVICE ALERT: R25-GAMMA2012;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462391534] SERVICE ALERT: R25-EPSILON2012;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462391534] SERVICE ALERT: R25-GAMMA2012;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462391535] SERVICE ALERT: R25-BETA2012;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462391536] SERVICE ALERT: R25-BETA2012;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462391536] SERVICE ALERT: R25-EPSILON2012;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462391536] SERVICE ALERT: R25-GAMMA2012;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462391537] SERVICE ALERT: R25-BETA2012;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462391553] SERVICE ALERT: R25-GAMMA2012;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462391553] SERVICE ALERT: R25-EPSILON2012;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462391595] SERVICE ALERT: R25-BETA2012;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462391596] SERVICE ALERT: R25-EPSILON2012;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462391596] SERVICE ALERT: R25-BETA2012;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462391611] SERVICE ALERT: R25-GAMMA2012;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462391655] SERVICE ALERT: R25-BETA2012;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462391657] SERVICE ALERT: R25-BETA2012;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1462391673] SERVICE ALERT: R25-GAMMA2012;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462391673] SERVICE ALERT: R25-EPSILON2012;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462391715] SERVICE ALERT: R25-BETA2012;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462391716] SERVICE ALERT: R25-EPSILON2012;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462391716] SERVICE ALERT: R25-BETA2012;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462391775] SERVICE ALERT: R25-BETA2012;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462391777] SERVICE ALERT: R25-BETA2012;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1462391933] SERVICE ALERT: DB01-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1462392053] SERVICE ALERT: DB01-2012;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1462392173] SERVICE ALERT: DB01-2012;CPU Load;OK;SOFT;3;CPU Load 0% (15 min average)
[1462394123] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1462394243] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1462394493] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1462395093] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462395213] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462395333] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 91% (15 min average)
[1462395933] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462396043] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1462396053] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462396163] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 91% (15 min average)
[1462396173] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 58% (15 min average)
[1462396773] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462396893] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462397013] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1462397613] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 91% (15 min average)
[1462398213] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462398332] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 81% (15 min average)
[1462398932] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462399052] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462399172] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462402162] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1462402172] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 68% (15 min average)
[1462402282] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 91% (15 min average)
[1462402772] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462402882] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462402892] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 87% (15 min average)
[1462403002] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 90% (15 min average)
[1462403492] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462403612] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 89% (15 min average)
[1462404212] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462404332] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462404452] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462405402] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462405522] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 82% (15 min average)
[1462405613] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462405672] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462405732] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462406516] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1462406577] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 60% - OK!
[1462406932] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Ready
[1462409852] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 88% (15 min average)
[1462410452] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1462410572] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462410692] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1462411892] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462412492] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462413322] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1462413442] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1462414292] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462414892] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462415400] HOST DOWNTIME ALERT: Echo-01;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: ADAUTH-RMAD;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: ADDC3;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: FTK;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: HDBACKUPS;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: ADDC33;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: IMAGES;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: MBAM;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: ADDC4;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: MBAMDB;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: ADDC44;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: NESSUS;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: NETENG-PA;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: ADAUTHDC1-old;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: ADDC5;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: OFFICEKMS;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: OPIE;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: ADDC55;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: PNCAPP-2012;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: PNCDB-2012;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: LANSWEEPER;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: ADCA1;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: ADWSUSCLIENTS;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: ADAUTHDC1;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: PNCWEB-2012;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: PROCARE;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: ADWU1;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: R25-ALPHA2012;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: R25-DELTA2012;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: ADCA2;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: CBORD;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: R25-ZETA2012;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: AD-RMAD;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: REFWIN;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: CEDARON;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: SANDBOX-DB;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: SCCM-PRIM01;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: ADDC1;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: COE-WEBDAV;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: SP2010;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: SP2010DB;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: DIRECTACCESS;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: TABLEAU;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: TI22;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: ADDC2;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: DPS-SECURE02;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: TI22-COE;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: TI22-PSYCLINIC;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: EMU-VALUE01;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: UHSPHARM-2012;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: WINKMS;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: EMUSCRIPT;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: XDM;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: EMUSQLEXPRESS;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: ADARC;STARTED; Host has entered a period of scheduled downtime
[1462415400] HOST DOWNTIME ALERT: ADDC22;STARTED; Host has entered a period of scheduled downtime
[1462415492] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462416092] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462416442] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462416562] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1462416682] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462417882] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462418482] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1462418492] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462419082] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1462419082] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STARTED; Service appears to have started flapping (24.1% change >= 20.0% threshold)
[1462419092] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)