[1464753648] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1464753708] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1464755608] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464756465] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1464757065] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464758608] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 612 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464758728] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464758739] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464759339] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 528 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464759459] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464760059] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 519 found since last run.
[1464760064] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 81% (15 min average)
[1464760179] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464760664] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1464760779] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 501 found since last run.
[1464760783] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 75% (15 min average)
[1464760899] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464761383] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464761504] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464761623] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 85% (15 min average)
[1464762223] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464762343] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464762463] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 86% (15 min average)
[1464762699] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 512 found since last run.
[1464762819] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464763063] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464763183] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1464763303] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464764019] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 563 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464764139] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464765939] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 543 found since last run.
[1464766059] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464766303] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 82% (15 min average)
[1464766903] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464767022] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464767142] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464768459] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 541 found since last run.
[1464768579] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464768942] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 58% (15 min average)
[1464769179] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 549 found since last run.
[1464769299] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464769529] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 500 found since last run.
[1464769648] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464769899] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 588 found since last run.
[1464770019] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464770248] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 531 found since last run.
[1464770368] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464770619] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 636 found since last run.
[1464770739] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464770742] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1464770862] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 77% (15 min average)
[1464770968] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 509 found since last run.
[1464771088] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464771339] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 556 found since last run.
[1464771459] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464771462] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464771582] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1464771702] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 92% (15 min average)
[1464772059] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 614 found since last run.
[1464772179] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464772302] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1464772422] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 59% (15 min average)
[1464772779] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 632 found since last run.
[1464772888] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 504 found since last run.
[1464772899] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464773008] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464773499] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 670 found since last run.
[1464773608] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 513 found since last run.
[1464773619] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464773728] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464774219] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 524 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 657 found since last run.
[1464774222] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1464774339] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464774342] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 30% (15 min average)
[1464774928] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 528 found since last run.
[1464774939] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 569 found since last run.
[1464775048] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464775059] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464775648] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 510 found since last run.
[1464775660] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 502 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 638 found since last run.
[1464775768] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464775779] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464776142] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464776262] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464776368] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 533 found since last run.
[1464776378] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 558 found since last run.
[1464776382] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464776488] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464776498] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464777088] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 625 found since last run.
[1464777098] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1001 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 719 found since last run.
[1464777208] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464777218] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464777808] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 3844 found since last run. oneLevelSearchOps OK. searchOps is too high! 3868 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 3029 found since last run.
[1464777818] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 2183 found since last run. oneLevelSearchOps OK. searchOps is too high! 2188 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 3136 found since last run.
[1464777928] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1371 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 787 found since last run.
[1464777938] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 697 found since last run.
[1464778048] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 1892 found since last run. oneLevelSearchOps OK. searchOps is too high! 1910 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 1130 found since last run.
[1464778058] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 627 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 1658 found since last run.
[1464778182] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (15 min average)
[1464778782] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464778902] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464779022] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464780822] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 37% (15 min average)
[1464781422] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1464781542] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464781662] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1464781905] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1464781908] SERVICE ALERT: markjeffersonpk;Device Status;OK;HARD;3;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1464781965] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 51% - OK!
[1464782668] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1464782728] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1464782788] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1464783389] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 99% - OK!
[1464785262] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 59% (15 min average)
[1464786538] SERVICE ALERT: owen101lab1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1464786598] SERVICE ALERT: owen101lab1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1464786658] SERVICE ALERT: owen101lab1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1464787062] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1464787182] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1464787258] SERVICE ALERT: owen101lab1;Device Status;OK;HARD;3;Ready
[1464787948] SERVICE ALERT: NETIDADPROXY;Service: Windows Time;OK;HARD;3;W32Time: Started
[1464788382] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1464788502] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464788548] SERVICE ALERT: NETIDADPROXY;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1464788622] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1464788668] SERVICE ALERT: NETIDADPROXY;Service: Windows Time;CRITICAL;SOFT;2;W32Time: Stopped
[1464788788] SERVICE ALERT: NETIDADPROXY;Service: Windows Time;OK;SOFT;3;W32Time: Started
[1464791480] SERVICE ALERT: INTLTESTDB;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464791482] SERVICE ALERT: INTLTESTDB;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464791510] HOST ALERT: INTLTESTDB;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1464791568] SERVICE ALERT: INTLTESTDB;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464791568] SERVICE ALERT: INTLTESTDB;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464791598] HOST ALERT: INTLTESTDB;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1464791600] SERVICE ALERT: INTLTESTDB;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464791602] SERVICE ALERT: INTLTESTDB;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464791630] HOST ALERT: INTLTESTDB;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1464791647] SERVICE ALERT: INTLTESTDB;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464791651] SERVICE ALERT: INTLTESTDB;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464791677] SERVICE ALERT: INTLTESTDB;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464791677] HOST ALERT: INTLTESTDB;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1464791688] SERVICE ALERT: INTLTESTDB;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464791688] SERVICE ALERT: INTLTESTDB;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464791718] HOST ALERT: INTLTESTDB;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1464791720] SERVICE ALERT: INTLTESTDB;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464791722] SERVICE ALERT: INTLTESTDB;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464791748] SERVICE ALERT: INTLTESTDB;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464791748] SERVICE ALERT: INTLTESTDB;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464791750] HOST ALERT: INTLTESTDB;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1464791754] SERVICE ALERT: INTLTESTDB;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464791767] SERVICE ALERT: INTLTESTDB;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464791771] SERVICE ALERT: INTLTESTDB;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464791784] HOST ALERT: INTLTESTDB;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1464791796] SERVICE ALERT: INTLTESTDB;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464791808] SERVICE ALERT: INTLTESTDB;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464791808] SERVICE ALERT: INTLTESTDB;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464791815] SERVICE ALERT: INTLTESTDB;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464791823] SERVICE ALERT: INTLTESTDB;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464791826] HOST ALERT: INTLTESTDB;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1464791868] SERVICE ALERT: INTLTESTDB;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464791868] SERVICE ALERT: INTLTESTDB;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464791874] SERVICE ALERT: INTLTESTDB;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464791887] SERVICE ALERT: INTLTESTDB;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464791891] SERVICE ALERT: INTLTESTDB;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464791898] HOST ALERT: INTLTESTDB;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1464791916] SERVICE ALERT: INTLTESTDB;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464791928] SERVICE ALERT: INTLTESTDB;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464791928] SERVICE ALERT: INTLTESTDB;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464791935] SERVICE ALERT: INTLTESTDB;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464791943] SERVICE ALERT: INTLTESTDB;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464791946] HOST ALERT: INTLTESTDB;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1464791988] SERVICE ALERT: INTLTESTDB;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464791988] SERVICE ALERT: INTLTESTDB;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464791994] SERVICE ALERT: INTLTESTDB;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464792048] SERVICE ALERT: INTLTESTDB;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464792048] SERVICE ALERT: INTLTESTDB;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464792055] SERVICE ALERT: INTLTESTDB;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464792063] SERVICE ALERT: INTLTESTDB;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464792168] SERVICE ALERT: INTLTESTDB;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464792168] SERVICE ALERT: INTLTESTDB;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464792222] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 88% (15 min average)
[1464792448] SERVICE ALERT: mckennypk;Device Status;OK;HARD;3;BAD DUPLEXER CONNECTION 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1464794022] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464794142] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464794262] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464794868] SERVICE ALERT: INTLTESTDB;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1464794871] SERVICE ALERT: INTLTESTDB;Service: Windows Time;OK;HARD;3;W32Time: Started
[1464794896] SERVICE ALERT: INTLTESTDB;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 1 minute(s)
[1464794968] SERVICE ALERT: INTLTESTDB;VMWare Tools;OK;HARD;3;VMTools: Started
[1464794968] SERVICE ALERT: INTLTESTDB;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1464794974] SERVICE ALERT: INTLTESTDB;CPU Load;OK;HARD;3;CPU Load 51% (15 min average)
[1464795035] SERVICE ALERT: INTLTESTDB;Service: RPC;OK;HARD;3;RpcSs: Started
[1464795043] SERVICE ALERT: INTLTESTDB;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1464795148] SERVICE ALERT: INTLTESTDB;C:\ Drive Space;OK;HARD;3;c:\ - total: 84.22 Gb - used: 71.59 Gb (85%) - free 12.63 Gb (15%)
[1464795148] SERVICE ALERT: INTLTESTDB;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1464795162] HOST ALERT: INTLTESTDB;UP;HARD;10;PING OK - Packet loss = 66%, RTA = 0.30 ms
[1464795300] SERVICE ALERT: INTLTESTDB;Memory Usage;OK;HARD;3;Memory usage: total:14311.42 MB - used: 1329.00 MB (9%) - free: 12982.42 MB (91%)
[1464795302] SERVICE ALERT: INTLTESTDB;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1464795388] SERVICE ALERT: INTLTESTDB;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1464795388] SERVICE ALERT: INTLTESTDB;Service: WMI;OK;HARD;3;Winmgmt: Started
[1464795462] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 86% (15 min average)
[1464796062] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1464796182] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1464796302] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464796902] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1464797502] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1464797622] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464797742] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 92% (15 min average)
[1464798942] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1464799062] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464799182] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1464799371] SERVICE ALERT: MAPS01TEST;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.102.77 and port 12489: No route to host
[1464799374] HOST ALERT: MAPS01TEST;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.77)
[1464799437] HOST ALERT: MAPS01TEST;DOWN;SOFT;2;CRITICAL - Host Unreachable (164.76.102.77)
[1464799471] SERVICE ALERT: MAPS01TEST;Service: NetLogin;CRITICAL;SOFT;1;connect to address 164.76.102.77 and port 12489: No route to host
[1464799471] SERVICE ALERT: MAPS01TEST;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.102.77 and port 12489: No route to host
[1464799474] HOST ALERT: MAPS01TEST;DOWN;SOFT;3;CRITICAL - Host Unreachable (164.76.102.77)
[1464799482] SERVICE ALERT: MAPS01TEST;Service: DNS Client;CRITICAL;SOFT;2;connect to address 164.76.102.77 and port 12489: No route to host
[1464799509] SERVICE ALERT: MAPS01TEST;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 164.76.102.77 and port 12489: No route to host
[1464799509] SERVICE ALERT: MAPS01TEST;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.102.77 and port 12489: No route to host
[1464799512] HOST ALERT: MAPS01TEST;DOWN;SOFT;4;CRITICAL - Host Unreachable (164.76.102.77)
[1464799512] SERVICE ALERT: MAPS01TEST;VMWare Tools;CRITICAL;SOFT;1;connect to address 164.76.102.77 and port 12489: No route to host
[1464799515] SERVICE ALERT: MAPS01TEST;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.102.77 and port 12489: No route to host
[1464799535] SERVICE ALERT: MAPS01TEST;Service: User Profile;CRITICAL;SOFT;1;connect to address 164.76.102.77 and port 12489: No route to host
[1464799538] HOST ALERT: MAPS01TEST;DOWN;SOFT;5;CRITICAL - Host Unreachable (164.76.102.77)
[1464799552] SERVICE ALERT: MAPS01TEST;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.102.77 and port 12489: No route to host
[1464799552] SERVICE ALERT: MAPS01TEST;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.102.77 and port 12489: No route to host
[1464799555] SERVICE ALERT: MAPS01TEST;Uptime;CRITICAL;SOFT;1;connect to address 164.76.102.77 and port 12489: No route to host
[1464799555] HOST ALERT: MAPS01TEST;DOWN;SOFT;6;CRITICAL - Host Unreachable (164.76.102.77)
[1464799591] SERVICE ALERT: MAPS01TEST;Service: NetLogin;CRITICAL;SOFT;2;connect to address 164.76.102.77 and port 12489: No route to host
[1464799591] SERVICE ALERT: MAPS01TEST;Service: RPC;CRITICAL;SOFT;2;connect to address 164.76.102.77 and port 12489: No route to host
[1464799594] HOST ALERT: MAPS01TEST;DOWN;SOFT;7;CRITICAL - Host Unreachable (164.76.102.77)
[1464799602] SERVICE ALERT: MAPS01TEST;Service: DNS Client;CRITICAL;HARD;3;connect to address 164.76.102.77 and port 12489: No route to host
[1464799619] SERVICE ALERT: MAPS01TEST;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.102.77 and port 12489: No route to host
[1464799622] HOST ALERT: MAPS01TEST;DOWN;SOFT;8;CRITICAL - Host Unreachable (164.76.102.77)
[1464799629] SERVICE ALERT: MAPS01TEST;C:\ Drive Space;CRITICAL;SOFT;2;connect to address 164.76.102.77 and port 12489: No route to host
[1464799632] SERVICE ALERT: MAPS01TEST;VMWare Tools;CRITICAL;SOFT;2;connect to address 164.76.102.77 and port 12489: No route to host
[1464799637] SERVICE ALERT: MAPS01TEST;Service: Firewall;CRITICAL;SOFT;2;connect to address 164.76.102.77 and port 12489: No route to host
[1464799637] SERVICE ALERT: MAPS01TEST;Service: WMI;CRITICAL;SOFT;2;connect to address 164.76.102.77 and port 12489: No route to host
[1464799640] HOST ALERT: MAPS01TEST;DOWN;SOFT;9;CRITICAL - Host Unreachable (164.76.102.77)
[1464799655] SERVICE ALERT: MAPS01TEST;Service: User Profile;CRITICAL;SOFT;2;connect to address 164.76.102.77 and port 12489: No route to host
[1464799658] HOST ALERT: MAPS01TEST;DOWN;HARD;10;CRITICAL - Host Unreachable (164.76.102.77)
[1464799672] SERVICE ALERT: MAPS01TEST;Service: Group Policy;CRITICAL;SOFT;2;connect to address 164.76.102.77 and port 12489: No route to host
[1464799672] SERVICE ALERT: MAPS01TEST;NSClient++ Version;CRITICAL;SOFT;2;connect to address 164.76.102.77 and port 12489: No route to host
[1464799675] SERVICE ALERT: MAPS01TEST;Uptime;CRITICAL;SOFT;2;connect to address 164.76.102.77 and port 12489: No route to host
[1464799711] SERVICE ALERT: MAPS01TEST;Service: NetLogin;CRITICAL;HARD;3;connect to address 164.76.102.77 and port 12489: No route to host
[1464799711] SERVICE ALERT: MAPS01TEST;Service: RPC;CRITICAL;HARD;3;connect to address 164.76.102.77 and port 12489: No route to host
[1464799739] SERVICE ALERT: MAPS01TEST;Service: Windows Time;CRITICAL;SOFT;2;connect to address 164.76.102.77 and port 12489: No route to host
[1464799749] SERVICE ALERT: MAPS01TEST;C:\ Drive Space;CRITICAL;HARD;3;connect to address 164.76.102.77 and port 12489: No route to host
[1464799749] SERVICE ALERT: MAPS01TEST;Service: Firewall;CRITICAL;HARD;3;connect to address 164.76.102.77 and port 12489: No route to host
[1464799754] SERVICE ALERT: MAPS01TEST;VMWare Tools;CRITICAL;HARD;3;connect to address 164.76.102.77 and port 12489: No route to host
[1464799760] SERVICE ALERT: MAPS01TEST;Service: WMI;CRITICAL;HARD;3;connect to address 164.76.102.77 and port 12489: No route to host
[1464799775] SERVICE ALERT: MAPS01TEST;Service: User Profile;CRITICAL;HARD;3;connect to address 164.76.102.77 and port 12489: No route to host
[1464799782] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 75% (15 min average)
[1464799790] SERVICE ALERT: MAPS01TEST;NSClient++ Version;CRITICAL;HARD;3;connect to address 164.76.102.77 and port 12489: No route to host
[1464799790] SERVICE ALERT: MAPS01TEST;CPU Load;CRITICAL;SOFT;1;connect to address 164.76.102.77 and port 12489: No route to host
[1464799790] SERVICE ALERT: MAPS01TEST;Service: Group Policy;CRITICAL;HARD;3;connect to address 164.76.102.77 and port 12489: No route to host
[1464799797] SERVICE ALERT: MAPS01TEST;Uptime;CRITICAL;HARD;3;connect to address 164.76.102.77 and port 12489: No route to host
[1464799859] SERVICE ALERT: MAPS01TEST;Service: Windows Time;CRITICAL;HARD;3;connect to address 164.76.102.77 and port 12489: No route to host
[1464799891] SERVICE ALERT: MAPS01TEST;Memory Usage;CRITICAL;SOFT;1;connect to address 164.76.102.77 and port 12489: No route to host
[1464799910] SERVICE ALERT: MAPS01TEST;CPU Load;CRITICAL;SOFT;2;connect to address 164.76.102.77 and port 12489: No route to host
[1464800011] SERVICE ALERT: MAPS01TEST;Memory Usage;CRITICAL;SOFT;2;connect to address 164.76.102.77 and port 12489: No route to host
[1464800019] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1464800030] SERVICE ALERT: MAPS01TEST;CPU Load;CRITICAL;HARD;3;connect to address 164.76.102.77 and port 12489: No route to host
[1464800131] SERVICE ALERT: MAPS01TEST;Memory Usage;CRITICAL;HARD;3;connect to address 164.76.102.77 and port 12489: No route to host
[1464800505] HOST DOWNTIME ALERT: EVENTSTEST-2012;STARTED; Host has entered a period of scheduled downtime
[1464800508] SERVICE ALERT: EVENTSTEST-2012;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464800538] HOST ALERT: EVENTSTEST-2012;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1464800579] SERVICE ALERT: EVENTSTEST-2012;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464800609] HOST ALERT: EVENTSTEST-2012;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1464800616] SERVICE ALERT: EVENTSTEST-2012;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464800628] SERVICE ALERT: EVENTSTEST-2012;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464800646] HOST ALERT: EVENTSTEST-2012;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1464800699] SERVICE ALERT: EVENTSTEST-2012;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464800729] HOST ALERT: EVENTSTEST-2012;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1464800736] SERVICE ALERT: EVENTSTEST-2012;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464800748] SERVICE ALERT: EVENTSTEST-2012;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464800766] HOST ALERT: EVENTSTEST-2012;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1464800813] SERVICE ALERT: EVENTSTEST-2012;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464800819] SERVICE ALERT: EVENTSTEST-2012;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464800826] SERVICE ALERT: EVENTSTEST-2012;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464800832] SERVICE ALERT: EVENTSTEST-2012;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464800843] HOST ALERT: EVENTSTEST-2012;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1464800856] SERVICE ALERT: EVENTSTEST-2012;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464800886] HOST ALERT: EVENTSTEST-2012;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1464800910] SERVICE ALERT: EVENTSTEST-2012;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464800928] SERVICE ALERT: EVENTSTEST-2012;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464800933] SERVICE ALERT: EVENTSTEST-2012;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464800940] HOST ALERT: EVENTSTEST-2012;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1464800946] SERVICE ALERT: EVENTSTEST-2012;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464800952] SERVICE ALERT: EVENTSTEST-2012;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464800976] HOST ALERT: EVENTSTEST-2012;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1464800982] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464800986] SERVICE ALERT: EVENTSTEST-2012;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464801002] SERVICE ALERT: EVENTSTEST-2012;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464801002] SERVICE ALERT: EVENTSTEST-2012;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464801007] SERVICE ALERT: EVENTSTEST-2012;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464801007] SERVICE ALERT: EVENTSTEST-2012;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464801016] HOST ALERT: EVENTSTEST-2012;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1464801030] SERVICE ALERT: EVENTSTEST-2012;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464801035] SERVICE ALERT: EVENTSTEST-2012;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464801048] SERVICE ALERT: EVENTSTEST-2012;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464801053] SERVICE ALERT: EVENTSTEST-2012;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464801066] SERVICE ALERT: EVENTSTEST-2012;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464801072] SERVICE ALERT: EVENTSTEST-2012;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464801102] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464801106] SERVICE ALERT: EVENTSTEST-2012;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464801122] SERVICE ALERT: EVENTSTEST-2012;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464801122] SERVICE ALERT: EVENTSTEST-2012;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464801127] SERVICE ALERT: EVENTSTEST-2012;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464801127] SERVICE ALERT: EVENTSTEST-2012;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464801150] SERVICE ALERT: EVENTSTEST-2012;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464801155] SERVICE ALERT: EVENTSTEST-2012;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464801168] SERVICE ALERT: EVENTSTEST-2012;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464801220] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464801226] SERVICE ALERT: EVENTSTEST-2012;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464801242] SERVICE ALERT: EVENTSTEST-2012;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464801242] SERVICE ALERT: EVENTSTEST-2012;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464801247] SERVICE ALERT: EVENTSTEST-2012;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464801247] SERVICE ALERT: EVENTSTEST-2012;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464801275] SERVICE ALERT: EVENTSTEST-2012;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464802844] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (41.4% change >= 20.0% threshold)
[1464802844] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1464802844] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1464802844] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1464802844] HOST DOWNTIME ALERT: VCENTERSSO;STARTED; Host has entered a period of scheduled downtime
[1464802844] HOST DOWNTIME ALERT: EVENTSTEST-2012;STARTED; Host has entered a period of scheduled downtime
[1464802844] SERVICE DOWNTIME ALERT: NETIDADPROXY;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1464802844] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1464802844] HOST DOWNTIME ALERT: VC;STARTED; Host has entered a period of scheduled downtime
[1464802844] HOST DOWNTIME ALERT: VUM;STARTED; Host has entered a period of scheduled downtime
[1464802853] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (41.4% change >= 20.0% threshold)
[1464802854] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1464802854] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1464802854] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1464802854] HOST DOWNTIME ALERT: VCENTERSSO;STARTED; Host has entered a period of scheduled downtime
[1464802854] HOST DOWNTIME ALERT: EVENTSTEST-2012;STARTED; Host has entered a period of scheduled downtime
[1464802854] SERVICE DOWNTIME ALERT: NETIDADPROXY;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1464802854] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1464802854] HOST DOWNTIME ALERT: VC;STARTED; Host has entered a period of scheduled downtime
[1464802854] HOST DOWNTIME ALERT: VUM;STARTED; Host has entered a period of scheduled downtime
[1464802856] SERVICE ALERT: MAPS01TEST;Service: Windows Time;OK;HARD;3;W32Time: Started
[1464802860] HOST ALERT: MAPS01TEST;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.29 ms
[1464803019] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1464803027] SERVICE ALERT: MAPS01TEST;CPU Load;OK;HARD;3;CPU Load 0% (15 min average)
[1464803125] HOST ALERT: mckennypk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1464803128] SERVICE ALERT: MAPS01TEST;Memory Usage;OK;HARD;3;Memory usage: total:14311.42 MB - used: 1226.56 MB (9%) - free: 13084.86 MB (91%)
[1464803189] HOST ALERT: mckennypk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.94 ms
[1464803199] SERVICE ALERT: MAPS01TEST;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1464803249] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1464803308] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1464803308] SERVICE ALERT: MAPS01TEST;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1464803308] SERVICE ALERT: MAPS01TEST;Service: RPC;OK;HARD;3;RpcSs: Started
[1464803346] SERVICE ALERT: MAPS01TEST;C:\ Drive Space;OK;HARD;3;c:\ - total: 59.90 Gb - used: 33.67 Gb (56%) - free 26.23 Gb (44%)
[1464803346] SERVICE ALERT: MAPS01TEST;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1464803351] SERVICE ALERT: MAPS01TEST;VMWare Tools;OK;HARD;3;VMTools: Started
[1464803354] SERVICE ALERT: MAPS01TEST;Service: WMI;OK;HARD;3;Winmgmt: Started
[1464803368] SERVICE ALERT: mckennypk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1464803372] SERVICE ALERT: MAPS01TEST;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1464803389] SERVICE ALERT: MAPS01TEST;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1464803389] SERVICE ALERT: MAPS01TEST;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1464803394] SERVICE ALERT: MAPS01TEST;Uptime;OK;HARD;3;System Uptime - 18 day(s) 21 hour(s) 24 minute(s)
[1464803619] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464804003] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (51.7% change >= 20.0% threshold)
[1464804003] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1464804003] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1464804003] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1464804003] HOST DOWNTIME ALERT: VCENTERSSO;STARTED; Host has entered a period of scheduled downtime
[1464804003] HOST DOWNTIME ALERT: EVENTSTEST-2012;STARTED; Host has entered a period of scheduled downtime
[1464804003] SERVICE DOWNTIME ALERT: NETIDADPROXY;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1464804003] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1464804003] HOST DOWNTIME ALERT: VC;STARTED; Host has entered a period of scheduled downtime
[1464804003] HOST DOWNTIME ALERT: VUM;STARTED; Host has entered a period of scheduled downtime
[1464804130] SERVICE ALERT: EVENTSTEST-2012;Service: WMI;OK;HARD;3;Winmgmt: Started
[1464804134] HOST ALERT: EVENTSTEST-2012;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1464804139] SERVICE ALERT: EVENTSTEST-2012;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1464804148] SERVICE ALERT: EVENTSTEST-2012;Service: RPC;OK;HARD;3;RpcSs: Started
[1464804206] SERVICE ALERT: EVENTSTEST-2012;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1464804222] SERVICE ALERT: EVENTSTEST-2012;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1464804222] SERVICE ALERT: EVENTSTEST-2012;Uptime;OK;HARD;3;System Uptime - 18 day(s) 21 hour(s) 59 minute(s)
[1464804225] SERVICE ALERT: EVENTSTEST-2012;Service: Windows Time;OK;HARD;3;W32Time: Started
[1464804227] SERVICE ALERT: EVENTSTEST-2012;C:\ Drive Space;OK;HARD;3;c:\ - total: 39.66 Gb - used: 29.04 Gb (73%) - free 10.61 Gb (27%)
[1464804227] SERVICE ALERT: EVENTSTEST-2012;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1464804255] SERVICE ALERT: EVENTSTEST-2012;Memory Usage;OK;HARD;3;Memory usage: total:5503.60 MB - used: 1754.65 MB (32%) - free: 3748.95 MB (68%)
[1464804328] SERVICE ALERT: EVENTSTEST-2012;VMWare Tools;OK;HARD;3;VMTools: Started
[1464804389] SERVICE ALERT: EVENTSTEST-2012;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1464804399] SERVICE ALERT: EVENTSTEST-2012;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1464804436] SERVICE ALERT: EVENTSTEST-2012;CPU Load;OK;HARD;3;CPU Load 9% (15 min average)
[1464805114] SERVICE ALERT: halleground2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1464805144] HOST ALERT: halleground2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1464805154] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1464805174] SERVICE ALERT: halleground2;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1464805174] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1464805184] HOST ALERT: halleground2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1464805214] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1464805223] HOST ALERT: halleground2;UP;SOFT;3;PING OK - Packet loss = 50%, RTA = 0.55 ms
[1464805228] SERVICE ALERT: halleground2;Device Status;OK;SOFT;3;57.03 ERROR  To continue turn off then on 57.03 ERROR  46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1464805228] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 3% - OK!
[1464805269] SERVICE ALERT: halleground2;Paper Tray Status;OK;SOFT;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1464807218] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1464807756] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;SOFT;1;WARNING: Replication with ADDC5-OLD failed!
[1464807818] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1464807876] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;SOFT;2;WARNING: Replication with ADDC5-OLD failed!
[1464807938] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 92% (15 min average)
[1464807995] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;HARD;3;WARNING: Replication with ADDC5-OLD failed!
[1464808538] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464808658] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464808778] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 92% (15 min average)
[1464809978] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464810098] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464810218] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464810451] SERVICE ALERT: ADRODC-LIV;AD Replication;OK;HARD;3;OK: replication is up and running
[1464811051] HOST ALERT: halleground2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1464811141] HOST ALERT: halleground2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1464811205] HOST ALERT: halleground2;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1464811418] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 83% (15 min average)
[1464812018] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1464812138] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1464812258] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1464812858] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 59% (15 min average)
[1464814058] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464814178] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1464814298] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1464814898] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1464815498] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464815618] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1464816218] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1464816338] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1464816458] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464816887] SERVICE ALERT: hallefirst3;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1464816948] SERVICE ALERT: hallefirst3;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1464817007] SERVICE ALERT: hallefirst3;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1464817058] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1464817658] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464817778] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464817898] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464818125] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 48871 found since last run. oneLevelSearchOps OK. searchOps is too high! 48872 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464818245] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464826896] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 88% (15 min average)
[1464827496] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1464827616] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1464827735] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1464828334] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 78% (15 min average)
[1464828934] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1464829054] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464829174] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464831829] SERVICE ALERT: mckennypk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1464831859] HOST ALERT: mckennypk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1464831890] SERVICE ALERT: mckennypk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1464831920] HOST ALERT: mckennypk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1464831949] SERVICE ALERT: mckennypk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1464831979] HOST ALERT: mckennypk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1464831993] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1464832023] HOST ALERT: mckennypk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1464832053] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1464832083] HOST ALERT: mckennypk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1464832113] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1464832143] HOST ALERT: mckennypk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1464832173] SERVICE ALERT: mckennypk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1464832203] HOST ALERT: mckennypk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1464832293] HOST ALERT: mckennypk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1464832365] HOST ALERT: mckennypk;UP;SOFT;9;PING OK - Packet loss = 61%, RTA = 0.52 ms
[1464832579] HOST ALERT: mckennypk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1464832669] HOST ALERT: mckennypk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1464832743] HOST ALERT: mckennypk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1464832803] HOST ALERT: mckennypk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1464832893] HOST ALERT: mckennypk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1464832983] HOST ALERT: mckennypk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1464833073] HOST ALERT: mckennypk;DOWN;SOFT;7;(Host check timed out after 30.00 seconds)
[1464833163] HOST ALERT: mckennypk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1464833253] HOST ALERT: mckennypk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1464833343] HOST ALERT: mckennypk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1464833846] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464834446] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 718 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464834483] SERVICE ALERT: PNCDB-2012;D:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464834566] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464834583] SERVICE ALERT: PNCDB-2012;D:\ Drive Space;OK;SOFT;2;d:\ - total: 450.00 Gb - used: 115.01 Gb (26%) - free 334.99 Gb (74%)
[1464834600] HOST DOWNTIME ALERT: Echo-01;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: FTK;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: ADCA2;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: XDM;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: ADDC1;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: IMAGES;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: LANSWEEPER;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: ADDC2;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: MBAM;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: MBAMDB;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: ADDC3;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: NESSUS;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: NETENG-PA;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: AD-RMAD;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: ADDC4;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: OFFICEKMS;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: OPIE;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: HDBACKUPS;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: ADDC5;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: PNCAPP-2012;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: PNCDB-2012;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: ADARC;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: ADWSUSCLIENTS;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: PNCWEB-2012;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: PROCARE;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: ADWU1;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: R25-ALPHA2012;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: R25-DELTA2012;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: ADAUTH-RMAD;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: CBORD;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: R25-INTERFACE;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: R25-WEBSERVICES;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: CEDARON;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: R25-WEBVIEWER;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: R25-ZETA2012;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: ADAUTHDC1;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: COE-WEBDAV;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: REFWIN;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: SANDBOX-DB;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: DIRECTACCESS;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: SCCM-PRIM01;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: SP2010;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: ADAUTHDC1-OLD;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: DPS-SECURE02;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: SP2010DB;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: TABLEAU;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: EMU-VALUE01;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: TI22;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: TI22-COE;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: ADCA1;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: EMUSCRIPT;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: TI22-PSYCLINIC;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: UHSPHARM-2012;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: EMUSQLEXPRESS;STARTED; Host has entered a period of scheduled downtime
[1464834600] HOST DOWNTIME ALERT: WINKMS;STARTED; Host has entered a period of scheduled downtime
[1464835166] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 994 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464835286] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464835886] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 685 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464836006] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464836606] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 574 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464836726] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464837326] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 655 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464837446] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464838046] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 579 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464838166] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464838766] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 559 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464838886] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464839486] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 581 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464839606] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.