[1457416223] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 85% (5 min average)
[1457416824] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457416942] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1457417062] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1457424862] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 84% (5 min average)
[1457425462] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457425582] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1457425702] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1457426302] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 80% (5 min average)
[1457426902] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457427022] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1457427142] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1457428342] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 93% (5 min average)
[1457428942] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1457429540] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 58% (5 min average)
[1457430740] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1457430860] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 97% (5 min average)
[1457430980] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 70% (5 min average)
[1457431580] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1457431701] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1457431820] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1457432420] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 77% (5 min average)
[1457433020] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1457433140] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1457433260] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 93% (5 min average)
[1457433860] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 97% (5 min average)
[1457439490] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;SOFT;1;WARNING: Replication with ADDC5 failed!
[1457439610] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;SOFT;2;WARNING: Replication with ADDC5 failed!
[1457439730] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;HARD;3;WARNING: Replication with ADDC5 failed!
[1457440459] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 85% (5 min average)
[1457441254] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;OK;HARD;3;Magenta Cartridge HP CE263A is at 100% - OK!
[1457441659] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457441708] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1457441768] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1457441779] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1457441828] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1457441899] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1457442053] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1457442112] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1457442135] SERVICE ALERT: ADRODC-LIV;AD Replication;OK;HARD;3;OK: replication is up and running
[1457442172] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1457442428] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1457442772] SERVICE ALERT: markjeffersonpk;Device Status;OK;HARD;3;TRAY 3 OPEN  46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Processing job from tray 2
[1457444017] SERVICE ALERT: ADWU1;Windows Update Services Event 364;UNKNOWN;SOFT;1;Failed to process command : EvtQuery failed: 1717: The interface is unknown.
[1457444137] SERVICE ALERT: ADWU1;Windows Update Services Event 364;OK;SOFT;2;OK: No entries found
[1457444274] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 93% (5 min average)
[1457444874] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1457446172] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1457447373] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1457448423] SERVICE ALERT: MBAM;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457448474] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 51% (5 min average)
[1457448543] SERVICE ALERT: MBAM;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1457449074] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457449194] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 93% (5 min average)
[1457449314] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1457449314] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (32.0% change >= 20.0% threshold)
[1457449376] SERVICE ALERT: HDBACKUPS;Service: Group Policy;CRITICAL;SOFT;1;(Return code of 13 is out of bounds)
[1457449376] SERVICE ALERT: SANDBOX-DB;CPU Load;CRITICAL;SOFT;1;(Return code of 13 is out of bounds)
[1457449376] HOST ALERT: HDBACKUPS;DOWN;SOFT;1;(Return code of 13 is out of bounds)
[1457449376] HOST ALERT: SANDBOX-DB;DOWN;SOFT;1;(Return code of 13 is out of bounds)
[1457449376] SERVICE ALERT: ADCA2;Uptime;CRITICAL;SOFT;1;(Return code of 13 is out of bounds)
[1457449376] HOST ALERT: ADCA2;DOWN;SOFT;1;(Return code of 13 is out of bounds)
[1457449376] SERVICE ALERT: hallegroundcolor;Device Status;CRITICAL;SOFT;1;(Return code of 13 is out of bounds)
[1457449376] HOST ALERT: hallegroundcolor;DOWN;SOFT;1;(Return code of 13 is out of bounds)
[1457449377] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (32.0% change >= 20.0% threshold)
[1457449398] HOST ALERT: hallegroundcolor;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.40 ms
[1457449401] HOST ALERT: HDBACKUPS;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1457449421] HOST ALERT: ADCA2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.41 ms
[1457449423] HOST ALERT: SANDBOX-DB;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.40 ms
[1457449436] SERVICE ALERT: hallegroundcolor;Device Status;OK;SOFT;2;Sleep mode on
[1457449496] SERVICE ALERT: SANDBOX-DB;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1457449496] SERVICE ALERT: HDBACKUPS;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1457449496] SERVICE ALERT: ADCA2;Uptime;OK;SOFT;2;System Uptime - 3 day(s) 15 hour(s) 44 minute(s)
[1457449896] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (32.0% change >= 20.0% threshold)
[1457450527] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (30.4% change >= 20.0% threshold)
[1457456547] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1457456607] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1457456666] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1457457543] SERVICE ALERT: MBAM;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457457663] SERVICE ALERT: MBAM;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1457457866] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1457458185] SERVICE ALERT: ADCA2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457458305] SERVICE ALERT: ADCA2;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1457459293] SERVICE ALERT: rooseveltpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1457459352] SERVICE ALERT: rooseveltpk;Device Status;OK;SOFT;2;TRAY 3 EMPTY PLAIN LETTER   Processing job from tray 2
[1457459952] SERVICE ALERT: rooseveltpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1457460012] SERVICE ALERT: rooseveltpk;Device Status;OK;SOFT;2;LOAD TRAY 3 PLAIN LETTER To use another tray press OK LOAD TRAY 3 PLAIN LETTER 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1457460172] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1457460232] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1457460292] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1457460892] SERVICE ALERT: markjeffersonpk;Device Status;OK;HARD;3;50.1 FUSER ERROR  46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1457461108] SERVICE ALERT: owen101lab3;Device Status;OK;HARD;3;Processing job
[1457461112] HOST ALERT: owen101lab3;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.62 ms
[1457461449] SERVICE ALERT: owen101lab3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1457461624] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 98% - OK!
[1457462464] SERVICE ALERT: MBAM;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457462583] SERVICE ALERT: MBAM;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1457462699] SERVICE ALERT: markjeffersonpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1457462729] HOST ALERT: markjeffersonpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1457462752] SERVICE ALERT: markjeffersonpk;Device Status;OK;SOFT;2;50.1 FUSER ERROR  46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1457462756] HOST ALERT: markjeffersonpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.57 ms
[1457462908] SERVICE ALERT: ADCA1;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457463027] SERVICE ALERT: ADCA1;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1457463521] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 199.90 Gb - used: 160.50 Gb (80%) - free 39.40 Gb (20%)
[1457463641] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 199.90 Gb - used: 160.51 Gb (80%) - free 39.39 Gb (20%)
[1457463694] HOST ALERT: markjeffersonpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1457463712] SERVICE ALERT: INTLAPP-ISSMAP;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 63.90 Gb - used: 54.97 Gb (86%) - free 8.93 Gb (14%)
[1457463745] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1457463758] HOST ALERT: markjeffersonpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.68 ms
[1457463758] HOST FLAPPING ALERT: markjeffersonpk;STARTED; Host appears to have started flapping (22.2% change > 20.0% threshold)
[1457463761] SERVICE ALERT: BOEAPP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 199.90 Gb - used: 160.58 Gb (80%) - free 39.32 Gb (20%)
[1457463805] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1457463832] SERVICE ALERT: INTLAPP-ISSMAP;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 63.90 Gb - used: 54.97 Gb (86%) - free 8.93 Gb (14%)
[1457463865] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;SOFT;3;Black Print Cartridge HP C8543X is at 29% - OK!
[1457463952] SERVICE ALERT: INTLAPP-ISSMAP;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 63.90 Gb - used: 54.97 Gb (86%) - free 8.93 Gb (14%)
[1457464443] HOST ALERT: hallefirst2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1457464471] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1457464501] HOST ALERT: hallefirst2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1457464531] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1457464561] HOST ALERT: hallefirst2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1457464586] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1457464591] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1457464616] HOST ALERT: hallefirst2;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1457464639] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1457464646] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1457464669] HOST ALERT: hallefirst2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1457464699] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1457464706] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1457464721] SERVICE ALERT: ADDC5;GPO Replication;WARNING;SOFT;1;Warning! 1 policies are not replicating properly!
[1457464729] HOST ALERT: hallefirst2;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1457464759] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1457464789] HOST ALERT: hallefirst2;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1457464879] HOST ALERT: hallefirst2;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1457464969] HOST ALERT: hallefirst2;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1457465031] HOST ALERT: BBTS;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1457465034] HOST ALERT: BBTS_DB;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1457465059] HOST ALERT: hallefirst2;DOWN;SOFT;2;(Host check timed out after 30.00 seconds)
[1457465069] HOST ALERT: BBTS;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1457465070] HOST ALERT: BBTS_DB;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1457465114] SERVICE ALERT: ADDC5;GPO Replication;WARNING;SOFT;2;Warning! 1 policies are not replicating properly!
[1457465116] HOST ALERT: BBTS_DB;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1457465116] HOST ALERT: BBTS;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1457465149] HOST ALERT: hallefirst2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1457465163] HOST ALERT: BBTS;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1457465163] HOST ALERT: BBTS_DB;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1457465209] HOST ALERT: BBTS;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1457465211] HOST ALERT: BBTS_DB;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1457465221] HOST ALERT: hallefirst2;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1457465256] HOST ALERT: BBTS;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1457465257] HOST ALERT: BBTS_DB;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1457465274] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 84.90 Gb - used: 75.35 Gb (89%) - free 9.55 Gb (11%)
[1457465303] HOST ALERT: BBTS;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1457465304] HOST ALERT: BBTS_DB;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1457465310] HOST ALERT: hallefirst2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1457465350] HOST ALERT: BBTS;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1457465350] HOST ALERT: BBTS_DB;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1457465365] HOST ALERT: markjeffersonpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1457465390] HOST ALERT: hallefirst2;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1457465394] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 84.90 Gb - used: 75.35 Gb (89%) - free 9.55 Gb (11%)
[1457465397] HOST ALERT: BBTS_DB;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1457465397] HOST ALERT: BBTS;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1457465410] SERVICE ALERT: ADDC5;GPO Replication;WARNING;HARD;3;Warning! 1 policies are not replicating properly!
[1457465429] HOST ALERT: markjeffersonpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.71 ms
[1457465429] HOST FLAPPING ALERT: markjeffersonpk;STARTED; Host appears to have started flapping (20.8% change > 20.0% threshold)
[1457465444] HOST ALERT: BBTS;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1457465444] HOST ALERT: BBTS_DB;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1457465479] HOST ALERT: hallefirst2;DOWN;SOFT;7;(Host check timed out after 30.00 seconds)
[1457465514] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 84.90 Gb - used: 75.35 Gb (89%) - free 9.55 Gb (11%)
[1457465569] HOST ALERT: hallefirst2;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1457465659] HOST ALERT: hallefirst2;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1457465749] HOST ALERT: hallefirst2;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1457465995] SERVICE ALERT: halleground1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1457466055] SERVICE ALERT: halleground1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1457466115] SERVICE ALERT: halleground1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1457466715] SERVICE ALERT: halleground1;Device Status;OK;HARD;3;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1457466769] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1457466829] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 100% - OK!
[1457467216] SERVICE ALERT: INTLTESTAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457467226] SERVICE ALERT: ADCA1;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457467336] SERVICE ALERT: INTLTESTAPP;CPU Load;OK;SOFT;2;CPU Load 3% (5 min average)
[1457467346] SERVICE ALERT: ADCA1;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1457467557] SERVICE ALERT: markjeffersonpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1457467583] HOST ALERT: markjeffersonpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1457467617] SERVICE ALERT: markjeffersonpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1457467633] SERVICE ALERT: markjeffersonpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1457467647] HOST ALERT: markjeffersonpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1457467677] SERVICE ALERT: markjeffersonpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1457467686] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1457467693] SERVICE ALERT: markjeffersonpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1457467707] HOST ALERT: markjeffersonpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1457467746] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1457467753] SERVICE ALERT: markjeffersonpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1457467776] HOST ALERT: markjeffersonpk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1457467806] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1457467836] HOST ALERT: markjeffersonpk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1457467926] HOST ALERT: markjeffersonpk;DOWN;SOFT;6;(Host check timed out after 30.00 seconds)
[1457468016] HOST ALERT: markjeffersonpk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1457468099] HOST ALERT: markjeffersonpk;UP;SOFT;8;PING OK - Packet loss = 79%, RTA = 0.73 ms
[1457468271] SERVICE ALERT: markjeffersonpk;Device Status;OK;HARD;3;Processing job from tray 3
[1457468271] SERVICE FLAPPING ALERT: markjeffersonpk;Device Status;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1457468348] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1457468400] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 21% - OK!
[1457468546] SERVICE ALERT: ADCA1;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1457468665] SERVICE ALERT: ADCA1;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1457468692] HOST ALERT: hallefirst2;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1457468785] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 29% - OK!
[1457468900] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1457468952] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1457469321] SERVICE FLAPPING ALERT: markjeffersonpk;Device Status;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1457473226] SERVICE ALERT: strongpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1457473285] SERVICE ALERT: strongpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1457473346] SERVICE ALERT: strongpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1457474347] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1457474407] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1457474467] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1457475146] SERVICE ALERT: strongpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1457476267] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1457476267] SERVICE FLAPPING ALERT: markjeffersonpk;Paper Tray Status;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1457476671] SERVICE ALERT: ADDC5;GPO Replication;OK;HARD;3;OK! All GPOs are properly replicated.
[1457476867] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1457476927] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1457476987] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1457477592] SERVICE ALERT: markjeffersonpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1457477622] HOST ALERT: markjeffersonpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1457477712] HOST ALERT: markjeffersonpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1457477802] HOST ALERT: markjeffersonpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1457477866] HOST ALERT: markjeffersonpk;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.65 ms
[1457478186] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1457489093] SERVICE ALERT: DB01-2012;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1457489213] SERVICE ALERT: DB01-2012;CPU Load;OK;SOFT;2;CPU Load 32% (5 min average)
[1457489585] SERVICE FLAPPING ALERT: markjeffersonpk;Paper Tray Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1457489695] SERVICE ALERT: SANDBOX-DB;CPU Load;WARNING;SOFT;1;CPU Load 91% (5 min average)
[1457489815] SERVICE ALERT: SANDBOX-DB;CPU Load;OK;SOFT;2;CPU Load 83% (5 min average)
[1457490762] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 90% (5 min average)
[1457491362] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1457497166] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1457497286] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 87% (5 min average)
[1457498953] SERVICE ALERT: hallefirst2;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1457499013] SERVICE ALERT: hallefirst2;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1457499073] SERVICE ALERT: hallefirst2;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty