[1459397463] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459397523] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459397584] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459398183] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Open tray indicated.
[1459398600] HOST DOWNTIME ALERT: ADWU1;STOPPED; Host has exited from a period of scheduled downtime
[1459398600] HOST DOWNTIME ALERT: PNCDB-2012;STOPPED; Host has exited from a period of scheduled downtime
[1459398600] HOST DOWNTIME ALERT: CBORD;STOPPED; Host has exited from a period of scheduled downtime
[1459398600] HOST DOWNTIME ALERT: PNCAPP-2012;STOPPED; Host has exited from a period of scheduled downtime
[1459398600] HOST DOWNTIME ALERT: OPIE;STOPPED; Host has exited from a period of scheduled downtime
[1459398600] HOST DOWNTIME ALERT: OFFICEKMS;STOPPED; Host has exited from a period of scheduled downtime
[1459398600] HOST DOWNTIME ALERT: IMAGES;STOPPED; Host has exited from a period of scheduled downtime
[1459398600] HOST DOWNTIME ALERT: MBAM;STOPPED; Host has exited from a period of scheduled downtime
[1459398600] HOST DOWNTIME ALERT: XDM;STOPPED; Host has exited from a period of scheduled downtime
[1459398600] HOST DOWNTIME ALERT: UHSPHARM-2012;STOPPED; Host has exited from a period of scheduled downtime
[1459398600] HOST DOWNTIME ALERT: EMUSCRIPT;STOPPED; Host has exited from a period of scheduled downtime
[1459398600] HOST DOWNTIME ALERT: TABLEAU;STOPPED; Host has exited from a period of scheduled downtime
[1459398600] HOST DOWNTIME ALERT: DPS-SECURE02;STOPPED; Host has exited from a period of scheduled downtime
[1459398600] HOST DOWNTIME ALERT: SANDBOX-DB;STOPPED; Host has exited from a period of scheduled downtime
[1459398600] HOST DOWNTIME ALERT: EMUSQLEXPRESS;STOPPED; Host has exited from a period of scheduled downtime
[1459398600] HOST DOWNTIME ALERT: EMU-VALUE01;STOPPED; Host has exited from a period of scheduled downtime
[1459398600] HOST DOWNTIME ALERT: TI22-PSYCLINIC;STOPPED; Host has exited from a period of scheduled downtime
[1459398600] HOST DOWNTIME ALERT: ADAUTH-RMAD;STOPPED; Host has exited from a period of scheduled downtime
[1459398600] HOST DOWNTIME ALERT: PNCWEB-2012;STOPPED; Host has exited from a period of scheduled downtime
[1459398600] HOST DOWNTIME ALERT: AD-RMAD;STOPPED; Host has exited from a period of scheduled downtime
[1459398600] HOST DOWNTIME ALERT: WINKMS;STOPPED; Host has exited from a period of scheduled downtime
[1459398600] HOST DOWNTIME ALERT: PROCARE;STOPPED; Host has exited from a period of scheduled downtime
[1459398600] HOST DOWNTIME ALERT: COE-WEBDAV;STOPPED; Host has exited from a period of scheduled downtime
[1459398600] HOST DOWNTIME ALERT: DIRECTACCESS;STOPPED; Host has exited from a period of scheduled downtime
[1459398600] HOST DOWNTIME ALERT: NETENG-PA;STOPPED; Host has exited from a period of scheduled downtime
[1459398600] HOST DOWNTIME ALERT: TI22;STOPPED; Host has exited from a period of scheduled downtime
[1459398600] HOST DOWNTIME ALERT: ADWSUSCLIENTS;STOPPED; Host has exited from a period of scheduled downtime
[1459398600] HOST DOWNTIME ALERT: CEDARON;STOPPED; Host has exited from a period of scheduled downtime
[1459398600] HOST DOWNTIME ALERT: FTK;STOPPED; Host has exited from a period of scheduled downtime
[1459398600] HOST DOWNTIME ALERT: LANSWEEPER;STOPPED; Host has exited from a period of scheduled downtime
[1459398600] HOST DOWNTIME ALERT: Echo-01;STOPPED; Host has exited from a period of scheduled downtime
[1459398600] HOST DOWNTIME ALERT: HDBACKUPS;STOPPED; Host has exited from a period of scheduled downtime
[1459398600] HOST DOWNTIME ALERT: MBAMDB;STOPPED; Host has exited from a period of scheduled downtime
[1459398600] HOST DOWNTIME ALERT: REFWIN;STOPPED; Host has exited from a period of scheduled downtime
[1459398600] HOST DOWNTIME ALERT: NESSUS;STOPPED; Host has exited from a period of scheduled downtime
[1459398783] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459398843] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459398903] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459399503] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Open tray indicated.
[1459400703] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459400763] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459400823] SERVICE ALERT: sillpk;Device Status;OK;SOFT;3;Open tray indicated.
[1459401733] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1459402023] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459402083] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459402143] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459402743] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Open tray indicated.
[1459403244] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1459403364] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 95% (5 min average)
[1459403484] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 83% (5 min average)
[1459403944] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459404003] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459404063] SERVICE ALERT: sillpk;Device Status;OK;SOFT;3;Open tray indicated.
[1459404091] SERVICE ALERT: XDM;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459404211] SERVICE ALERT: XDM;CPU Load;OK;SOFT;2;CPU Load 2% (5 min average)
[1459404921] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 84.90 Gb - used: 79.21 Gb (93%) - free 5.69 Gb (7%)
[1459405041] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 84.90 Gb - used: 79.21 Gb (93%) - free 5.69 Gb (7%)
[1459405161] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 84.90 Gb - used: 79.20 Gb (93%) - free 5.70 Gb (7%)
[1459405263] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459405323] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459405383] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459405884] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1459405983] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Open tray indicated.
[1459406004] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 86% (5 min average)
[1459407183] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459407243] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459407303] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459407861] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1459407903] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Open tray indicated.
[1459407981] SERVICE ALERT: BOEAPPTEST;CPU Load;OK;SOFT;2;CPU Load 21% (5 min average)
[1459409103] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459409163] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459409224] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459409823] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Open tray indicated.
[1459410423] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459410483] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459410543] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459410733] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 65% (5 min average)
[1459411143] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Open tray indicated.
[1459411333] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1459411453] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1459411573] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1459412343] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459412403] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459412464] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459413063] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Open tray indicated.
[1459413373] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 86% (5 min average)
[1459413663] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459413723] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459413783] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459413973] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1459414093] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 97% (5 min average)
[1459414213] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1459414213] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1459414383] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Open tray indicated.
[1459414813] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 79% (5 min average)
[1459415413] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1459415533] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1459415583] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459415644] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459415653] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 84% (5 min average)
[1459415703] SERVICE ALERT: sillpk;Device Status;OK;SOFT;3;Open tray indicated.
[1459415898] SERVICE ALERT: NETIDADPROXY;Service: DNS Client;CRITICAL;SOFT;1;Dnscache: Stopped
[1459416017] SERVICE ALERT: NETIDADPROXY;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1459416853] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459416903] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459416963] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459416973] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1459417023] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459417093] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 67% (5 min average)
[1459417623] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Open tray indicated.
[1459418824] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459418883] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459418893] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459418943] SERVICE ALERT: sillpk;Device Status;OK;SOFT;3;Open tray indicated.
[1459419013] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1459419133] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1459420143] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459420203] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459420263] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459420333] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 69% (5 min average)
[1459420863] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Open tray indicated.
[1459421533] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459421653] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1459421773] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1459422063] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459422123] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459422183] SERVICE ALERT: sillpk;Device Status;OK;SOFT;3;Open tray indicated.
[1459423383] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459423443] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459423503] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459424103] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Open tray indicated.
[1459424173] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 87% (5 min average)
[1459424309] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 1 Trays Empty
[1459424370] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 1 Trays Empty
[1459424429] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 1 Trays Empty
[1459424524] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1459424773] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1459424893] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1459425013] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1459425303] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459425363] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459425423] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459425499] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459425559] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459425589] HOST ALERT: prayharroldpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459425619] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459425637] HOST ALERT: prayharroldpk;UP;SOFT;2;PING OK - Packet loss = 73%, RTA = 0.65 ms
[1459426023] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Open tray indicated.
[1459426143] SERVICE ALERT: ADRODC-LIV;Service: Group Policy;CRITICAL;SOFT;1;gpsvc: Stopped
[1459426213] SERVICE ALERT: prayharroldpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459426283] SERVICE ALERT: ADRODC-LIV;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1459426313] HOST ALERT: ADRODC-LIV;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459426323] SERVICE ALERT: ADRODC-LIV;Service: MDT Monitor;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459426347] SERVICE ALERT: ADRODC-LIV;Service: Intersite Messaging;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459426353] HOST ALERT: ADRODC-LIV;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1459426361] SERVICE ALERT: ADRODC-LIV;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 207.74.163.140 and port 12489: Connection refused
[1459426364] SERVICE ALERT: ADRODC-LIV;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459426378] HOST ALERT: ADRODC-LIV;UP;SOFT;3;PING OK - Packet loss = 72%, RTA = 6.77 ms
[1459426392] SERVICE ALERT: ADRODC-LIV;Service: Group Policy;OK;SOFT;3;gpsvc: Started
[1459426423] SERVICE ALERT: ADRODC-LIV;Service: MDT Monitor;OK;SOFT;2;MDT_Monitor: Started
[1459426447] SERVICE ALERT: ADRODC-LIV;Service: Intersite Messaging;OK;SOFT;2;IsmServ: Started
[1459426464] SERVICE ALERT: ADRODC-LIV;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1459426478] SERVICE ALERT: ADRODC-LIV;C:\ Drive Space;OK;SOFT;2;c:\ - total: 1115.16 Gb - used: 557.59 Gb (50%) - free 557.56 Gb (50%)
[1459426579] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459426623] SERVICE ALERT: INTLAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459426623] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459426683] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459426699] SERVICE ALERT: BOEAPPTEST;CPU Load;OK;SOFT;2;CPU Load 20% (5 min average)
[1459426743] SERVICE ALERT: INTLAPP;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1459426743] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459426794] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459426854] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459426914] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459427343] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Open tray indicated.
[1459427358] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 84.90 Gb - used: 80.72 Gb (95%) - free 4.18 Gb (5%)
[1459427413] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1459427413] SERVICE ALERT: prayharroldpk;Device Status;OK;HARD;3;Ready
[1459427428] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1459427524] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1459427584] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1459427644] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1459427733] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1459427793] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1459427853] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1459428013] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 52% (5 min average)
[1459428114] SERVICE ALERT: markjeffersonpk;Device Status;OK;HARD;3;66.00.15 EXTERNAL DEVICE FAILURE Turn off, check connection, turn on 66.00.15 EXTERNAL DEVICE FAILURE 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1459428244] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1459428453] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;CRITICAL;HARD;3;Black Print Cartridge HP C8543X is at 1% - CRITICAL!
[1459428543] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459428603] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459428613] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1459428663] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459428733] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 83% (5 min average)
[1459428980] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459429040] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459429101] SERVICE ALERT: mckennypk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459429263] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Open tray indicated.
[1459429653] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1459429700] SERVICE ALERT: mckennypk;Device Status;OK;HARD;3;ORDER CARTRIDGE LESS THAN 1200 PAGES 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1459429819] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459429869] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459429879] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459429911] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459429929] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459429933] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1459429939] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459429939] SERVICE FLAPPING ALERT: prayharroldpk;Device Status;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1459429971] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459429989] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459430031] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459430053] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1459430173] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 97% (5 min average)
[1459430463] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459430523] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459430584] SERVICE ALERT: sillpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459430773] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 36% (5 min average)
[1459430965] HOST ALERT: prayharroldpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459431029] HOST ALERT: prayharroldpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.68 ms
[1459431133] SERVICE ALERT: prayharroldpk;Device Status;OK;HARD;3;MENUS  RETRIEVE JOB  INFORMATION  PAPER HANDLING
[1459431183] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Open tray indicated.
[1459431184] SERVICE ALERT: prayharroldpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1459431225] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 21% - OK!
[1459431373] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459431493] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1459431499] SERVICE ALERT: hallefirst3;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1459431558] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 84.90 Gb - used: 80.32 Gb (95%) - free 4.58 Gb (5%)
[1459431559] SERVICE ALERT: hallefirst3;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1459431613] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1459431619] SERVICE ALERT: hallefirst3;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1459431619] SERVICE ALERT: hallefirst3;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1459431669] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1459431679] SERVICE ALERT: hallefirst3;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1459431729] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1459431739] SERVICE ALERT: hallefirst3;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1459431783] SERVICE ALERT: sillpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459431789] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1459431843] SERVICE ALERT: sillpk;Device Status;OK;SOFT;2;MENUS  SHOW ME HOW  RETRIEVE JOB  INFORMATION
[1459431969] SERVICE ALERT: alexanderpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.142.38. Is the printer off?  Is the network down?
[1459432023] SERVICE ALERT: alexanderpk;Device Status;OK;SOFT;2;TRAY 3 EMPTY ANY TYPE LETTER   Ready
[1459432143] SERVICE ALERT: C4SAS-SE;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459432176] SERVICE ALERT: rooseveltpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459432236] SERVICE ALERT: rooseveltpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459432263] SERVICE ALERT: C4SAS-SE;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1459432296] SERVICE ALERT: rooseveltpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459432314] SERVICE ALERT: VCENTERSSO;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1459432434] SERVICE ALERT: VCENTERSSO;CPU Load;OK;SOFT;2;CPU Load 32% (5 min average)
[1459432802] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (46.2% change >= 20.0% threshold)
[1459432802] SERVICE FLAPPING ALERT: prayharroldpk;Device Status;STARTED; Service appears to have started flapping (26.6% change >= 20.0% threshold)
[1459432802] SERVICE FLAPPING ALERT: sillpk;Device Status;STARTED; Service appears to have started flapping (59.5% change >= 20.0% threshold)
[1459432802] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1459432802] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1459432802] SERVICE DOWNTIME ALERT: BOEAPPTEST;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1459432813] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 66% (5 min average)
[1459432896] SERVICE ALERT: rooseveltpk;Device Status;OK;HARD;3;Open tray indicated.
[1459433413] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459433414] SERVICE ALERT: hallefirst3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1459433533] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1459433533] SERVICE ALERT: hallefirst3;Device Status;OK;HARD;3;ORDER CARTRIDGE LESS THAN 2200 PAGES 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Processing job from tray 4
[1459433583] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 7% - OK!
[1459433653] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1459433674] SERVICE ALERT: ADCA2;CPU Load;WARNING;SOFT;1;CPU Load 94% (5 min average)
[1459433794] SERVICE ALERT: ADCA2;CPU Load;OK;SOFT;2;CPU Load 49% (5 min average)
[1459434096] SERVICE ALERT: rooseveltpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459434155] SERVICE ALERT: rooseveltpk;Device Status;OK;SOFT;2;Open tray indicated.
[1459434695] SERVICE ALERT: TI22-COE;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1459434815] SERVICE ALERT: TI22-COE;CPU Load;CRITICAL;SOFT;2;CPU Load 96% (5 min average)
[1459434935] SERVICE ALERT: TI22-COE;CPU Load;OK;SOFT;3;CPU Load 39% (5 min average)
[1459435023] SERVICE ALERT: TI22-COE;Service: RPC;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1459435143] SERVICE ALERT: TI22-COE;Service: RPC;OK;SOFT;2;RpcSs: Started
[1459435158] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;OK;HARD;3;c:\ - total: 84.90 Gb - used: 75.88 Gb (89%) - free 9.02 Gb (11%)
[1459435453] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 86% (5 min average)
[1459435765] HOST ALERT: sillpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459435833] HOST ALERT: sillpk;UP;SOFT;2;PING OK - Packet loss = 44%, RTA = 0.46 ms
[1459436053] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1459436124] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;CRITICAL;HARD;3;Black Print Cartridge HP C8543X is at 1% - CRITICAL!
[1459436173] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1459436293] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1459436513] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459436555] SERVICE ALERT: rooseveltpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459436573] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459436615] SERVICE ALERT: rooseveltpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459436633] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459436676] SERVICE ALERT: rooseveltpk;Device Status;OK;SOFT;3;Open tray indicated.
[1459436735] SERVICE ALERT: TI22-COE;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1459436855] SERVICE ALERT: TI22-COE;CPU Load;OK;SOFT;2;CPU Load 7% (5 min average)
[1459437234] SERVICE ALERT: markjeffersonpk;Device Status;OK;HARD;3;Remove any jammed paper.
[1459437234] SERVICE FLAPPING ALERT: markjeffersonpk;Device Status;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1459437705] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;ORDER CARTRIDGE LESS THAN 1300 PAGES 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1459437709] HOST ALERT: hallefirst2;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.46 ms
[1459437969] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 5% - OK!
[1459438055] SERVICE ALERT: TI22-COE;CPU Load;WARNING;SOFT;1;CPU Load 94% (5 min average)
[1459438094] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1459438175] SERVICE ALERT: TI22-COE;CPU Load;OK;SOFT;2;CPU Load 70% (5 min average)
[1459438938] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459438989] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459438998] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459439031] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459439049] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459439058] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459439091] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459439109] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459439152] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459439331] SERVICE ALERT: mckennypk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1459439390] SERVICE ALERT: mckennypk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1459439450] SERVICE ALERT: mckennypk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1459439652] SERVICE ALERT: prayharroldpk;Device Status;OK;HARD;3;Ready
[1459439704] SERVICE ALERT: prayharroldpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1459439704] SERVICE FLAPPING ALERT: prayharroldpk;Paper Tray Status;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1459439745] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 21% - OK!
[1459439745] SERVICE FLAPPING ALERT: prayharroldpk;Black Print Cartridge HP C8543X;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1459441773] SERVICE ALERT: ADDC5;GPO Replication;UNKNOWN;SOFT;1;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1459441793] SERVICE ALERT: ADDC5;GPO Replication;OK;SOFT;2;Export-CSV : The process cannot access the file 'C:\Program
[1459442010] SERVICE ALERT: XDM;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1459442130] SERVICE ALERT: XDM;CPU Load;OK;SOFT;2;CPU Load 21% (5 min average)
[1459442356] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1459442416] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1459442475] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1459442644] SERVICE FLAPPING ALERT: sillpk;Device Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1459443650] SERVICE ALERT: mckennypk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1459444276] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1459444310] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459444370] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459444430] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459444533] SERVICE ALERT: halleground1;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1459444575] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459444593] SERVICE ALERT: halleground1;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1459444635] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459444653] SERVICE ALERT: halleground1;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1459444659] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1459444693] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 87% (5 min average)
[1459444695] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459444699] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459444719] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1459444759] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459444779] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1459444789] HOST ALERT: hallefirst2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459444809] HOST ALERT: hallefirst1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459444814] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;SOFT;3;Printer Paper Is OK! All Trays Still Full
[1459444818] HOST ALERT: hallefirst2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.52 ms
[1459444818] SERVICE ALERT: hallefirst3;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1459444848] HOST ALERT: hallefirst3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459444869] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1459444876] HOST ALERT: hallefirst1;UP;SOFT;2;PING OK - Packet loss = 37%, RTA = 0.59 ms
[1459444878] SERVICE ALERT: hallefirst3;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1459444907] HOST ALERT: hallefirst3;UP;SOFT;2;PING WARNING - Packet loss = 83%, RTA = 0.54 ms
[1459444929] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1459444938] SERVICE ALERT: hallefirst3;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1459444939] SERVICE ALERT: hallefirst3;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1459444959] HOST ALERT: hallefirst1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459444988] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1459444989] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1459444998] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1459444999] SERVICE ALERT: hallefirst3;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1459445019] HOST ALERT: hallefirst1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1459445048] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1459445058] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1459445059] SERVICE ALERT: hallefirst3;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1459445078] HOST ALERT: hallefirst3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459445088] HOST ALERT: hallefirst1;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1459445102] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;OK;SOFT;3;Black Print Cartridge HP C8543X is at 6% - OK!
[1459445106] HOST ALERT: hallefirst3;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.55 ms
[1459445106] HOST FLAPPING ALERT: hallefirst3;STARTED; Host appears to have started flapping (21.6% change > 20.0% threshold)
[1459445118] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1459445148] HOST ALERT: hallefirst1;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1459445229] HOST ALERT: hallefirst1;UP;SOFT;5;PING OK - Packet loss = 77%, RTA = 0.56 ms
[1459445229] HOST FLAPPING ALERT: hallefirst1;STARTED; Host appears to have started flapping (21.1% change > 20.0% threshold)
[1459445293] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1459445413] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1459445419] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459445436] HOST ALERT: hallefirst3;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1459445479] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459445500] HOST ALERT: hallefirst3;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1459445532] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 78% (5 min average)
[1459445539] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1459445708] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1459445738] HOST ALERT: hallefirst3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459445748] HOST ALERT: hallefirst1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459445768] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1459445775] HOST ALERT: hallefirst3;UP;SOFT;2;PING OK - Packet loss = 37%, RTA = 0.64 ms
[1459445812] HOST ALERT: hallefirst1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.57 ms
[1459445822] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;OK;SOFT;3;Black Print Cartridge HP C8543X is at 6% - OK!
[1459445889] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 4% - OK!
[1459445899] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459445973] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 80% - OK!
[1459446019] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1459446132] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459446133] SERVICE ALERT: hallefirst3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1459446133] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1459446139] SERVICE ALERT: BOEAPPTEST;CPU Load;OK;SOFT;3;CPU Load 21% (5 min average)
[1459446184] SERVICE ALERT: hallefirst1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1459446224] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;ORDER CARTRIDGE LESS THAN 1100 PAGES 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1459446252] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1459446253] SERVICE ALERT: hallefirst3;Device Status;OK;HARD;3;ORDER CARTRIDGE LESS THAN 1900 PAGES 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1459446312] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;Ready
[1459446372] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1459446594] SERVICE ALERT: TI22-COE;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459446675] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1459446694] SERVICE ALERT: TI22-COE;CPU Load;OK;SOFT;2;CPU Load 1% (5 min average)
[1459446736] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1459446796] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1459446972] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 92% (5 min average)
[1459447395] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1459447395] SERVICE FLAPPING ALERT: alexanderpk;Paper Tray Status;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1459447933] HOST FLAPPING ALERT: hallefirst3;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1459447939] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459448059] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1459448112] HOST FLAPPING ALERT: hallefirst1;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1459448173] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1459448179] SERVICE ALERT: BOEAPPTEST;CPU Load;OK;SOFT;3;CPU Load 23% (5 min average)
[1459448633] SERVICE FLAPPING ALERT: markjeffersonpk;Device Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1459448853] SERVICE ALERT: halleground1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1459449972] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 86% (5 min average)
[1459450228] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 1 Trays Empty
[1459450288] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 1 Trays Empty
[1459450349] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 1 Trays Empty
[1459450579] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1459450689] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1459450700] SERVICE ALERT: BOEAPPTEST;CPU Load;OK;SOFT;2;CPU Load 35% (5 min average)
[1459450748] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1459450804] SERVICE ALERT: IMAGES;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459450808] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1459450923] SERVICE ALERT: IMAGES;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1459451053] SERVICE FLAPPING ALERT: prayharroldpk;Device Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1459451104] SERVICE FLAPPING ALERT: prayharroldpk;Paper Tray Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1459451136] HOST ALERT: TABLEAU;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459451144] SERVICE FLAPPING ALERT: prayharroldpk;Black Print Cartridge HP C8543X;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1459451172] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1459451208] HOST ALERT: TABLEAU;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1459451228] HOST ALERT: TABLEAU;UP;SOFT;3;PING OK - Packet loss = 16%, RTA = 0.37 ms
[1459451292] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1459451314] SERVICE ALERT: TABLEAU;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1459451412] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1459451434] SERVICE ALERT: TABLEAU;CPU Load;OK;SOFT;2;CPU Load 26% (5 min average)
[1459451910] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459451970] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459452030] SERVICE ALERT: halleground2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459452748] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1459452925] SERVICE ALERT: EMUFILES;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459453045] SERVICE ALERT: EMUFILES;CPU Load;OK;SOFT;2;CPU Load 8% (5 min average)
[1459453424] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459453484] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459453544] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459453861] SERVICE ALERT: SP2010DB;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459453981] SERVICE ALERT: SP2010DB;CPU Load;OK;SOFT;2;CPU Load 4% (5 min average)
[1459454298] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1459454418] SERVICE ALERT: BOEAPPTEST;CPU Load;WARNING;SOFT;2;CPU Load 94% (5 min average)
[1459454430] SERVICE ALERT: halleground2;Device Status;OK;HARD;3;TRAY 2 OPEN  46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1459454538] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1459455137] SERVICE ALERT: BOEAPPTEST;CPU Load;OK;HARD;3;CPU Load 4% (5 min average)
[1459455258] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459455309] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459455318] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459455344] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Processing job from tray 4
[1459455344] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1459455350] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459455369] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459455378] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459455410] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459455429] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459455470] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459455973] SERVICE ALERT: prayharroldpk;Device Status;OK;HARD;3;Ready
[1459456023] SERVICE ALERT: prayharroldpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1459456065] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 21% - OK!
[1459456210] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 82% (5 min average)
[1459456432] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459456492] SERVICE ALERT: markjeffersonpk;Device Status;OK;SOFT;2;CLOSE RIGHT DOOR  46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1459456802] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1459456810] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1459456862] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 62% - OK!
[1459456930] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1459457050] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1459457092] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459457152] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459457213] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459457812] SERVICE ALERT: markjeffersonpk;Device Status;OK;HARD;3;Warming up
[1459458062] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1459458122] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1459458137] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1459458173] SERVICE ALERT: rooseveltpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1459458182] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;OK;SOFT;3;Black Print Cartridge HP C8543X is at 62% - OK!
[1459458232] SERVICE ALERT: rooseveltpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1459458257] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1459458293] SERVICE ALERT: rooseveltpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1459458377] SERVICE ALERT: BOEAPPTEST;CPU Load;OK;SOFT;3;CPU Load 8% (5 min average)
[1459458796] SERVICE FLAPPING ALERT: alexanderpk;Paper Tray Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1459458850] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 77% (5 min average)
[1459459450] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1459459570] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1459459690] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1459461244] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1459461304] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1459461364] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1459461490] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 63% (5 min average)
[1459462090] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1459462210] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1459462330] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1459463093] SERVICE ALERT: rooseveltpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1459464573] SERVICE ALERT: owen101labcolor;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 1 Trays Empty
[1459464615] SERVICE ALERT: NETIDADPROXY;Service: DNS Client;CRITICAL;SOFT;1;Dnscache: Stopped
[1459464632] SERVICE ALERT: owen101labcolor;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! All Trays Still Full
[1459464735] SERVICE ALERT: NETIDADPROXY;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1459465544] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1459465929] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (5 min average)
[1459467129] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459467249] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1459467370] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1459474154] SERVICE ALERT: DB01-2012;CPU Load;WARNING;SOFT;1;CPU Load 94% (5 min average)
[1459474273] SERVICE ALERT: DB01-2012;CPU Load;OK;SOFT;2;CPU Load 25% (5 min average)
[1459475602] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1459475722] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 87% (5 min average)
[1459478168] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1459479322] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1459479442] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 89% (5 min average)
[1459480042] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1459480162] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 96% (5 min average)
[1459480282] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 92% (5 min average)
[1459480882] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 83% (5 min average)