[1463631393] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463631993] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (15 min average)
[1463632200] HOST DOWNTIME ALERT: ADCA2;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: PNCDB-2012;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: PNCAPP-2012;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: MBAMDB;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: HDBACKUPS;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: ADAUTH-RMAD;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: ADDC33;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: EMU-VALUE01;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: ADAUTHDC1-OLD;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: UHSPHARM-2012;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: R25-ALPHA2012;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: ADDC1;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: ADDC22;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: TABLEAU;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: ADDC3;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: ADDC11;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: R25-ZETA2012;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: R25-WEBSERVICES;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: SP2010;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: XDM;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: EMUSCRIPT;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: OFFICEKMS;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: R25-DELTA2012;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: WINKMS;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: DIRECTACCESS;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: CBORD;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: NESSUS;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: PROCARE;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: TI22-PSYCLINIC;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: REFWIN;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: AD-RMAD;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: ADWSUSCLIENTS;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: NETENG-PA;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: OPIE;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: TI22;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: FTK;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: R25-WEBVIEWER;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: ADDC55;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: LANSWEEPER;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: ADAUTHDC1;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: IMAGES;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: ADDC2;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: SCCM-PRIM01;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: ADCA1;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: TI22-COE;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: EMUSQLEXPRESS;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: DPS-SECURE02;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: ADARC;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: PNCWEB-2012;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: CEDARON;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: ADDC5;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: COE-WEBDAV;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: ADWU1;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: MBAM;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: SANDBOX-DB;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: Echo-01;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: ADDC4;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: R25-INTERFACE;STOPPED; Host has exited from a period of scheduled downtime
[1463632200] HOST DOWNTIME ALERT: SP2010DB;STOPPED; Host has exited from a period of scheduled downtime
[1463632593] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463632713] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1463632833] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463632931] SERVICE ALERT: ADWU1;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463633031] SERVICE ALERT: ADWU1;Memory Usage;OK;SOFT;2;Memory usage: total:22471.74 MB - used: 2469.25 MB (11%) - free: 20002.50 MB (89%)
[1463633068] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;1;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1463633188] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;2;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1463633308] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;HARD;3;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1463633433] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463634033] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463634633] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463637033] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463638233] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463638702] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1463638822] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 73% (15 min average)
[1463638833] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463639433] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463640508] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;OK;HARD;3;OK: No entries found
[1463640633] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463641833] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463642423] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1463642542] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463642662] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 83% (15 min average)
[1463646033] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 75% (15 min average)
[1463646633] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1463646753] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 85% (15 min average)
[1463647953] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463648073] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463648193] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463648793] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 81% (15 min average)
[1463649393] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1463649513] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1463649633] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 76% (15 min average)
[1463650233] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463650353] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1463650473] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 90% (15 min average)
[1463651073] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1463651193] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 92% (15 min average)
[1463651448] SERVICE ALERT: ADWU1;Service: Group Policy;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1463651464] SERVICE ALERT: ADWU1;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1463651568] SERVICE ALERT: ADWU1;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1463651581] SERVICE ALERT: ADWU1;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1463651793] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1463651908] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;1;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1463651913] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 86% (15 min average)
[1463652028] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;2;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1463652148] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;HARD;3;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1463652269] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 73% - OK!
[1463652273] HOST ALERT: hallesecondsouth;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.44 ms
[1463652513] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1463652583] SERVICE ALERT: hallesecondsouth;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1463652633] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1463652633] SERVICE ALERT: hallesecondsouth;Device Status;OK;HARD;3;Ready
[1463652753] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463653353] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463653953] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 73% (15 min average)
[1463654553] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463654672] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1463655872] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463655992] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 90% (15 min average)
[1463657744] SERVICE ALERT: honorspk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 51% - OK!
[1463657792] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463657912] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1463658032] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463658632] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 90% (15 min average)
[1463658748] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;OK;HARD;3;OK: No entries found
[1463658861] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463658981] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 85% (15 min average)
[1463659203] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (40.5% change >= 20.0% threshold)
[1463659203] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1463659203] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1463659203] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1463659203] SERVICE DOWNTIME ALERT: NETIDADPROXY;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1463659203] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1463659232] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1463659352] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463659472] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1463660672] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 37% (15 min average)
[1463663672] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463663792] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 90% (15 min average)
[1463663828] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors is too high! 5237 found since last run. securityErrors OK. wholeSubtreeSearchOps is too high! 5242 found since last run. oneLevelSearchOps OK. searchOps is too high! 5246 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463663948] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463664392] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1463664512] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1463664548] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5240 found since last run. oneLevelSearchOps OK. searchOps is too high! 5241 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463664668] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463665112] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463665232] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463665268] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5145 found since last run. oneLevelSearchOps OK. searchOps is too high! 5165 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463665352] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463665388] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463665881] SERVICE ALERT: ADDC5;GPO Replication;UNKNOWN;SOFT;1;CHECK_NRPE: Received 0 bytes from daemon.  Check the remote server logs for error messages.
[1463665988] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10471 found since last run. oneLevelSearchOps OK. searchOps is too high! 10479 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463666013] SERVICE ALERT: ADDC5;GPO Replication;OK;SOFT;2;Export-CSV : The process cannot access the file 'C:\Program
[1463666108] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10469 found since last run. oneLevelSearchOps OK. searchOps is too high! 10470 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463666228] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463666781] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1463666828] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 15718 found since last run. oneLevelSearchOps OK. searchOps is too high! 15723 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463666901] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 86% (15 min average)
[1463666948] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;2;errors is too high! 5221 found since last run. securityErrors OK. wholeSubtreeSearchOps is too high! 5234 found since last run. oneLevelSearchOps OK. searchOps is too high! 5238 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463667068] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463667501] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1463667621] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1463667668] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10479 found since last run. oneLevelSearchOps OK. searchOps is too high! 10482 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463667741] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 94% (15 min average)
[1463667788] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463668341] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463668352] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 83% (15 min average)
[1463668388] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors is too high! 5223 found since last run. securityErrors OK. wholeSubtreeSearchOps is too high! 20941 found since last run. oneLevelSearchOps OK. searchOps is too high! 20946 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463668461] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1463668508] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10469 found since last run. oneLevelSearchOps OK. searchOps is too high! 10470 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463668580] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 93% (15 min average)
[1463668628] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463669229] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 60261 found since last run. oneLevelSearchOps OK. searchOps is too high! 60267 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463669348] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463669949] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 60260 found since last run. oneLevelSearchOps OK. searchOps is too high! 60261 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463670070] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463670752] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1463670872] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463670992] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463671268] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 13449 found since last run. oneLevelSearchOps OK. searchOps is too high! 13450 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463671388] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463671444] SERVICE ALERT: INTLAPP-ISSMAP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1463671566] SERVICE ALERT: INTLAPP-ISSMAP;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1463671592] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 67% (15 min average)
[1463671685] SERVICE ALERT: INTLAPP-ISSMAP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463671988] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 13911 found since last run. oneLevelSearchOps OK. searchOps is too high! 13912 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463672108] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463672245] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (26.2% change >= 20.0% threshold)
[1463672245] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1463672245] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1463672245] HOST DOWNTIME ALERT: VIEW-CS;STARTED; Host has entered a period of scheduled downtime
[1463672245] SERVICE DOWNTIME ALERT: NETIDADPROXY;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1463672245] HOST DOWNTIME ALERT: VCOMPOSER;STARTED; Host has entered a period of scheduled downtime
[1463672284] SERVICE ALERT: INTLAPP-ISSMAP;CPU Load;OK;HARD;3;CPU Load 78% (15 min average)
[1463672708] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 14382 found since last run. oneLevelSearchOps OK. searchOps is too high! 14383 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463672792] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463672829] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463672912] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463673032] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463673428] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 13918 found since last run. oneLevelSearchOps OK. searchOps is too high! 13919 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463673548] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 2593 found since last run. oneLevelSearchOps OK. searchOps is too high! 2594 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463673632] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463673668] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463674232] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1463674268] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 16503 found since last run. oneLevelSearchOps OK. searchOps is too high! 16504 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463674388] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463674988] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 36686 found since last run. oneLevelSearchOps OK. searchOps is too high! 36687 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463675109] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 24097 found since last run. oneLevelSearchOps OK. searchOps is too high! 24098 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463675228] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463675432] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 77% (15 min average)
[1463675828] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 46881 found since last run. oneLevelSearchOps OK. searchOps is too high! 46882 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463675948] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463676032] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1463676152] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463676272] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1463676498] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1463676549] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 2963 found since last run. oneLevelSearchOps OK. searchOps is too high! 2964 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463676558] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1463676599] SERVICE ALERT: hallefirst3;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1463676618] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1463676659] SERVICE ALERT: hallefirst3;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1463676659] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1463676668] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463676679] SERVICE ALERT: hallefirst3;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1463676689] HOST ALERT: hallefirst1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1463676714] SERVICE ALERT: hallefirst1;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! All Trays Still Full
[1463676718] HOST ALERT: hallefirst1;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.53 ms
[1463676719] SERVICE ALERT: hallefirst3;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1463676739] SERVICE ALERT: hallefirst3;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1463676769] HOST ALERT: hallefirst3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1463676799] SERVICE ALERT: hallefirst3;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1463676807] HOST ALERT: hallefirst3;UP;SOFT;2;PING OK - Packet loss = 44%, RTA = 0.49 ms
[1463677212] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 37% - OK!
[1463677268] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 3395 found since last run. oneLevelSearchOps OK. searchOps is too high! 3396 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463677313] SERVICE ALERT: hallefirst3;Device Status;OK;HARD;3;PERFORM PRINTER MAINTENANCE 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1463677388] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463677394] SERVICE ALERT: hallefirst3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1463677472] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 91% (15 min average)
[1463677988] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 55012 found since last run. oneLevelSearchOps OK. searchOps is too high! 55013 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463678109] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 66986 found since last run. oneLevelSearchOps OK. searchOps is too high! 66987 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463678229] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 66661 found since last run. oneLevelSearchOps OK. searchOps is too high! 66662 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463678672] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463678792] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1463678912] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463679378] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1463679428] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463679498] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1463679618] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 74% (15 min average)
[1463679744] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;1;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1463679862] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;2;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1463679982] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;HARD;3;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1463680052] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 60 seconds.
[1463680172] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 60 seconds.
[1463680263] SERVICE ALERT: idm1.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 1
[1463680280] SERVICE ALERT: idm2.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 5
[1463680292] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 60 seconds.
[1463680384] SERVICE ALERT: idm1.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 1
[1463680400] SERVICE ALERT: idm2.emich.edu;NSDS Errors;OK;SOFT;2;OK: Total errors: 0
[1463680503] SERVICE ALERT: idm1.emich.edu;NSDS Errors;OK;SOFT;3;OK: 0 Errors Reported
[1463680712] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463680834] SERVICE ALERT: idm1.emich.edu;AD Auth Driver;OK;HARD;3;OK: cn=ADAUTH Driver,cn=driverset1,o=system is running.
[1463681228] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 150678 found since last run. oneLevelSearchOps OK. searchOps is too high! 150679 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463681312] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463681349] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 43053 found since last run. oneLevelSearchOps OK. searchOps is too high! 43054 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463681470] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 39776 found since last run. oneLevelSearchOps OK. searchOps is too high! 39777 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463682668] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463682668] SERVICE FLAPPING ALERT: idm1.emich.edu;LDAP Stats;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1463682800] SERVICE ALERT: idm2.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 6
[1463682822] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 6
[1463682844] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 6
[1463682904] SERVICE ALERT: idm1.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 6
[1463682920] SERVICE ALERT: idm2.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 6
[1463682942] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 6
[1463682963] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 6
[1463683024] SERVICE ALERT: idm1.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 6
[1463683040] SERVICE ALERT: idm2.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 6
[1463683062] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 6
[1463683083] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 6
[1463683112] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1463683143] SERVICE ALERT: idm1.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 6
[1463683227] HOST ALERT: inqsit;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1463683291] HOST ALERT: inqsit;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.36 ms
[1463683712] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463683832] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463683868] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors is too high! 32475 found since last run. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463683952] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463683989] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;2;errors is too high! 31346 found since last run. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463684108] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;HARD;3;errors is too high! 28644 found since last run. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463684133] HOST ALERT: duoauth1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1463684221] HOST ALERT: duoauth1;UP;SOFT;2;PING WARNING - Packet loss = 82%, RTA = 0.33 ms
[1463684343] SERVICE ALERT: idm1.emich.edu;NSDS Errors;OK;HARD;3;OK: 0 Errors Reported
[1463684552] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1463684840] SERVICE ALERT: idm2.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1463684862] SERVICE ALERT: idm4.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1463684883] SERVICE ALERT: idm3.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1463685752] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463685908] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463686040] SERVICE ALERT: idm2.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 8
[1463686062] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 8
[1463686084] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 8
[1463686144] SERVICE ALERT: idm1.emich.edu;NSDS Errors;WARNING;SOFT;1;WARNING: Total errors: 16
[1463686160] SERVICE ALERT: idm2.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 8
[1463686182] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 8
[1463686203] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 8
[1463686263] SERVICE ALERT: idm1.emich.edu;NSDS Errors;WARNING;SOFT;2;WARNING: Total errors: 16
[1463686281] SERVICE ALERT: idm2.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 8
[1463686302] SERVICE ALERT: idm4.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 8
[1463686323] SERVICE ALERT: idm3.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 8
[1463686352] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463686383] SERVICE ALERT: idm1.emich.edu;NSDS Errors;WARNING;HARD;3;WARNING: Total errors: 12
[1463686952] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 84% (15 min average)
[1463687235] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;OK;HARD;3;OK: No entries found
[1463687552] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463687672] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1463687792] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1463688680] SERVICE ALERT: idm2.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1463688680] SERVICE FLAPPING ALERT: idm2.emich.edu;NSDS Errors;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1463688702] SERVICE ALERT: idm4.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1463688702] SERVICE FLAPPING ALERT: idm4.emich.edu;NSDS Errors;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1463688723] SERVICE ALERT: idm3.emich.edu;NSDS Errors;OK;HARD;3;OK: Total errors: 0
[1463688723] SERVICE FLAPPING ALERT: idm3.emich.edu;NSDS Errors;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1463688784] SERVICE ALERT: idm1.emich.edu;NSDS Errors;OK;HARD;3;OK: 0 Errors Reported
[1463688784] SERVICE FLAPPING ALERT: idm1.emich.edu;NSDS Errors;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1463688992] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 83% (15 min average)
[1463690192] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1463690312] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1463690432] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1463691032] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 77% (15 min average)
[1463691254] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 1 Trays Empty
[1463691310] SERVICE ALERT: REFWIN;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463691313] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 1 Trays Empty
[1463691340] HOST ALERT: REFWIN;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1463691342] SERVICE ALERT: REFWIN;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463691371] SERVICE ALERT: REFWIN;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463691372] HOST ALERT: REFWIN;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1463691373] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 1 Trays Empty
[1463691382] SERVICE ALERT: REFWIN;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463691406] SERVICE ALERT: REFWIN;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463691412] HOST ALERT: REFWIN;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1463691418] SERVICE ALERT: REFWIN;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463691419] SERVICE ALERT: REFWIN;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1463691427] SERVICE ALERT: REFWIN;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1463691430] SERVICE ALERT: REFWIN;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1463691442] SERVICE ALERT: REFWIN;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1463691447] HOST ALERT: REFWIN;UP;SOFT;4;PING WARNING - Packet loss = 83%, RTA = 0.34 ms
[1463691471] SERVICE ALERT: REFWIN;C:\ Drive Space;OK;SOFT;2;c:\ - total: 59.90 Gb - used: 41.53 Gb (69%) - free 18.37 Gb (31%)
[1463691482] SERVICE ALERT: REFWIN;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1463691506] SERVICE ALERT: REFWIN;VMWare Tools;OK;SOFT;2;VMTools: Started
[1463691518] SERVICE ALERT: REFWIN;Service: RPC;OK;SOFT;2;RpcSs: Started
[1463691519] SERVICE ALERT: REFWIN;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1463691530] SERVICE ALERT: REFWIN;Service: Group Policy;OK;SOFT;3;gpsvc: Started
[1463691547] SERVICE ALERT: REFWIN;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1463691632] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463691752] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463691872] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463692021] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;CRITICAL;HARD;3;Black Print Cartridge HP C8543X is at 1% - CRITICAL!
[1463693041] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1463693072] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1463693101] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1463693161] SERVICE ALERT: hallefirst1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1463693672] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1463693792] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463693912] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463694420] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1463696108] SERVICE FLAPPING ALERT: idm1.emich.edu;LDAP Stats;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1463696161] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;There is no help available for this item. 54 6F 20 65 78 69 74 20 70 72 65 73 73 20 1F 2E
[1463696312] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 91% (15 min average)
[1463696912] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1463697032] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1463697152] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 72% (15 min average)
[1463697752] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463697872] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463697992] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463698279] SERVICE FLAPPING ALERT: idm2.emich.edu;NSDS Errors;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1463698301] SERVICE FLAPPING ALERT: idm4.emich.edu;NSDS Errors;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1463698323] SERVICE FLAPPING ALERT: idm3.emich.edu;NSDS Errors;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1463698382] SERVICE FLAPPING ALERT: idm1.emich.edu;NSDS Errors;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1463700392] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 84% (15 min average)
[1463700992] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1463701112] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 75% (15 min average)
[1463701560] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1463701620] SERVICE ALERT: hallefirst1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1463701680] SERVICE ALERT: hallefirst1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1463701712] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1463701832] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1463701952] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1463703111] SERVICE ALERT: hallefirst3;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1463703171] SERVICE ALERT: hallefirst3;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1463703232] SERVICE ALERT: hallefirst3;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1463703752] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1463704352] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463704472] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463704592] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463705107] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 64546 found since last run. oneLevelSearchOps OK. searchOps is too high! 64547 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463705227] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 29578 found since last run. oneLevelSearchOps OK. searchOps is too high! 29579 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463705348] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 29824 found since last run. oneLevelSearchOps OK. searchOps is too high! 29825 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1463705631] SERVICE ALERT: hallefirst3;Device Status;OK;HARD;3;PERFORM PRINTER MAINTENANCE 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1463706392] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 73% (15 min average)
[1463706479] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1463706992] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463707112] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 70% (15 min average)
[1463707145] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1463707712] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463707832] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 83% (15 min average)
[1463708432] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463708552] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463708672] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463710472] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 73% (15 min average)
[1463711072] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1463711192] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1463711312] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463712512] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1463714312] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1463714912] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)