[1469764877] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 513 found since last run.
[1469764997] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469765597] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 586 found since last run.
[1469765717] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469766007] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 91% (15 min average)
[1469766317] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 534 found since last run.
[1469766437] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469767037] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 510 found since last run.
[1469767157] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469767207] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1469767327] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1469767447] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1469768357] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 513 found since last run.
[1469768477] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469768647] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1469769077] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 567 found since last run.
[1469769197] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469769246] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1469769366] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1469769486] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 94% (15 min average)
[1469769786] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 522 found since last run.
[1469769797] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 528 found since last run.
[1469769906] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469769917] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469770086] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1469770206] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 92% (15 min average)
[1469770506] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 506 found since last run.
[1469770626] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469771117] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 504 found since last run.
[1469771226] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 506 found since last run.
[1469771237] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469771346] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469771578] SERVICE ALERT: EMUFILES;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1469771697] SERVICE ALERT: EMUFILES;CPU Load;OK;SOFT;2;CPU Load 5% (15 min average)
[1469771837] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 182 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 630 found since last run.
[1469771883] SERVICE ALERT: INTLAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1469771946] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 197 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 656 found since last run.
[1469771957] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469772003] SERVICE ALERT: INTLAPP;CPU Load;OK;SOFT;2;CPU Load 4% (15 min average)
[1469772066] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469772606] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1469772726] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1469772846] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1469772866] SERVICE ALERT: FITNESSGRAM;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1469772986] SERVICE ALERT: FITNESSGRAM;CPU Load;OK;SOFT;2;CPU Load 6% (15 min average)
[1469773266] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 501 found since last run.
[1469773386] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469774357] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 517 found since last run.
[1469774477] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469774646] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 73% (15 min average)
[1469775246] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1469775366] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1469775486] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 87% (15 min average)
[1469776086] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1469776206] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1469776326] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1469776386] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 528 found since last run.
[1469776506] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469777106] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 522 found since last run.
[1469777226] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469777800] SERVICE ALERT: hallegroundcolor;Toner Collection Unit HP CE265A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1469777855] SERVICE ALERT: hallegroundcolor;Toner Collection Unit HP CE265A;OK;SOFT;2;
[1469778077] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 514 found since last run.
[1469778197] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469778726] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 91% (15 min average)
[1469779326] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1469779446] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1469779566] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 92% (15 min average)
[1469779998] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 501 found since last run.
[1469780117] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469780826] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 583 found since last run.
[1469780946] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469781316] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 535 found since last run.
[1469781436] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469781546] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 563 found since last run.
[1469781666] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469781966] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1469782035] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 566 found since last run.
[1469782086] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 84% (15 min average)
[1469782155] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469782266] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 536 found since last run.
[1469782386] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469782755] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 530 found since last run.
[1469782875] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469782986] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 603 found since last run.
[1469783106] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469783475] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 560 found since last run.
[1469783595] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469783706] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 572 found since last run.
[1469783827] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469784195] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 538 found since last run.
[1469784315] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469784426] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 616 found since last run.
[1469784546] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469784915] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 104 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 638 found since last run.
[1469785035] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469785146] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 126 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 641 found since last run.
[1469785266] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469785635] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 129 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 634 found since last run.
[1469785755] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469785866] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 585 found since last run.
[1469785918] SERVICE ALERT: ADWU1;Service: User Profile;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1469785986] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469786038] SERVICE ALERT: ADWU1;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1469786355] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 559 found since last run.
[1469786421] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;1;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1469786475] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469786541] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;2;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1469786586] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 598 found since last run.
[1469786661] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;HARD;3;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1469786706] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469786886] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1469787006] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1469787075] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 584 found since last run.
[1469787126] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1469787195] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469787306] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 537 found since last run.
[1469787426] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469787726] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 85% (15 min average)
[1469787795] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 547 found since last run.
[1469787915] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469788026] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 547 found since last run.
[1469788146] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469788515] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 534 found since last run.
[1469788635] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469788746] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 566 found since last run.
[1469788866] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469788926] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1469789046] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1469789166] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 94% (15 min average)
[1469789235] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 532 found since last run.
[1469789355] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469789466] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 538 found since last run.
[1469789586] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469789955] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 561 found since last run.
[1469790075] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469790186] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 122 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 601 found since last run.
[1469790306] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469790366] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1469790486] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1469790606] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1469790675] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 551 found since last run.
[1469790795] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469790906] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 106 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 575 found since last run.
[1469791026] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469791206] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 68% (15 min average)
[1469791395] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 111 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 665 found since last run.
[1469791515] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469791626] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 553 found since last run.
[1469791746] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469792115] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 103 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 596 found since last run.
[1469792235] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469792346] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 123 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 665 found since last run.
[1469792466] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469792835] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 123 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 705 found since last run.
[1469792955] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469793066] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 101 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 628 found since last run.
[1469793186] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469793261] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;OK;HARD;3;OK: No entries found
[1469793555] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 623 found since last run.
[1469793675] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469793786] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1178 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 601 found since last run.
[1469793906] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469794275] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 523 found since last run.
[1469794395] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469794407] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1469794419] SERVICE ALERT: owen101lab3;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1469794437] HOST ALERT: owen101lab3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1469794462] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 43% - OK!
[1469794466] HOST ALERT: owen101lab3;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.47 ms
[1469794473] SERVICE ALERT: owen101lab3;Device Status;OK;SOFT;2;Warming up
[1469794506] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1389 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 680 found since last run.
[1469794626] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469794996] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1470 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 1327 found since last run.
[1469795115] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469795226] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 109 found since last run. wholeSubtreeSearchOps is too high! 1831 found since last run. oneLevelSearchOps OK. searchOps is too high! 2029 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 1612 found since last run.
[1469795346] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469795373] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 1264 found since last run. oneLevelSearchOps OK. searchOps is too high! 1265 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1469795493] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469795714] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 100 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1095 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 651 found since last run.
[1469795834] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469795947] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 108 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 722 found since last run.
[1469796006] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1469796066] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469796126] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 41% (15 min average)
[1469796434] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1006 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 588 found since last run.
[1469796554] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469796665] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 673 found since last run.
[1469796785] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469797154] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1445 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 563 found since last run.
[1469797274] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469797385] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 670 found since last run.
[1469797505] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469797874] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 609 found since last run.
[1469797994] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469798105] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 631 found since last run.
[1469798225] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469798443] HOST ALERT: owen101lab3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1469798533] HOST ALERT: owen101lab3;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1469798594] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 586 found since last run.
[1469798623] HOST ALERT: owen101lab3;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1469798666] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1469798679] SERVICE ALERT: owen101lab3;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1469798696] HOST ALERT: owen101lab3;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1469798714] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469798726] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1469798739] SERVICE ALERT: owen101lab3;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1469798756] HOST ALERT: owen101lab3;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1469798786] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1469798799] SERVICE ALERT: owen101lab3;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1469798816] HOST ALERT: owen101lab3;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1469798825] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 657 found since last run.
[1469798863] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1469798893] HOST ALERT: owen101lab3;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1469798923] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1469798945] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469798953] HOST ALERT: owen101lab3;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1469798983] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1469799013] HOST ALERT: owen101lab3;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1469799103] HOST ALERT: owen101lab3;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1469799314] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 574 found since last run.
[1469799434] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469799545] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 1502 found since last run. oneLevelSearchOps OK. searchOps is too high! 1541 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 708 found since last run.
[1469799665] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469800034] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1059 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 653 found since last run.
[1469800154] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469800265] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1540 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 704 found since last run.
[1469800385] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469800754] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 580 found since last run.
[1469800874] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469800985] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 1755 found since last run. oneLevelSearchOps OK. searchOps is too high! 1895 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 736 found since last run.
[1469801105] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469801474] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1254 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 649 found since last run.
[1469801594] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469801705] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 103 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 741 found since last run.
[1469801825] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469802190] SERVICE ALERT: IMAGES;Service: MDT Monitor;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469802194] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 640 found since last run.
[1469802314] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469802425] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 122 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1281 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 796 found since last run.
[1469802545] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469802725] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1469802770] SERVICE ALERT: IMAGES;Service: MDT Monitor;UNKNOWN;HARD;3;Failed to open service MDT_Monitor: 1060: The specified service does not exist as an installed service.
[1469802914] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 648 found since last run.
[1469803034] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469803145] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 2551 found since last run. oneLevelSearchOps OK. searchOps is too high! 2634 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 697 found since last run.
[1469803265] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469803634] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1429 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 653 found since last run.
[1469803754] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469803865] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1031 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 660 found since last run.
[1469803985] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469804354] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 1558 found since last run. oneLevelSearchOps OK. searchOps is too high! 1566 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 630 found since last run.
[1469804474] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469804585] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1008 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 634 found since last run.
[1469804705] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469805074] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 110 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 733 found since last run.
[1469805194] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469805305] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1261 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 676 found since last run.
[1469805425] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469805794] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 1568 found since last run. oneLevelSearchOps OK. searchOps is too high! 1583 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 704 found since last run.
[1469805827] SERVICE ALERT: epprogsql2;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.4.204 and port 12489: Connection refused
[1469805836] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469805914] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469805936] SERVICE ALERT: epprogsql2;CPU Load;OK;SOFT;2;CPU Load 0% (15 min average)
[1469805947] SERVICE ALERT: epprogsql2;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1469806025] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 868 found since last run.
[1469806145] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469806514] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1009 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 759 found since last run.
[1469806634] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469806745] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1093 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 605 found since last run.
[1469806865] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469807101] HOST DOWNTIME ALERT: TARGETXTEST2012;STARTED; Host has entered a period of scheduled downtime
[1469807101] HOST DOWNTIME ALERT: BOEAPPTEST;STARTED; Host has entered a period of scheduled downtime
[1469807101] HOST DOWNTIME ALERT: ITRFTEST;STARTED; Host has entered a period of scheduled downtime
[1469807101] HOST DOWNTIME ALERT: EVENTSTEST-2012;STARTED; Host has entered a period of scheduled downtime
[1469807101] HOST DOWNTIME ALERT: AXIOM-TEST-2012;STARTED; Host has entered a period of scheduled downtime
[1469807101] HOST DOWNTIME ALERT: AX-QUAL;STARTED; Host has entered a period of scheduled downtime
[1469807101] HOST DOWNTIME ALERT: AX-DEVL;STARTED; Host has entered a period of scheduled downtime
[1469807101] HOST DOWNTIME ALERT: INTLTESTAPP;STARTED; Host has entered a period of scheduled downtime
[1469807101] HOST DOWNTIME ALERT: SANDBOX-APP;STARTED; Host has entered a period of scheduled downtime
[1469807101] HOST DOWNTIME ALERT: AXLIC-DEVL;STARTED; Host has entered a period of scheduled downtime
[1469807101] HOST DOWNTIME ALERT: INTLTESTDB;STARTED; Host has entered a period of scheduled downtime
[1469807101] HOST DOWNTIME ALERT: MAPS01TEST;STARTED; Host has entered a period of scheduled downtime
[1469807234] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 109 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1309 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 693 found since last run.
[1469807354] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469807465] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 111 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1189 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 646 found since last run.
[1469807585] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469807954] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 104 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 775 found since last run.
[1469808074] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469808125] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1469808185] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 104 found since last run. wholeSubtreeSearchOps is too high! 1861 found since last run. oneLevelSearchOps OK. searchOps is too high! 2040 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 759 found since last run.
[1469808245] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 44% (15 min average)
[1469808305] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469808674] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1423 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 583 found since last run.
[1469808795] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469808845] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1469808905] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1464 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 641 found since last run.
[1469808965] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 71% (15 min average)
[1469809025] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469809394] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 631 found since last run.
[1469809514] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469809625] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 1619 found since last run. oneLevelSearchOps OK. searchOps is too high! 1655 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 660 found since last run.
[1469809745] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469810114] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 685 found since last run.
[1469810165] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1469810234] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469810285] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 71% (15 min average)
[1469810345] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1535 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 671 found since last run.
[1469810465] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469810834] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 1765 found since last run. oneLevelSearchOps OK. searchOps is too high! 1872 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 632 found since last run.
[1469810954] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469811064] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1140 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 640 found since last run.
[1469811184] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469811554] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1393 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 649 found since last run.
[1469811674] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469811783] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1234 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 641 found since last run.
[1469811903] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469812163] SERVICE ALERT: TARGETXTEST2012;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1469812265] SERVICE ALERT: ITRFTEST;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469812268] SERVICE ALERT: ITRFTEST;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469812268] SERVICE ALERT: ITRFTEST;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469812269] SERVICE ALERT: ITRFTEST;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469812269] SERVICE ALERT: ITRFTEST;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469812271] SERVICE ALERT: ITRFTEST;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469812271] SERVICE ALERT: ITRFTEST;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469812271] SERVICE ALERT: ITRFTEST;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469812271] SERVICE ALERT: ITRFTEST;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469812272] SERVICE ALERT: ITRFTEST;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469812272] SERVICE ALERT: ITRFTEST;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469812273] SERVICE ALERT: ITRFTEST;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469812274] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 2132 found since last run. oneLevelSearchOps OK. searchOps is too high! 2235 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 623 found since last run.
[1469812281] SERVICE ALERT: ITRFTEST;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469812283] SERVICE ALERT: TARGETXTEST2012;CPU Load;OK;SOFT;2;CPU Load 65% (15 min average)
[1469812295] HOST ALERT: ITRFTEST;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1469812384] SERVICE ALERT: ITRFTEST;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469812385] HOST ALERT: ITRFTEST;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1469812388] SERVICE ALERT: ITRFTEST;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469812388] SERVICE ALERT: ITRFTEST;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469812389] SERVICE ALERT: ITRFTEST;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469812389] SERVICE ALERT: ITRFTEST;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469812391] SERVICE ALERT: ITRFTEST;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469812391] SERVICE ALERT: ITRFTEST;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469812391] SERVICE ALERT: ITRFTEST;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469812391] SERVICE ALERT: ITRFTEST;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469812392] SERVICE ALERT: ITRFTEST;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469812392] SERVICE ALERT: ITRFTEST;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469812393] SERVICE ALERT: ITRFTEST;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469812394] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469812401] SERVICE ALERT: ITRFTEST;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469812418] HOST ALERT: ITRFTEST;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1469812491] SERVICE ALERT: INTLTESTAPP;Service: Group Policy;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1469812493] SERVICE ALERT: INTLTESTAPP;Service: NetLogin;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1469812493] SERVICE ALERT: INTLTESTAPP;Service: WMI;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1469812493] SERVICE ALERT: INTLTESTAPP;Service: Firewall;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1469812493] SERVICE ALERT: INTLTESTDB;Service: WMI;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1469812493] SERVICE ALERT: INTLTESTAPP;Service: RPC;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1469812493] SERVICE ALERT: INTLTESTDB;Service: NetLogin;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1469812493] SERVICE ALERT: INTLTESTAPP;VMWare Tools;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1469812503] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 677 found since last run.
[1469812504] SERVICE ALERT: ITRFTEST;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469812508] SERVICE ALERT: ITRFTEST;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469812508] HOST ALERT: ITRFTEST;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1469812508] SERVICE ALERT: ITRFTEST;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469812509] SERVICE ALERT: ITRFTEST;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469812509] SERVICE ALERT: ITRFTEST;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469812511] SERVICE ALERT: ITRFTEST;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469812511] SERVICE ALERT: ITRFTEST;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469812511] SERVICE ALERT: ITRFTEST;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469812511] SERVICE ALERT: ITRFTEST;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469812512] SERVICE ALERT: ITRFTEST;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469812512] SERVICE ALERT: ITRFTEST;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469812513] SERVICE ALERT: ITRFTEST;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469812521] SERVICE ALERT: ITRFTEST;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469812538] HOST ALERT: ITRFTEST;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1469812602] HOST ALERT: ITRFTEST;UP;SOFT;6;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1469812611] SERVICE ALERT: INTLTESTAPP;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1469812613] SERVICE ALERT: INTLTESTDB;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1469812613] SERVICE ALERT: INTLTESTAPP;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1469812613] SERVICE ALERT: INTLTESTAPP;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1469812613] SERVICE ALERT: INTLTESTAPP;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1469812613] SERVICE ALERT: INTLTESTDB;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1469812613] SERVICE ALERT: INTLTESTAPP;Service: RPC;OK;SOFT;2;RpcSs: Started
[1469812613] SERVICE ALERT: INTLTESTAPP;VMWare Tools;OK;SOFT;2;VMTools: Started
[1469812623] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469812994] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1029 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 704 found since last run.
[1469813066] SERVICE ALERT: INTLAPP-ISSMAP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1469813085] SERVICE ALERT: ITRFTEST;Memory Usage;OK;HARD;3;Memory usage: total:8188.80 MB - used: 1923.93 MB (23%) - free: 6264.87 MB (77%)
[1469813088] SERVICE ALERT: ITRFTEST;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 9 minute(s)
[1469813088] SERVICE ALERT: ITRFTEST;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1469813089] SERVICE ALERT: ITRFTEST;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1469813089] SERVICE ALERT: ITRFTEST;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1469813091] SERVICE ALERT: ITRFTEST;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1469813091] SERVICE ALERT: ITRFTEST;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1469813091] SERVICE ALERT: ITRFTEST;Service: Windows Time;OK;HARD;3;W32Time: Started
[1469813091] SERVICE ALERT: ITRFTEST;Service: WMI;OK;HARD;3;Winmgmt: Started
[1469813092] SERVICE ALERT: ITRFTEST;Service: RPC;OK;HARD;3;RpcSs: Started
[1469813092] SERVICE ALERT: ITRFTEST;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1469813093] SERVICE ALERT: ITRFTEST;CPU Load;OK;HARD;3;CPU Load 1% (15 min average)
[1469813101] SERVICE ALERT: ITRFTEST;C:\ Drive Space;OK;HARD;3;c:\ - total: 408.28 Gb - used: 40.40 Gb (10%) - free 367.87 Gb (90%)
[1469813114] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469813186] SERVICE ALERT: INTLAPP-ISSMAP;CPU Load;OK;SOFT;2;CPU Load 89% (15 min average)
[1469813223] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 116 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1095 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 711 found since last run.
[1469813285] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1469813343] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469813405] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 87% (15 min average)
[1469813412] SERVICE ALERT: AX-QUAL;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.7.143 and port 12489: Connection refused
[1469813529] SERVICE ALERT: AX-QUAL;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1469813714] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 125 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1448 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 722 found since last run.
[1469813834] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469813943] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 1738 found since last run. oneLevelSearchOps OK. searchOps is too high! 1997 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 780 found since last run.
[1469814063] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469814080] SERVICE ALERT: AXIOM-TEST-2012;Service: Windows Time;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1469814101] SERVICE ALERT: AXIOM-TEST-2012;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.4.172 and port 12489: Connection refused
[1469814200] SERVICE ALERT: AXIOM-TEST-2012;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1469814221] SERVICE ALERT: AXIOM-TEST-2012;Service: RPC;OK;SOFT;2;RpcSs: Started
[1469814338] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1469814343] SERVICE ALERT: SANDBOX-APP;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.4.163 and port 12489: No route to host
[1469814343] SERVICE ALERT: SANDBOX-APP;Service: User Profile;CRITICAL;SOFT;1;connect to address 164.76.4.163 and port 12489: No route to host
[1469814343] SERVICE ALERT: SANDBOX-APP;Service: Windows Time;CRITICAL;SOFT;1;connect to address 164.76.4.163 and port 12489: No route to host
[1469814346] SERVICE ALERT: SANDBOX-APP;VMWare Tools;CRITICAL;SOFT;1;connect to address 164.76.4.163 and port 12489: No route to host
[1469814346] HOST ALERT: SANDBOX-APP;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.4.163)
[1469814351] SERVICE ALERT: SANDBOX-APP;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469814398] HOST ALERT: SANDBOX-APP;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1469814419] HOST ALERT: SANDBOX-APP;UP;SOFT;3;PING OK - Packet loss = 37%, RTA = 0.33 ms
[1469814434] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1282 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 712 found since last run.
[1469814451] SERVICE ALERT: SANDBOX-APP;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1469814453] SERVICE ALERT: SANDBOX-APP;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1469814458] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1469814459] SERVICE ALERT: SANDBOX-APP;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1469814459] SERVICE ALERT: SANDBOX-APP;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1469814461] SERVICE ALERT: SANDBOX-APP;VMWare Tools;OK;SOFT;2;VMTools: Started
[1469814474] SERVICE ALERT: BOEAPPTEST;Service: NetLogin;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1469814554] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469814571] SERVICE ALERT: SANDBOX-APP;CPU Load;OK;SOFT;3;CPU Load 5% (15 min average)
[1469814578] SERVICE ALERT: BOEAPPTEST;CPU Load;OK;SOFT;3;CPU Load 2% (15 min average)
[1469814593] SERVICE ALERT: BOEAPPTEST;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1469814605] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1469814663] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 1583 found since last run. oneLevelSearchOps OK. searchOps is too high! 1769 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 663 found since last run.
[1469814725] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1469814783] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469814845] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 93% (15 min average)
[1469815154] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 686 found since last run.
[1469815252] SERVICE ALERT: AX-QUAL;Service: NetLogin;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1469815274] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469815338] SERVICE ALERT: AX-DEVL;Service: WMI;CRITICAL;SOFT;1;connect to address 164.76.7.141 and port 12489: Connection refused
[1469815372] SERVICE ALERT: AX-QUAL;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1469815383] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 107 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1700 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 702 found since last run.
[1469815458] SERVICE ALERT: AX-DEVL;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1469815503] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469815874] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 101 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1379 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 657 found since last run.
[1469815994] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469816103] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 606 found since last run.
[1469816223] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469816594] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1492 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 677 found since last run.
[1469816714] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469816823] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 621 found since last run.
[1469816943] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469817314] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 2150 found since last run. oneLevelSearchOps OK. searchOps is too high! 2359 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 748 found since last run.
[1469817434] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469817543] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 669 found since last run.
[1469817663] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469817773] SERVICE ALERT: AXLIC-DEVL;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1469817796] HOST ALERT: ITRFTEST;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1469817886] HOST ALERT: ITRFTEST;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1469817893] SERVICE ALERT: AXLIC-DEVL;CPU Load;OK;SOFT;2;CPU Load 49% (15 min average)
[1469817903] SERVICE ALERT: ITRFTEST;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469817908] SERVICE ALERT: ITRFTEST;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469817908] SERVICE ALERT: ITRFTEST;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469817909] SERVICE ALERT: ITRFTEST;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469817909] SERVICE ALERT: ITRFTEST;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469817910] SERVICE ALERT: ITRFTEST;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469817910] SERVICE ALERT: ITRFTEST;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469817910] SERVICE ALERT: ITRFTEST;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469817910] SERVICE ALERT: ITRFTEST;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469817910] SERVICE ALERT: ITRFTEST;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469817910] SERVICE ALERT: ITRFTEST;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469817913] SERVICE ALERT: ITRFTEST;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469817921] SERVICE ALERT: ITRFTEST;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469817933] HOST ALERT: ITRFTEST;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1469817967] SERVICE ALERT: AXLIC-DEVL;Service: Group Policy;CRITICAL;SOFT;1;gpsvc: Stopped
[1469818023] SERVICE ALERT: ITRFTEST;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469818023] HOST ALERT: ITRFTEST;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1469818028] SERVICE ALERT: ITRFTEST;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469818028] SERVICE ALERT: ITRFTEST;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469818029] SERVICE ALERT: ITRFTEST;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469818029] SERVICE ALERT: ITRFTEST;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469818030] SERVICE ALERT: ITRFTEST;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469818030] SERVICE ALERT: ITRFTEST;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469818030] SERVICE ALERT: ITRFTEST;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469818030] SERVICE ALERT: ITRFTEST;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469818030] SERVICE ALERT: ITRFTEST;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469818030] SERVICE ALERT: ITRFTEST;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469818033] SERVICE ALERT: ITRFTEST;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469818034] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1490 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 675 found since last run.
[1469818041] SERVICE ALERT: ITRFTEST;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469818058] HOST ALERT: ITRFTEST;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1469818087] SERVICE ALERT: AXLIC-DEVL;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1469818140] SERVICE ALERT: ITRFTEST;Service: WMI;OK;SOFT;3;Winmgmt: Started
[1469818140] SERVICE ALERT: ITRFTEST;Service: User Profile;OK;SOFT;3;ProfSvc: Started
[1469818140] SERVICE ALERT: ITRFTEST;Service: RPC;OK;SOFT;3;RpcSs: Started
[1469818140] SERVICE ALERT: ITRFTEST;Service: Windows Time;CRITICAL;HARD;3;W32Time: Stopped
[1469818140] SERVICE ALERT: ITRFTEST;Service: DNS Client;OK;SOFT;3;Dnscache: Started
[1469818140] SERVICE ALERT: ITRFTEST;Service: Firewall;OK;SOFT;3;MpsSvc: Started
[1469818141] SERVICE ALERT: ITRFTEST;C:\ Drive Space;OK;SOFT;3;c:\ - total: 408.28 Gb - used: 39.96 Gb (10%) - free 368.31 Gb (90%)
[1469818142] SERVICE ALERT: ITRFTEST;CPU Load;OK;SOFT;3;CPU Load 20% (15 min average)
[1469818143] SERVICE ALERT: ITRFTEST;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469818143] HOST ALERT: ITRFTEST;UP;SOFT;6;PING WARNING - Packet loss = 80%, RTA = 0.34 ms
[1469818148] SERVICE ALERT: ITRFTEST;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469818148] SERVICE ALERT: ITRFTEST;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469818149] SERVICE ALERT: ITRFTEST;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469818149] SERVICE ALERT: ITRFTEST;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469818154] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469818263] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 625 found since last run.
[1469818383] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469818723] SERVICE ALERT: ITRFTEST;Memory Usage;OK;HARD;3;Memory usage: total:8188.80 MB - used: 1653.58 MB (20%) - free: 6535.22 MB (80%)
[1469818723] SERVICE FLAPPING ALERT: ITRFTEST;Memory Usage;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1469818727] SERVICE ALERT: ITRFTEST;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 10 minute(s)
[1469818727] SERVICE FLAPPING ALERT: ITRFTEST;Uptime;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1469818728] SERVICE ALERT: ITRFTEST;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1469818728] SERVICE FLAPPING ALERT: ITRFTEST;Service: Group Policy;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1469818729] SERVICE ALERT: ITRFTEST;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1469818729] SERVICE FLAPPING ALERT: ITRFTEST;NSClient++ Version;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1469818729] SERVICE ALERT: ITRFTEST;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1469818729] SERVICE FLAPPING ALERT: ITRFTEST;Service: NetLogin;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1469818730] SERVICE ALERT: ITRFTEST;Service: Windows Time;OK;HARD;3;W32Time: Started
[1469818730] SERVICE FLAPPING ALERT: ITRFTEST;Service: Windows Time;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1469818754] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1176 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 660 found since last run.
[1469818874] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469818983] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 1614 found since last run. oneLevelSearchOps OK. searchOps is too high! 1952 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 693 found since last run.
[1469819103] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469819474] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 1604 found since last run. oneLevelSearchOps OK. searchOps is too high! 1745 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 683 found since last run.
[1469819594] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469819703] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1196 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 601 found since last run.
[1469819823] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469820194] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 1753 found since last run. oneLevelSearchOps OK. searchOps is too high! 2008 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 593 found since last run.
[1469820314] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469820423] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 621 found since last run.
[1469820543] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469820914] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 1697 found since last run. oneLevelSearchOps OK. searchOps is too high! 1857 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 685 found since last run.
[1469821034] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469821143] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1760 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 682 found since last run.
[1469821263] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469821634] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1623 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 683 found since last run.
[1469821754] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469821863] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1015 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 528 found since last run.
[1469821983] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469822354] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 647 found since last run.
[1469822474] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469822583] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1121 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 558 found since last run.
[1469822703] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469823074] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 1560 found since last run. oneLevelSearchOps OK. searchOps is too high! 1748 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 662 found since last run.
[1469823194] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469823245] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1469823303] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 569 found since last run.
[1469823365] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1469823423] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469823485] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1469823794] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1400 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 661 found since last run.
[1469823914] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469824023] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1150 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 620 found since last run.
[1469824085] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 76% (15 min average)
[1469824143] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469824514] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1306 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 725 found since last run.
[1469824634] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469825101] HOST DOWNTIME ALERT: MAPS01TEST;STOPPED; Host has exited from a period of scheduled downtime
[1469825101] HOST DOWNTIME ALERT: TARGETXTEST2012;STOPPED; Host has exited from a period of scheduled downtime
[1469825101] HOST DOWNTIME ALERT: AXIOM-TEST-2012;STOPPED; Host has exited from a period of scheduled downtime
[1469825101] HOST DOWNTIME ALERT: EVENTSTEST-2012;STOPPED; Host has exited from a period of scheduled downtime
[1469825101] HOST DOWNTIME ALERT: BOEAPPTEST;STOPPED; Host has exited from a period of scheduled downtime
[1469825101] HOST DOWNTIME ALERT: ITRFTEST;STOPPED; Host has exited from a period of scheduled downtime
[1469825101] HOST DOWNTIME ALERT: INTLTESTDB;STOPPED; Host has exited from a period of scheduled downtime
[1469825101] HOST DOWNTIME ALERT: AX-QUAL;STOPPED; Host has exited from a period of scheduled downtime
[1469825101] HOST DOWNTIME ALERT: AXLIC-DEVL;STOPPED; Host has exited from a period of scheduled downtime
[1469825101] HOST DOWNTIME ALERT: INTLTESTAPP;STOPPED; Host has exited from a period of scheduled downtime
[1469825101] HOST DOWNTIME ALERT: AX-DEVL;STOPPED; Host has exited from a period of scheduled downtime
[1469825101] HOST DOWNTIME ALERT: SANDBOX-APP;STOPPED; Host has exited from a period of scheduled downtime
[1469825234] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 1507 found since last run. oneLevelSearchOps OK. searchOps is too high! 1676 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 625 found since last run.
[1469825343] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 523 found since last run.
[1469825355] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469825463] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469825885] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1469825954] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 642 found since last run.
[1469826005] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 89% (15 min average)
[1469826073] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469826408] HOST ALERT: jira01;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1469826418] HOST ALERT: jira01test;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1469826498] HOST ALERT: jira01;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1469826508] HOST ALERT: jira01test;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1469826588] HOST ALERT: jira01;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1469826598] HOST ALERT: jira01test;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1469826605] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1469826661] SERVICE ALERT: ldap.emich.edu;PING;CRITICAL;SOFT;1;PING CRITICAL - Packet loss = 60%, RTA = 0.29 ms
[1469826673] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 567 found since last run.
[1469826677] HOST ALERT: jira01;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1469826688] HOST ALERT: jira01test;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1469826725] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1469826767] HOST ALERT: jira01;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1469826775] SERVICE ALERT: ldap.emich.edu;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1469826778] HOST ALERT: jira01test;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1469826793] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469826845] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 90% (15 min average)
[1469826857] HOST ALERT: jira01;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1469826868] HOST ALERT: jira01test;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1469826947] HOST ALERT: jira01;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1469826958] HOST ALERT: jira01test;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1469827037] HOST ALERT: jira01;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1469827048] HOST ALERT: jira01test;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1469827127] HOST ALERT: jira01;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1469827138] HOST ALERT: jira01test;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1469827217] HOST ALERT: jira01;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1469827228] HOST ALERT: jira01test;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1469827263] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1122 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 669 found since last run.
[1469827383] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469827393] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 688 found since last run.
[1469827513] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469827802] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1469827802] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1469827802] HOST DOWNTIME ALERT: VUM;STARTED; Host has entered a period of scheduled downtime
[1469827802] SERVICE DOWNTIME ALERT: DIRECTACCESS;NSClient++ Version;STARTED; Service has entered a period of scheduled downtime
[1469827802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Windows Time;STARTED; Service has entered a period of scheduled downtime
[1469827802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Uptime;STARTED; Service has entered a period of scheduled downtime
[1469827802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: DNS Client;STARTED; Service has entered a period of scheduled downtime
[1469827802] SERVICE DOWNTIME ALERT: DIRECTACCESS;VMWare Tools;STARTED; Service has entered a period of scheduled downtime
[1469827802] SERVICE DOWNTIME ALERT: BOEAPP;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1469827802] SERVICE DOWNTIME ALERT: DIRECTACCESS;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1469827802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Firewall;STARTED; Service has entered a period of scheduled downtime
[1469827802] HOST DOWNTIME ALERT: INTLAPP-ISSMAP;STARTED; Host has entered a period of scheduled downtime
[1469827802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: Group Policy;STARTED; Service has entered a period of scheduled downtime
[1469827802] SERVICE DOWNTIME ALERT: NETIDADPROXY;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1469827802] SERVICE DOWNTIME ALERT: DIRECTACCESS;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1469827802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: NetLogin;STARTED; Service has entered a period of scheduled downtime
[1469827802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: RPC;STARTED; Service has entered a period of scheduled downtime
[1469827802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Memory Usage;STARTED; Service has entered a period of scheduled downtime
[1469827802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: User Profile;STARTED; Service has entered a period of scheduled downtime
[1469827802] SERVICE DOWNTIME ALERT: DIRECTACCESS;Service: WMI;STARTED; Service has entered a period of scheduled downtime
[1469827983] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 550 found since last run.
[1469828103] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469828113] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 504 found since last run.
[1469828233] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 580 found since last run.
[1469828353] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469828703] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 1905 found since last run. oneLevelSearchOps OK. searchOps is too high! 1936 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 1143 found since last run.
[1469828823] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469828953] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 520 found since last run.
[1469829073] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469829245] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1469829365] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1469829423] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 527 found since last run.
[1469829543] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469829673] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1059 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 564 found since last run.
[1469829793] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469829869] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1469829884] SERVICE ALERT: ADRODC-LIV;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469829886] SERVICE ALERT: ADRODC-LIV;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469829887] SERVICE ALERT: BBTS;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469829892] HOST ALERT: Google;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1469829892] HOST ALERT: porterpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1469829893] SERVICE ALERT: ADRODC-DET;Service: PXE Boot;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469829898] SERVICE ALERT: rackhampk;Black Cartridge HP CC364X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.141.212. Is the printer off?  Is the network down?
[1469829899] SERVICE ALERT: owen101labcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1469829899] SERVICE ALERT: sillpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1469829899] SERVICE ALERT: eateriespk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.59.141. Is the printer off?  Is the network down?
[1469829899] HOST ALERT: hallefirstcolor;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1469829900] SERVICE ALERT: hallefirstcolor;Toner Collection Unit HP CE265A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1469829900] SERVICE ALERT: ADRODC-LIV;Service: MDT Monitor;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469829901] HOST ALERT: studentcenterpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1469829908] HOST ALERT: marshallpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1469829909] SERVICE ALERT: owen101labcolor;Black Cartridge HP CE260X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1469829914] HOST ALERT: ADRODC-LIV;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1469829917] HOST ALERT: BBTS;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1469829921] SERVICE ALERT: BBTS_DB;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469829923] SERVICE ALERT: hallesouthcommons;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.111.220. Is the printer off?  Is the network down?
[1469829923] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1469829923] SERVICE ALERT: ADRODC-LIV;Service: Key Distribution Center;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469829923] HOST ALERT: ADRODC-DET;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1469829928] HOST ALERT: rackhampk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1469829928] SERVICE ALERT: hallegroundcolor;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1469829929] SERVICE ALERT: marshallpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.137.49. Is the printer off?  Is the network down?
[1469829929] SERVICE ALERT: alexanderpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.142.94. Is the printer off?  Is the network down?
[1469829929] HOST ALERT: owen101labcolor;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1469829929] HOST ALERT: sillpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1469829929] HOST ALERT: owen101lab2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1469829929] HOST ALERT: eateriespk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1469829929] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1469829929] SERVICE ALERT: hallegroundcolor;Yellow Cartridge HP CE262A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1469829930] HOST ALERT: hallefirstcolor;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1469829933] SERVICE ALERT: ADRODC-LIV;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469829934] SERVICE ALERT: ADRODC-LIV;Service: Distributed File System Replication;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469829937] SERVICE ALERT: ADRODC-LIV;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469829943] HOST ALERT: strongpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1469829943] SERVICE ALERT: ADRODC-DET;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469829949] HOST ALERT: halleground3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1469829949] HOST ALERT: prayharroldpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1469829950] SERVICE ALERT: BBTS_DB;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469829951] HOST ALERT: BBTS_DB;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1469829953] HOST ALERT: hallesouthcommons;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1469829953] HOST ALERT: hallefirst3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1469829953] HOST ALERT: ADRODC-LIV;DOWN;SOFT;2;(Host check timed out after 30.01 seconds)
[1469829955] HOST ALERT: hallefirst2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1469829957] SERVICE ALERT: strongpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.130.98. Is the printer off?  Is the network down?
[1469829958] SERVICE ALERT: owen101labcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1469829958] HOST ALERT: hallegroundcolor;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1469829958] SERVICE ALERT: rackhampk;Black Cartridge HP CC364X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.141.212. Is the printer off?  Is the network down?
[1469829959] HOST ALERT: marshallpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1469829959] SERVICE ALERT: hallefirstcolor;Toner Collection Unit HP CE265A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1469829959] HOST ALERT: alexanderpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1469829959] SERVICE ALERT: owen101labcolor;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1469829959] SERVICE ALERT: sillpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1469829959] SERVICE ALERT: eateriespk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.59.141. Is the printer off?  Is the network down?
[1469829960] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.150.252. Is the printer off?  Is the network down?
[1469829960] SERVICE ALERT: markjeffersonpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1469829960] SERVICE ALERT: BBTS;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469829968] SERVICE ALERT: sillpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1469829969] SERVICE ALERT: owen101labcolor;Black Cartridge HP CE260X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1469829970] SERVICE ALERT: ADRODC-LIV;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469829972] SERVICE ALERT: BBTS;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469829973] HOST ALERT: ADRODC-DET;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1469829974] HOST ALERT: owen101lab1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1469829982] HOST ALERT: Google;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1469829982] HOST ALERT: porterpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1469829983] SERVICE ALERT: hallesouthcommons;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.111.220. Is the printer off?  Is the network down?
[1469829983] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1469829984] HOST ALERT: ADRODC-WEMU;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1469829985] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1469829987] HOST ALERT: strongpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1469829988] SERVICE ALERT: marshallpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.137.49. Is the printer off?  Is the network down?
[1469829988] HOST ALERT: owen101labcolor;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1469829988] SERVICE ALERT: hallegroundcolor;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1469829988] HOST ALERT: rackhampk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1469829989] HOST ALERT: hallefirstcolor;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1469829989] SERVICE ALERT: alexanderpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.142.94. Is the printer off?  Is the network down?
[1469829989] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1469829989] HOST ALERT: sillpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1469829989] HOST ALERT: eateriespk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1469829989] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1469829990] SERVICE ALERT: hallegroundcolor;Yellow Cartridge HP CE262A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1469829990] HOST ALERT: studentcenterpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1469829990] HOST ALERT: markjeffersonpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1469829990] HOST ALERT: BBTS;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1469829993] SERVICE ALERT: ADRODC-DET;Service: Distributed File System Replication;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469829998] HOST ALERT: crossroadspk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1469829998] HOST ALERT: hallefirst1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1469829999] SERVICE ALERT: hallegroundcolor;Cyan Cartridge HP CE261A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1469830000] HOST ALERT: ADRODC-LIV;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1469830002] HOST ALERT: honorspk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1469830003] SERVICE ALERT: ADRODC-DET;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830004] SERVICE ALERT: ADRODC-LIV;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830006] SERVICE ALERT: ADRODC-LIV;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830007] SERVICE ALERT: BBTS;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830007] SERVICE ALERT: ADRODC-LIV;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830007] SERVICE ALERT: ADRODC-DET;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830009] SERVICE ALERT: ADRODC-LIV;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830013] HOST ALERT: hallesouthcommons;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1469830013] HOST ALERT: hallefirst3;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1469830013] SERVICE ALERT: ADRODC-DET;Service: MDT Monitor;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830013] SERVICE ALERT: ADRODC-DET;Service: PXE Boot;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830014] SERVICE ALERT: eateriespk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.59.141. Is the printer off?  Is the network down?
[1469830015] HOST ALERT: prayharroldpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1469830015] SERVICE ALERT: ADRODC-DET;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830017] SERVICE ALERT: strongpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.130.98. Is the printer off?  Is the network down?
[1469830018] HOST ALERT: marshallpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1469830018] SERVICE ALERT: hallefirstcolor;Toner Collection Unit HP CE265A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1469830018] SERVICE ALERT: owen101labcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1469830018] SERVICE ALERT: hallesecondsouth;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1469830018] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1469830018] HOST ALERT: hallegroundcolor;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1469830018] SERVICE ALERT: rackhampk;Black Cartridge HP CC364X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.141.212. Is the printer off?  Is the network down?
[1469830019] HOST ALERT: alexanderpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1469830019] SERVICE ALERT: owen101labcolor;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1469830019] SERVICE ALERT: sillpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1469830019] HOST ALERT: owen101lab2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1469830019] SERVICE ALERT: eateriespk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.59.141. Is the printer off?  Is the network down?
[1469830019] HOST ALERT: hallefirstcolor;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1469830020] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.150.252. Is the printer off?  Is the network down?
[1469830020] SERVICE ALERT: markjeffersonpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1469830020] SERVICE ALERT: ADRODC-LIV;Service: MDT Monitor;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830020] SERVICE ALERT: ADRODC-WEMU;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830023] HOST ALERT: ADRODC-DET;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1469830023] SERVICE ALERT: ADRODC-LIV;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830027] SERVICE ALERT: ADRODC-DET;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830028] SERVICE ALERT: sillpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1469830029] SERVICE ALERT: owen101labcolor;Magenta Cartridge HP CE263A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1469830029] SERVICE ALERT: owen101labcolor;Black Cartridge HP CE260X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1469830031] SERVICE ALERT: ADRODC-LIV;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830034] HOST ALERT: ADRODC-LIV;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1469830036] SERVICE ALERT: ADRODC-DET;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830037] HOST ALERT: BBTS;DOWN;SOFT;3;(Host check timed out after 30.00 seconds)
[1469830039] HOST ALERT: halleground3;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1469830039] SERVICE ALERT: hallesecondsouth;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1469830041] HOST ALERT: mckennypk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1469830041] HOST ALERT: BBTS_DB;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1469830041] SERVICE ALERT: BBTS_DB;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830042] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1469830043] SERVICE ALERT: hallesouthcommons;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.111.220. Is the printer off?  Is the network down?
[1469830043] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1469830043] SERVICE ALERT: ADRODC-LIV;Service: Key Distribution Center;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830043] SERVICE ALERT: ADRODC-DET;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830043] SERVICE ALERT: ADRODC-LIV;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830043] SERVICE ALERT: ADRODC-DET;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830043] SERVICE ALERT: ADRODC-LIV;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830043] SERVICE ALERT: ADRODC-DET;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830043] SERVICE ALERT: ADRODC-DET;Service: AD Web Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830043] SERVICE ALERT: ADRODC-DET;Service: Remote Management;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830043] SERVICE ALERT: ADRODC-DET;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830043] SERVICE ALERT: ADRODC-DET;Service: AD Domain Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830043] SERVICE ALERT: ADRODC-DET;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830044] HOST ALERT: eateriespk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1469830045] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1469830045] HOST ALERT: hallefirst2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1469830045] SERVICE ALERT: hallefirst3;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1469830046] SERVICE ALERT: owen101labcolor;Toner Collection Unit HP CE265A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1469830047] HOST ALERT: strongpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1469830047] HOST ALERT: rooseveltpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1469830048] SERVICE ALERT: marshallpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.137.49. Is the printer off?  Is the network down?
[1469830048] HOST ALERT: owen101labcolor;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1469830048] HOST ALERT: hallesecondsouth;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1469830048] HOST ALERT: prayharroldpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1469830048] SERVICE ALERT: hallegroundcolor;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1469830048] HOST ALERT: rackhampk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1469830048] SERVICE ALERT: ADRODC-DET;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830048] SERVICE ALERT: ADRODC-LIV;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830049] SERVICE ALERT: hallegroundcolor;Yellow Cartridge HP CE262A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1469830049] SERVICE ALERT: alexanderpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.142.94. Is the printer off?  Is the network down?
[1469830049] SERVICE ALERT: crossroadspk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1469830049] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1469830049] SERVICE ALERT: halleground3;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1469830049] HOST ALERT: sillpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1469830049] SERVICE ALERT: owen101labcolor;Image Transfer Kit HP CE249A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1469830050] HOST ALERT: studentcenterpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1469830050] HOST ALERT: markjeffersonpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1469830050] HOST ALERT: ADRODC-WEMU;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1469830053] SERVICE ALERT: ADRODC-LIV;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830054] SERVICE ALERT: ADRODC-LIV;Service: Distributed File System Replication;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830055] SERVICE ALERT: ADRODC-LIV;Service: PXE Boot;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830057] HOST ALERT: ADRODC-DET;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1469830057] SERVICE ALERT: ADRODC-LIV;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830058] SERVICE ALERT: ADRODC-LIV;Service: DFS Namespace;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830059] SERVICE ALERT: hallegroundcolor;Toner Collection Unit HP CE265A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1469830059] SERVICE ALERT: hallegroundcolor;Cyan Cartridge HP CE261A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1469830061] SERVICE ALERT: BBTS_DB;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830061] SERVICE ALERT: ADRODC-LIV;Service: Intersite Messaging;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830063] SERVICE ALERT: ADRODC-DET;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830064] HOST ALERT: owen101lab1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1469830064] HOST ALERT: halleground2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1469830070] SERVICE ALERT: sillpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1469830070] SERVICE ALERT: BBTS_DB;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830071] HOST ALERT: BBTS_DB;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1469830072] HOST ALERT: Google;DOWN;HARD;3;PING CRITICAL - Packet loss = 100%
[1469830072] HOST ALERT: porterpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1469830073] HOST ALERT: hallesouthcommons;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1469830073] HOST ALERT: hallefirst3;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1469830073] HOST ALERT: ADRODC-LIV;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1469830074] SERVICE ALERT: eateriespk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.59.141. Is the printer off?  Is the network down?
[1469830077] SERVICE ALERT: strongpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.130.98. Is the printer off?  Is the network down?
[1469830078] HOST ALERT: marshallpk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1469830078] SERVICE ALERT: hallesecondsouth;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1469830078] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1469830078] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1469830078] HOST ALERT: hallegroundcolor;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1469830079] HOST ALERT: alexanderpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1469830079] SERVICE ALERT: owen101labcolor;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1469830079] HOST ALERT: crossroadspk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1469830079] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1469830079] HOST ALERT: prayharroldpk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1469830079] HOST ALERT: halleground3;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1469830079] HOST ALERT: owen101labcolor;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1469830080] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.150.252. Is the printer off?  Is the network down?
[1469830080] SERVICE ALERT: markjeffersonpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1469830080] SERVICE ALERT: BBTS;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830083] SERVICE ALERT: ADRODC-DET;AD Replication;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 30 seconds.
[1469830088] SERVICE ALERT: sillpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1469830088] HOST ALERT: hallefirst1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1469830088] SERVICE ALERT: ADRODC-LIV;Service: AD Domain Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830089] SERVICE ALERT: strongpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.130.98. Is the printer off?  Is the network down?
[1469830089] SERVICE ALERT: hallefirst3;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1469830089] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1469830089] SERVICE ALERT: owen101labcolor;Magenta Cartridge HP CE263A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1469830090] SERVICE ALERT: mckennypk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1469830090] SERVICE ALERT: ADRODC-LIV;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830091] SERVICE ALERT: BBTS_DB;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830092] HOST ALERT: honorspk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1469830092] SERVICE ALERT: BBTS;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830093] HOST ALERT: ADRODC-DET;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1469830099] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1469830099] SERVICE ALERT: hallesecondsouth;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1469830100] HOST ALERT: sillpk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1469830102] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1469830103] SERVICE ALERT: BBTS_DB;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830104] HOST ALERT: eateriespk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1469830105] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1469830105] SERVICE ALERT: hallefirst3;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1469830106] SERVICE ALERT: owen101labcolor;Toner Collection Unit HP CE265A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1469830107] HOST ALERT: strongpk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1469830108] SERVICE ALERT: eateriespk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.59.141. Is the printer off?  Is the network down?
[1469830108] HOST ALERT: hallesecondsouth;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1469830108] HOST ALERT: hallefirst2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1469830109] SERVICE ALERT: crossroadspk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1469830109] HOST ALERT: halleground1;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1469830109] HOST ALERT: owen101lab2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1469830109] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1469830109] SERVICE ALERT: halleground3;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1469830109] HOST ALERT: hallefirstcolor;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1469830110] SERVICE ALERT: owen101labcolor;Image Transfer Kit HP CE249A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1469830110] HOST ALERT: studentcenterpk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1469830110] HOST ALERT: markjeffersonpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1469830110] HOST ALERT: BBTS;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1469830113] SERVICE ALERT: ADRODC-DET;Service: Distributed File System Replication;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830118] HOST ALERT: ADRODC-LIV;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1469830119] HOST ALERT: hallefirst3;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1469830119] HOST ALERT: mckennypk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1469830119] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1469830119] HOST ALERT: owen101labcolor;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1469830119] SERVICE ALERT: hallegroundcolor;Toner Collection Unit HP CE265A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1469830119] SERVICE ALERT: hallegroundcolor;Cyan Cartridge HP CE261A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1469830121] HOST ALERT: BBTS_DB;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1469830123] SERVICE ALERT: ADRODC-DET;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830124] SERVICE ALERT: ADRODC-LIV;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830126] SERVICE ALERT: ADRODC-LIV;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830127] SERVICE ALERT: BBTS;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830127] SERVICE ALERT: ADRODC-LIV;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830127] SERVICE ALERT: ADRODC-DET;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830127] SERVICE ALERT: ADRODC-LIV;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830129] SERVICE ALERT: sillpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1469830129] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1469830129] HOST ALERT: hallefirst1;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1469830129] SERVICE ALERT: ADRODC-LIV;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830132] HOST ALERT: owen101lab1;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1469830133] SERVICE ALERT: ADRODC-DET;Service: MDT Monitor;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830133] SERVICE ALERT: ADRODC-DET;Service: PXE Boot;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830134] SERVICE ALERT: eateriespk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.59.141. Is the printer off?  Is the network down?
[1469830135] HOST ALERT: prayharroldpk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1469830135] SERVICE ALERT: ADRODC-DET;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830137] HOST ALERT: rooseveltpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1469830138] HOST ALERT: eateriespk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1469830138] SERVICE ALERT: strongpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.130.98. Is the printer off?  Is the network down?
[1469830138] SERVICE ALERT: hallesecondsouth;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1469830138] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1469830138] HOST ALERT: rackhampk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1469830138] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1469830139] HOST ALERT: crossroadspk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1469830139] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1469830139] HOST ALERT: halleground3;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1469830140] HOST ALERT: ADRODC-WEMU;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1469830140] SERVICE ALERT: ADRODC-LIV;Service: MDT Monitor;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830140] SERVICE ALERT: ADRODC-WEMU;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830140] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1469830141] SERVICE ALERT: BBTS_DB;Size: IFINDX03.DBF;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1469830143] SERVICE ALERT: ADRODC-LIV;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830143] HOST ALERT: ADRODC-DET;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1469830147] SERVICE ALERT: ADRODC-DET;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830148] SERVICE ALERT: strongpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.130.98. Is the printer off?  Is the network down?
[1469830148] SERVICE ALERT: halleground3;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1469830149] SERVICE ALERT: hallefirst3;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1469830149] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1469830149] HOST ALERT: hallefirst2;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1469830149] SERVICE ALERT: owen101labcolor;Magenta Cartridge HP CE263A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1469830149] HOST ALERT: hallegroundcolor;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1469830150] SERVICE ALERT: mckennypk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1469830151] SERVICE ALERT: ADRODC-LIV;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830154] HOST ALERT: halleground2;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1469830154] HOST ALERT: ADRODC-LIV;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1469830156] SERVICE ALERT: ADRODC-DET;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830157] HOST ALERT: BBTS;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1469830158] SERVICE ALERT: rooseveltpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.239.219. Is the printer off?  Is the network down?
[1469830159] HOST ALERT: sillpk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1469830159] HOST ALERT: hallesecondsouth;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1469830159] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1469830159] SERVICE ALERT: hallesecondsouth;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1469830161] SERVICE ALERT: BBTS_DB;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830162] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1469830162] HOST ALERT: porterpk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1469830163] HOST ALERT: hallesouthcommons;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1469830163] SERVICE ALERT: ADRODC-LIV;Service: Key Distribution Center;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830163] SERVICE ALERT: ADRODC-DET;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830163] SERVICE ALERT: ADRODC-LIV;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830163] SERVICE ALERT: ADRODC-DET;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830163] SERVICE ALERT: ADRODC-LIV;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830163] SERVICE ALERT: ADRODC-DET;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830163] SERVICE ALERT: ADRODC-DET;Service: AD Web Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830163] SERVICE ALERT: ADRODC-DET;Service: Remote Management;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830163] SERVICE ALERT: ADRODC-DET;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830163] SERVICE ALERT: ADRODC-DET;Service: AD Domain Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830163] SERVICE ALERT: ADRODC-DET;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830165] SERVICE ALERT: hallefirst3;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1469830166] SERVICE ALERT: owen101labcolor;Toner Collection Unit HP CE265A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1469830168] HOST ALERT: marshallpk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1469830168] HOST ALERT: strongpk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1469830168] SERVICE ALERT: eateriespk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.59.141. Is the printer off?  Is the network down?
[1469830168] HOST ALERT: prayharroldpk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1469830168] SERVICE ALERT: ADRODC-DET;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830168] SERVICE ALERT: ADRODC-LIV;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830169] SERVICE ALERT: owen101labcolor;Image Transfer Kit HP CE249A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1469830169] HOST ALERT: alexanderpk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1469830169] SERVICE ALERT: crossroadspk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1469830169] SERVICE ALERT: halleground1;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1469830169] SERVICE ALERT: halleground3;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1469830170] HOST ALERT: ADRODC-WEMU;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1469830170] HOST ALERT: hallefirstcolor;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1469830171] HOST ALERT: BBTS_DB;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1469830173] SERVICE ALERT: ADRODC-LIV;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830174] SERVICE ALERT: ADRODC-LIV;Service: Distributed File System Replication;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830175] SERVICE ALERT: ADRODC-LIV;Service: PXE Boot;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830177] HOST ALERT: ADRODC-DET;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1469830177] SERVICE ALERT: ADRODC-LIV;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830177] SERVICE ALERT: BBTS;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830178] HOST ALERT: halleground3;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1469830178] SERVICE ALERT: ADRODC-LIV;Service: DFS Namespace;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830179] HOST ALERT: hallefirst3;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1469830179] HOST ALERT: mckennypk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1469830179] SERVICE ALERT: rooseveltpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.239.219. Is the printer off?  Is the network down?
[1469830179] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1469830179] SERVICE ALERT: owen101labcolor;Yellow Cartridge HP CE262A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1469830179] HOST ALERT: owen101labcolor;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1469830179] SERVICE ALERT: hallegroundcolor;Toner Collection Unit HP CE265A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1469830180] SERVICE ALERT: BBTS;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830181] SERVICE ALERT: BBTS_DB;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830181] SERVICE ALERT: ADRODC-LIV;Service: Intersite Messaging;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830182] HOST ALERT: honorspk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1469830183] SERVICE ALERT: ADRODC-DET;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830187] SERVICE ALERT: ADRODC-LIV;AD Replication;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 30 seconds.
[1469830188] SERVICE ALERT: sillpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.127.240. Is the printer off?  Is the network down?
[1469830188] HOST ALERT: rooseveltpk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1469830189] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1469830189] HOST ALERT: hallefirst1;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1469830189] HOST ALERT: hallesecondsouth;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1469830190] SERVICE ALERT: BBTS_DB;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830192] HOST ALERT: owen101lab1;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1469830193] HOST ALERT: ADRODC-LIV;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1469830196] SERVICE ALERT: marshallpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.137.49. Is the printer off?  Is the network down?
[1469830198] SERVICE ALERT: strongpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.130.98. Is the printer off?  Is the network down?
[1469830198] HOST ALERT: eateriespk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1469830198] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1469830199] HOST ALERT: crossroadspk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1469830199] SERVICE ALERT: hallefirst2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1469830199] HOST ALERT: halleground1;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1469830199] HOST ALERT: owen101lab2;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1469830200] HOST ALERT: studentcenterpk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1469830200] HOST ALERT: markjeffersonpk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1469830200] SERVICE ALERT: BBTS;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830200] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1469830203] SERVICE ALERT: ADRODC-DET;AD Replication;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 30 seconds.
[1469830207] HOST ALERT: BBTS;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1469830208] SERVICE ALERT: strongpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.130.98. Is the printer off?  Is the network down?
[1469830208] SERVICE ALERT: halleground3;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1469830208] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1469830208] SERVICE ALERT: ADRODC-LIV;Service: AD Domain Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830209] SERVICE ALERT: hallefirst3;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1469830209] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1469830209] HOST ALERT: hallefirst2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1469830209] HOST ALERT: hallegroundcolor;DOWN;SOFT;5;(Host check timed out after 30.01 seconds)
[1469830210] SERVICE ALERT: hallegroundcolor;Black Cartridge HP CE260X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1469830210] SERVICE ALERT: mckennypk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1469830210] SERVICE ALERT: ADRODC-LIV;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830211] HOST ALERT: BBTS_DB;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1469830211] SERVICE ALERT: BBTS_DB;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830212] SERVICE ALERT: BBTS;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830213] HOST ALERT: ADRODC-DET;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1469830218] HOST ALERT: sillpk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1469830219] SERVICE ALERT: rooseveltpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.239.219. Is the printer off?  Is the network down?
[1469830219] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1469830219] SERVICE ALERT: marshallpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.137.49. Is the printer off?  Is the network down?
[1469830221] SERVICE ALERT: hallefirstcolor;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1469830223] SERVICE ALERT: BBTS_DB;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830226] HOST ALERT: marshallpk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1469830228] HOST ALERT: strongpk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1469830228] SERVICE ALERT: eateriespk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.59.141. Is the printer off?  Is the network down?
[1469830228] HOST ALERT: rackhampk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1469830229] SERVICE ALERT: halleground1;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1469830230] HOST ALERT: hallefirstcolor;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1469830233] SERVICE ALERT: ADRODC-DET;Service: Distributed File System Replication;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830238] HOST ALERT: halleground3;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1469830238] HOST ALERT: hallefirst1;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1469830238] HOST ALERT: ADRODC-LIV;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1469830239] HOST ALERT: hallefirst3;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1469830239] SERVICE ALERT: rackhampk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.141.212. Is the printer off?  Is the network down?
[1469830239] HOST ALERT: mckennypk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1469830239] SERVICE ALERT: rooseveltpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.239.219. Is the printer off?  Is the network down?
[1469830239] SERVICE ALERT: hallefirst2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.163. Is the printer off?  Is the network down?
[1469830239] SERVICE ALERT: owen101labcolor;Yellow Cartridge HP CE262A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1469830240] HOST ALERT: hallegroundcolor;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1469830241] HOST ALERT: BBTS_DB;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1469830242] HOST ALERT: BBTS;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1469830243] SERVICE ALERT: ADRODC-DET;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830244] HOST ALERT: halleground2;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1469830247] SERVICE ALERT: ADRODC-LIV;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830247] SERVICE ALERT: ADRODC-DET;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830247] SERVICE ALERT: ADRODC-LIV;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830249] HOST ALERT: rooseveltpk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1469830249] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.107.53. Is the printer off?  Is the network down?
[1469830249] SERVICE ALERT: ADRODC-LIV;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830252] HOST ALERT: porterpk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1469830253] HOST ALERT: hallesouthcommons;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1469830253] SERVICE ALERT: ADRODC-DET;Service: MDT Monitor;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830255] SERVICE ALERT: ADRODC-DET;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830256] SERVICE ALERT: marshallpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.137.49. Is the printer off?  Is the network down?
[1469830258] HOST ALERT: prayharroldpk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1469830258] SERVICE ALERT: strongpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.130.98. Is the printer off?  Is the network down?
[1469830258] HOST ALERT: eateriespk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1469830259] HOST ALERT: alexanderpk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1469830259] HOST ALERT: halleground1;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1469830259] SERVICE ALERT: hallegroundcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1469830259] SERVICE ALERT: BBTS_DB;Size: TS_DATA02.DBF;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1469830259] SERVICE ALERT: honorspk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.125.96. Is the printer off?  Is the network down?
[1469830259] SERVICE ALERT: rackhampk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.141.212. Is the printer off?  Is the network down?
[1469830260] SERVICE ALERT: ADRODC-WEMU;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830260] HOST ALERT: ADRODC-WEMU;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1469830260] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1469830263] SERVICE ALERT: ADRODC-LIV;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830263] HOST ALERT: ADRODC-DET;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1469830267] SERVICE ALERT: ADRODC-DET;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830268] SERVICE ALERT: halleground3;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1469830268] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1469830269] HOST ALERT: rackhampk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1469830269] HOST ALERT: hallefirst2;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1469830269] SERVICE ALERT: BBTS_DB;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830269] HOST ALERT: owen101labcolor;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1469830270] SERVICE ALERT: hallegroundcolor;Black Cartridge HP CE260X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1469830270] SERVICE ALERT: BBTS_DB;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830270] SERVICE ALERT: ADRODC-WEMU;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830270] SERVICE ALERT: BBTS;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830271] SERVICE ALERT: ADRODC-LIV;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830271] SERVICE ALERT: BBTS;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830271] SERVICE ALERT: BBTS_DB;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830272] HOST ALERT: honorspk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1469830276] SERVICE ALERT: hallegroundcolor;Image Transfer Kit HP CE249A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1469830276] SERVICE ALERT: ADRODC-DET;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830277] HOST ALERT: ADRODC-LIV;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1469830278] SERVICE ALERT: rooseveltpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.239.219. Is the printer off?  Is the network down?
[1469830279] SERVICE ALERT: alexanderpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.142.94. Is the printer off?  Is the network down?
[1469830279] SERVICE ALERT: marshallpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.137.49. Is the printer off?  Is the network down?
[1469830279] HOST ALERT: hallesecondsouth;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1469830281] SERVICE ALERT: hallefirstcolor;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1469830282] SERVICE ALERT: porterpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1469830282] HOST ALERT: owen101lab1;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1469830283] SERVICE ALERT: ADRODC-DET;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830283] SERVICE ALERT: ADRODC-LIV;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830283] SERVICE ALERT: ADRODC-DET;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830283] SERVICE ALERT: ADRODC-LIV;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830283] SERVICE ALERT: ADRODC-DET;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830283] SERVICE ALERT: ADRODC-DET;Service: AD Web Service;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830283] SERVICE ALERT: ADRODC-DET;Service: Remote Management;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830283] SERVICE ALERT: ADRODC-DET;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830283] SERVICE ALERT: ADRODC-DET;Service: AD Domain Service;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830283] SERVICE ALERT: ADRODC-DET;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830286] HOST ALERT: marshallpk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1469830288] HOST ALERT: strongpk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1469830288] SERVICE ALERT: halleground2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1469830288] SERVICE ALERT: studentcenterpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.150.252. Is the printer off?  Is the network down?
[1469830288] SERVICE ALERT: ADRODC-DET;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830288] SERVICE ALERT: ADRODC-LIV;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830289] HOST ALERT: crossroadspk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1469830289] SERVICE ALERT: halleground1;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1469830289] HOST ALERT: hallegroundcolor;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1469830289] HOST ALERT: BBTS_DB;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1469830289] HOST ALERT: owen101lab2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1469830290] HOST ALERT: studentcenterpk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1469830290] HOST ALERT: markjeffersonpk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1469830290] HOST ALERT: hallefirstcolor;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1469830295] SERVICE ALERT: ADRODC-LIV;Service: PXE Boot;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830297] HOST ALERT: ADRODC-DET;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1469830297] SERVICE ALERT: BBTS;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830298] HOST ALERT: halleground3;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1469830298] HOST ALERT: hallefirst1;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1469830298] SERVICE ALERT: ADRODC-LIV;Service: DFS Namespace;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830299] SERVICE ALERT: hallegroundcolor;Magenta Cartridge HP CE263A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1469830299] SERVICE ALERT: rackhampk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.141.212. Is the printer off?  Is the network down?
[1469830299] SERVICE ALERT: rooseveltpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.239.219. Is the printer off?  Is the network down?
[1469830299] SERVICE ALERT: owen101labcolor;Yellow Cartridge HP CE262A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1469830300] SERVICE ALERT: BBTS;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830300] HOST ALERT: ADRODC-WEMU;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1469830300] HOST ALERT: BBTS;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1469830301] SERVICE ALERT: BBTS_DB;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830301] SERVICE ALERT: ADRODC-LIV;Service: Intersite Messaging;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830306] SERVICE ALERT: ADRODC-LIV;AD Replication;CRITICAL;SOFT;2;CHECK_NRPE: Socket timeout after 30 seconds.
[1469830308] HOST ALERT: sillpk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1469830308] HOST ALERT: rooseveltpk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1469830309] HOST ALERT: alexanderpk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1469830312] HOST ALERT: porterpk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1469830316] SERVICE ALERT: marshallpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.137.49. Is the printer off?  Is the network down?
[1469830317] SERVICE ALERT: BBTS_DB;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830318] HOST ALERT: halleground2;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1469830319] SERVICE ALERT: studentcenterpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.150.252. Is the printer off?  Is the network down?
[1469830319] HOST ALERT: halleground1;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1469830319] SERVICE ALERT: hallefirstcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1469830319] SERVICE ALERT: hallegroundcolor;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1469830319] SERVICE ALERT: hallegroundcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1469830319] SERVICE ALERT: crossroadspk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1469830319] SERVICE ALERT: honorspk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.125.96. Is the printer off?  Is the network down?
[1469830319] SERVICE ALERT: honorspk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.125.96. Is the printer off?  Is the network down?
[1469830319] SERVICE ALERT: rackhampk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.141.212. Is the printer off?  Is the network down?
[1469830323] SERVICE ALERT: ADRODC-DET;AD Replication;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 30 seconds.
[1469830328] SERVICE ALERT: hallefirstcolor;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1469830328] SERVICE ALERT: hallefirst1;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1469830328] SERVICE ALERT: owen101labcolor;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1469830328] SERVICE ALERT: ADRODC-LIV;Service: AD Domain Service;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830329] HOST ALERT: hallefirst3;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1469830329] HOST ALERT: hallegroundcolor;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1469830329] SERVICE ALERT: owen101labcolor;Cyan Cartridge HP CE261A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1469830329] HOST ALERT: rackhampk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1469830329] HOST ALERT: mckennypk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1469830329] SERVICE ALERT: owen101lab1;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1469830329] SERVICE ALERT: ADRODC-WEMU;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830329] HOST ALERT: owen101labcolor;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1469830330] SERVICE ALERT: hallegroundcolor;Black Cartridge HP CE260X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1469830330] SERVICE ALERT: owen101lab2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1469830330] SERVICE ALERT: ADRODC-WEMU;Service: Key Distribution Center;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830330] SERVICE ALERT: BBTS_DB;Size: IFINDX03.DBF;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1469830331] HOST ALERT: BBTS_DB;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1469830331] SERVICE ALERT: BBTS_DB;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830336] SERVICE ALERT: hallegroundcolor;Image Transfer Kit HP CE249A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1469830337] SERVICE ALERT: BBTS;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830339] SERVICE ALERT: alexanderpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.142.94. Is the printer off?  Is the network down?
[1469830339] SERVICE ALERT: marshallpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.137.49. Is the printer off?  Is the network down?
[1469830341] SERVICE ALERT: hallefirstcolor;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1469830342] SERVICE ALERT: porterpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1469830343] HOST ALERT: hallesouthcommons;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1469830343] SERVICE ALERT: BBTS_DB;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830346] HOST ALERT: marshallpk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1469830348] HOST ALERT: prayharroldpk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1469830348] HOST ALERT: eateriespk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1469830348] SERVICE ALERT: halleground2;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1469830348] SERVICE ALERT: studentcenterpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.150.252. Is the printer off?  Is the network down?
[1469830349] HOST ALERT: studentcenterpk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1469830349] HOST ALERT: hallefirstcolor;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1469830349] HOST ALERT: crossroadspk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1469830349] SERVICE ALERT: markjeffersonpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1469830349] HOST ALERT: honorspk;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1469830357] SERVICE ALERT: crossroadspk;Black Cartridge HP CC364X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1469830358] HOST ALERT: hallefirst1;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1469830359] SERVICE ALERT: hallegroundcolor;Magenta Cartridge HP CE263A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1469830359] SERVICE ALERT: rackhampk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.141.212. Is the printer off?  Is the network down?
[1469830359] HOST ALERT: owen101lab1;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1469830359] HOST ALERT: hallefirst2;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1469830359] HOST ALERT: ADRODC-WEMU;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1469830360] HOST ALERT: hallegroundcolor;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1469830360] HOST ALERT: owen101lab2;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1469830361] HOST ALERT: BBTS_DB;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1469830363] SERVICE ALERT: ADRODC-WEMU;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830363] SERVICE ALERT: ADRODC-WEMU;Service: AD Domain Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830363] SERVICE ALERT: ADRODC-WEMU;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830363] SERVICE ALERT: BBTS_DB;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830363] SERVICE ALERT: BBTS_DB;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830363] SERVICE ALERT: ADRODC-WEMU;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830367] HOST ALERT: BBTS;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1469830367] SERVICE ALERT: ADRODC-LIV;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830368] SERVICE ALERT: ADRODC-WEMU;Service: Intersite Messaging;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830369] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1469830369] SERVICE ALERT: rooseveltpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.239.219. Is the printer off?  Is the network down?
[1469830369] HOST ALERT: alexanderpk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1469830369] HOST ALERT: hallesecondsouth;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1469830372] HOST ALERT: porterpk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1469830372] SERVICE ALERT: ADRODC-WEMU;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830373] SERVICE ALERT: owen101lab1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1469830373] SERVICE ALERT: ADRODC-WEMU;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830373] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1469830373] SERVICE ALERT: ADRODC-LIV;Service: Remote Management;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830374] SERVICE ALERT: ADRODC-WEMU;Service: DNS Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830374] SERVICE ALERT: ADRODC-WEMU;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830376] SERVICE ALERT: ADRODC-WEMU;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830377] SERVICE ALERT: ADRODC-WEMU;Service: Distributed File System Replication;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830377] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1469830378] HOST ALERT: strongpk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1469830378] HOST ALERT: halleground2;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1469830379] SERVICE ALERT: rackhampk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.141.212. Is the printer off?  Is the network down?
[1469830379] SERVICE ALERT: studentcenterpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.150.252. Is the printer off?  Is the network down?
[1469830379] SERVICE ALERT: hallefirstcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1469830379] SERVICE ALERT: hallegroundcolor;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1469830379] SERVICE ALERT: crossroadspk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1469830379] HOST ALERT: markjeffersonpk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1469830379] SERVICE ALERT: hallegroundcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1469830379] SERVICE ALERT: ADRODC-LIV;Service: AD Web Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830379] SERVICE ALERT: alexanderpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.142.94. Is the printer off?  Is the network down?
[1469830379] SERVICE ALERT: honorspk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.125.96. Is the printer off?  Is the network down?
[1469830379] SERVICE ALERT: honorspk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.125.96. Is the printer off?  Is the network down?
[1469830380] SERVICE ALERT: ADRODC-WEMU;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830381] SERVICE ALERT: ADRODC-WEMU;Service: Remote Management;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830381] SERVICE ALERT: ADRODC-WEMU;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830381] SERVICE ALERT: ADRODC-WEMU;Service: DFS Namespace;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830382] SERVICE ALERT: ADRODC-WEMU;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830384] SERVICE ALERT: ADRODC-LIV;Service: DNS Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830384] SERVICE ALERT: ADRODC-DET;Service: DNS Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830384] SERVICE ALERT: ADRODC-DET;Service: Key Distribution Center;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830387] HOST ALERT: crossroadspk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1469830388] HOST ALERT: halleground3;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1469830388] SERVICE ALERT: hallefirstcolor;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1469830388] SERVICE ALERT: owen101labcolor;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1469830388] SERVICE ALERT: ADRODC-DET;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830388] SERVICE ALERT: ADRODC-DET;Service: Intersite Messaging;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830389] SERVICE ALERT: owen101labcolor;Cyan Cartridge HP CE261A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1469830389] HOST ALERT: rackhampk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1469830389] SERVICE ALERT: owen101lab1;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1469830389] SERVICE ALERT: BBTS_DB;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830389] SERVICE ALERT: hallesouthcommons;Black Cartridge HP CC364X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.111.220. Is the printer off?  Is the network down?
[1469830390] SERVICE ALERT: owen101lab2;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1469830390] SERVICE ALERT: BBTS;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830390] SERVICE ALERT: BBTS_DB;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830390] SERVICE ALERT: ADRODC-WEMU;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830390] SERVICE ALERT: BBTS;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830391] SERVICE ALERT: BBTS;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830391] SERVICE ALERT: BBTS_DB;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830393] HOST ALERT: ADRODC-WEMU;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1469830393] SERVICE ALERT: BBTS;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830393] SERVICE ALERT: ADRODC-DET;Service: DFS Namespace;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830396] SERVICE ALERT: hallegroundcolor;Image Transfer Kit HP CE249A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1469830398] HOST ALERT: sillpk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1469830398] HOST ALERT: rooseveltpk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1469830399] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1469830399] HOST ALERT: hallefirstcolor;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1469830399] SERVICE ALERT: alexanderpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.142.94. Is the printer off?  Is the network down?
[1469830399] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1469830399] SERVICE ALERT: honorspk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.125.96. Is the printer off?  Is the network down?
[1469830400] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1469830402] SERVICE ALERT: porterpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1469830403] HOST ALERT: owen101lab1;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1469830407] HOST ALERT: halleground1;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1469830408] SERVICE ALERT: halleground2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1469830408] SERVICE ALERT: studentcenterpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.150.252. Is the printer off?  Is the network down?
[1469830409] HOST ALERT: studentcenterpk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1469830409] SERVICE ALERT: owen101lab2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1469830409] HOST ALERT: hallegroundcolor;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1469830409] SERVICE ALERT: markjeffersonpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1469830409] HOST ALERT: alexanderpk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1469830409] HOST ALERT: honorspk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1469830411] SERVICE ALERT: BBTS;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830417] SERVICE ALERT: BBTS;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830417] SERVICE ALERT: crossroadspk;Black Cartridge HP CC364X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1469830418] HOST ALERT: owen101labcolor;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1469830419] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1469830419] HOST ALERT: hallefirst3;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1469830419] SERVICE ALERT: hallegroundcolor;Magenta Cartridge HP CE263A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1469830419] HOST ALERT: mckennypk;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1469830419] HOST ALERT: hallesouthcommons;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1469830420] HOST ALERT: owen101lab2;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1469830420] SERVICE ALERT: BBTS;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830420] HOST ALERT: BBTS;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1469830420] SERVICE ALERT: BBTS;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830426] SERVICE ALERT: ADRODC-LIV;AD Replication;CRITICAL;HARD;3;CHECK_NRPE: Socket timeout after 30 seconds.
[1469830428] SERVICE ALERT: porterpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1469830429] HOST ALERT: hallefirst1;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1469830429] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1469830429] SERVICE ALERT: rooseveltpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.239.219. Is the printer off?  Is the network down?
[1469830430] HOST ALERT: halleground3;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1469830432] SERVICE ALERT: owen101lab1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1469830432] HOST ALERT: porterpk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1469830433] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1469830434] SERVICE ALERT: BBTS;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830435] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1469830436] HOST ALERT: marshallpk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1469830437] SERVICE ALERT: BBTS_DB;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830437] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1469830438] HOST ALERT: prayharroldpk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1469830438] HOST ALERT: eateriespk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1469830438] HOST ALERT: halleground2;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1469830438] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1469830439] SERVICE ALERT: studentcenterpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.150.252. Is the printer off?  Is the network down?
[1469830439] SERVICE ALERT: hallefirstcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1469830439] SERVICE ALERT: hallegroundcolor;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1469830439] SERVICE ALERT: crossroadspk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1469830439] HOST ALERT: markjeffersonpk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1469830439] SERVICE ALERT: alexanderpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.142.94. Is the printer off?  Is the network down?
[1469830439] SERVICE ALERT: honorspk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.125.96. Is the printer off?  Is the network down?
[1469830447] HOST ALERT: crossroadspk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1469830447] SERVICE ALERT: BBTS;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830448] SERVICE ALERT: hallefirstcolor;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1469830448] SERVICE ALERT: owen101labcolor;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1469830448] SERVICE ALERT: BBTS_DB;Size: TS_DATA02.DBF;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1469830449] SERVICE ALERT: halleground1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1469830449] SERVICE ALERT: owen101labcolor;Cyan Cartridge HP CE261A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1469830449] SERVICE ALERT: mckennypk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1469830449] SERVICE ALERT: owen101lab1;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1469830449] HOST ALERT: hallefirst2;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1469830449] SERVICE ALERT: ADRODC-WEMU;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830449] SERVICE ALERT: porterpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1469830449] SERVICE ALERT: hallesouthcommons;Black Cartridge HP CC364X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.111.220. Is the printer off?  Is the network down?
[1469830450] SERVICE ALERT: owen101lab2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1469830450] SERVICE ALERT: ADRODC-WEMU;Service: Key Distribution Center;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830451] SERVICE ALERT: BBTS_DB;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830457] SERVICE ALERT: BBTS;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830459] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1469830459] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1469830459] SERVICE ALERT: hallesouthcommons;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.111.220. Is the printer off?  Is the network down?
[1469830459] HOST ALERT: rooseveltpk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1469830459] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1469830459] SERVICE ALERT: honorspk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.125.96. Is the printer off?  Is the network down?
[1469830459] HOST ALERT: hallesecondsouth;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1469830460] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1469830462] HOST ALERT: owen101lab1;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1469830465] SERVICE ALERT: BBTS_DB;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830466] SERVICE ALERT: BBTS_DB;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1469830467] HOST ALERT: halleground1;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1469830468] HOST ALERT: strongpk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1469830469] HOST ALERT: studentcenterpk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1469830469] SERVICE ALERT: owen101lab2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1469830469] SERVICE ALERT: markjeffersonpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1469830469] HOST ALERT: alexanderpk;DOWN;SOFT;9;(Host check timed out after 30.00 seconds)
[1469830469] HOST ALERT: honorspk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1469830477] SERVICE ALERT: crossroadspk;Black Cartridge HP CC364X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.58.6. Is the printer off?  Is the network down?
[1469830478] HOST ALERT: owen101labcolor;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1469830479] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1469830479] HOST ALERT: rackhampk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1469830479] HOST ALERT: mckennypk;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1469830479] HOST ALERT: ADRODC-WEMU;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1469830479] HOST ALERT: porterpk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1469830479] HOST ALERT: hallesouthcommons;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1469830480] HOST ALERT: owen101lab2;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1469830483] SERVICE ALERT: ADRODC-WEMU;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830483] SERVICE ALERT: ADRODC-WEMU;Service: AD Domain Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830483] SERVICE ALERT: ADRODC-WEMU;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830483] SERVICE ALERT: BBTS_DB;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830483] SERVICE ALERT: BBTS_DB;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830483] SERVICE ALERT: ADRODC-WEMU;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830488] HOST ALERT: sillpk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1469830488] SERVICE ALERT: porterpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1469830488] SERVICE ALERT: ADRODC-WEMU;Service: Intersite Messaging;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830489] HOST ALERT: hallefirst1;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1469830489] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1469830489] SERVICE ALERT: rooseveltpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.239.219. Is the printer off?  Is the network down?
[1469830490] HOST ALERT: halleground3;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1469830492] SERVICE ALERT: owen101lab1;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1469830492] SERVICE ALERT: ADRODC-WEMU;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830493] SERVICE ALERT: ADRODC-WEMU;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830493] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1469830493] SERVICE ALERT: ADRODC-LIV;Service: Remote Management;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830494] SERVICE ALERT: ADRODC-WEMU;Service: DNS Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830494] SERVICE ALERT: ADRODC-WEMU;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830495] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1469830496] SERVICE ALERT: ADRODC-WEMU;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830497] SERVICE ALERT: ADRODC-WEMU;Service: Distributed File System Replication;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830497] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1469830498] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1469830499] HOST ALERT: markjeffersonpk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1469830499] SERVICE ALERT: ADRODC-LIV;Service: AD Web Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830499] SERVICE ALERT: alexanderpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.142.94. Is the printer off?  Is the network down?
[1469830500] SERVICE ALERT: ADRODC-WEMU;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830501] SERVICE ALERT: ADRODC-WEMU;Service: Remote Management;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830501] SERVICE ALERT: ADRODC-WEMU;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830501] SERVICE ALERT: ADRODC-WEMU;Service: DFS Namespace;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830502] SERVICE ALERT: ADRODC-WEMU;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830504] SERVICE ALERT: ADRODC-LIV;Service: DNS Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830504] SERVICE ALERT: ADRODC-DET;Service: DNS Service;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830504] SERVICE ALERT: ADRODC-DET;Service: Key Distribution Center;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830507] HOST ALERT: crossroadspk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1469830508] SERVICE ALERT: ADRODC-DET;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830508] SERVICE ALERT: ADRODC-DET;Service: Intersite Messaging;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830509] HOST ALERT: halleground2;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1469830509] SERVICE ALERT: halleground1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1469830509] HOST ALERT: hallefirst3;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1469830509] SERVICE ALERT: mckennypk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1469830509] SERVICE ALERT: BBTS_DB;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830509] SERVICE ALERT: porterpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1469830509] SERVICE ALERT: hallesouthcommons;Black Cartridge HP CC364X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.111.220. Is the printer off?  Is the network down?
[1469830510] SERVICE ALERT: BBTS;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830510] SERVICE ALERT: BBTS_DB;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830510] SERVICE ALERT: ADRODC-WEMU;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830510] SERVICE ALERT: BBTS;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830511] SERVICE ALERT: BBTS;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830511] SERVICE ALERT: BBTS_DB;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830513] HOST ALERT: ADRODC-WEMU;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1469830513] SERVICE ALERT: BBTS;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830513] SERVICE ALERT: ADRODC-DET;Service: DFS Namespace;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830518] HOST ALERT: porterpk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1469830519] SERVICE ALERT: hallefirst1;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.232. Is the printer off?  Is the network down?
[1469830519] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1469830519] SERVICE ALERT: hallesouthcommons;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.111.220. Is the printer off?  Is the network down?
[1469830519] HOST ALERT: rooseveltpk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1469830519] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1469830519] SERVICE ALERT: honorspk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.125.96. Is the printer off?  Is the network down?
[1469830520] SERVICE ALERT: BBTS_DB;Size: IFINDX03.DBF;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1469830520] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.53. Is the printer off?  Is the network down?
[1469830522] HOST ALERT: owen101lab1;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1469830526] HOST ALERT: marshallpk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1469830527] HOST ALERT: halleground1;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1469830528] HOST ALERT: prayharroldpk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1469830528] HOST ALERT: eateriespk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1469830529] SERVICE ALERT: owen101lab2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1469830529] HOST ALERT: alexanderpk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1469830531] SERVICE ALERT: BBTS;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830539] SERVICE ALERT: halleground2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1469830539] HOST ALERT: mckennypk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1469830539] HOST ALERT: hallefirst2;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1469830539] HOST ALERT: hallesouthcommons;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1469830540] SERVICE ALERT: BBTS;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830548] SERVICE ALERT: porterpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1469830549] HOST ALERT: hallefirst1;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1469830549] HOST ALERT: owen101lab2;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1469830549] HOST ALERT: hallesecondsouth;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1469830549] HOST ALERT: honorspk;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1469830554] SERVICE ALERT: BBTS;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830555] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1469830557] SERVICE ALERT: BBTS_DB;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830558] HOST ALERT: strongpk;DOWN;HARD;10;(Host check timed out after 30.01 seconds)
[1469830558] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1469830559] HOST ALERT: studentcenterpk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1469830567] SERVICE ALERT: BBTS;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830569] SERVICE ALERT: halleground1;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.106.51. Is the printer off?  Is the network down?
[1469830569] HOST ALERT: halleground2;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1469830569] HOST ALERT: rackhampk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1469830569] SERVICE ALERT: mckennypk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1469830569] SERVICE ALERT: ADRODC-WEMU;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830569] SERVICE ALERT: porterpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.133.132. Is the printer off?  Is the network down?
[1469830570] SERVICE ALERT: ADRODC-WEMU;Service: Key Distribution Center;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830571] SERVICE ALERT: BBTS_DB;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830577] SERVICE ALERT: BBTS;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830577] HOST ALERT: sillpk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1469830579] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1469830579] SERVICE ALERT: hallesouthcommons;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.111.220. Is the printer off?  Is the network down?
[1469830585] SERVICE ALERT: BBTS_DB;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830586] SERVICE ALERT: BBTS_DB;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1469830588] HOST ALERT: markjeffersonpk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1469830597] HOST ALERT: crossroadspk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1469830599] HOST ALERT: hallefirst3;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1469830599] HOST ALERT: halleground1;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1469830599] HOST ALERT: mckennypk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1469830603] SERVICE ALERT: ADRODC-WEMU;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830603] SERVICE ALERT: ADRODC-WEMU;Service: AD Domain Service;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830603] SERVICE ALERT: ADRODC-WEMU;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830603] SERVICE ALERT: BBTS_DB;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830603] SERVICE ALERT: BBTS_DB;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830603] SERVICE ALERT: ADRODC-WEMU;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830608] SERVICE ALERT: ADRODC-WEMU;Service: Intersite Messaging;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830609] HOST ALERT: rooseveltpk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1469830609] HOST ALERT: hallesouthcommons;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1469830612] HOST ALERT: owen101lab1;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1469830612] SERVICE ALERT: ADRODC-WEMU;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830613] SERVICE ALERT: ADRODC-WEMU;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830613] SERVICE ALERT: ADRODC-LIV;Service: Remote Management;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830613] SERVICE ALERT: ADRODC-WEMU;Service: DNS Service;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830614] SERVICE ALERT: ADRODC-WEMU;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830616] SERVICE ALERT: ADRODC-WEMU;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830617] SERVICE ALERT: ADRODC-WEMU;Service: Distributed File System Replication;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830618] SERVICE ALERT: BBTS_DB;Size: IFDATA03.DBF;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1469830619] SERVICE ALERT: ADRODC-LIV;Service: AD Web Service;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830620] SERVICE ALERT: ADRODC-WEMU;Service: RPC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830621] SERVICE ALERT: ADRODC-WEMU;Service: Remote Management;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830621] SERVICE ALERT: ADRODC-WEMU;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830621] SERVICE ALERT: ADRODC-WEMU;Service: DFS Namespace;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830622] SERVICE ALERT: ADRODC-WEMU;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830624] SERVICE ALERT: ADRODC-LIV;Service: DNS Service;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830624] SERVICE ALERT: ADRODC-DET;Service: DNS Service;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830624] SERVICE ALERT: ADRODC-DET;Service: Key Distribution Center;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830628] SERVICE ALERT: ADRODC-DET;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830628] SERVICE ALERT: ADRODC-DET;Service: Intersite Messaging;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830629] HOST ALERT: hallefirst2;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1469830630] SERVICE ALERT: BBTS;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830633] SERVICE ALERT: BBTS;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830633] SERVICE ALERT: ADRODC-DET;Service: DFS Namespace;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830638] SERVICE ALERT: BBTS_DB;Size: TS_DATA02.DBF;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1469830639] HOST ALERT: owen101lab2;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1469830639] HOST ALERT: hallesecondsouth;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1469830639] HOST ALERT: honorspk;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1469830642] SERVICE ALERT: BBTS_DB;Size: TS_INDEX02.DBF;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1469830651] SERVICE ALERT: BBTS;VMWare Tools;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830659] HOST ALERT: halleground2;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1469830660] SERVICE ALERT: BBTS;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830674] SERVICE ALERT: BBTS;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830678] HOST ALERT: markjeffersonpk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1469830687] SERVICE ALERT: BBTS;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830689] HOST ALERT: halleground1;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1469830689] HOST ALERT: mckennypk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1469830691] SERVICE ALERT: BBTS_DB;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830699] HOST ALERT: rooseveltpk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1469830705] SERVICE ALERT: BBTS_DB;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830706] SERVICE ALERT: BBTS_DB;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1469830729] HOST ALERT: hallesecondsouth;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1469830729] HOST ALERT: honorspk;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1469830749] HOST ALERT: halleground2;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1469830779] HOST ALERT: halleground1;DOWN;HARD;10;(Host check timed out after 30.01 seconds)
[1469830808] SERVICE ALERT: BBTS_DB;Size: IFDATA03.DBF;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1469830832] SERVICE ALERT: BBTS_DB;Size: TS_INDEX02.DBF;CRITICAL;SOFT;2;(Return code of 255 is out of bounds)
[1469830997] SERVICE ALERT: BBTS_DB;Size: IFDATA03.DBF;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1469831021] SERVICE ALERT: BBTS_DB;Size: TS_INDEX02.DBF;CRITICAL;HARD;3;(Return code of 255 is out of bounds)
[1469834024] SERVICE ALERT: BBTS_DB;Size: IFINDX03.DBF;OK;HARD;3;OK: Size is  13540.0078125  Mb
[1469834025] HOST ALERT: hallegroundcolor;UP;HARD;10;PING OK - Packet loss = 28%, RTA = 0.89 ms
[1469834025] HOST ALERT: halleground2;UP;HARD;10;PING OK - Packet loss = 72%, RTA = 0.51 ms
[1469834025] HOST ALERT: markjeffersonpk;UP;HARD;10;PING OK - Packet loss = 68%, RTA = 0.51 ms
[1469834025] HOST ALERT: studentcenterpk;UP;HARD;10;PING OK - Packet loss = 72%, RTA = 0.47 ms
[1469834025] HOST ALERT: porterpk;UP;HARD;10;PING OK - Packet loss = 79%, RTA = 0.50 ms
[1469834026] HOST ALERT: Google;UP;HARD;3;PING WARNING - Packet loss = 80%, RTA = 16.94 ms
[1469834026] HOST ALERT: BBTS_DB;UP;HARD;10;PING WARNING - Packet loss = 82%, RTA = 0.44 ms
[1469834026] HOST ALERT: BBTS;UP;HARD;10;PING OK - Packet loss = 50%, RTA = 0.40 ms
[1469834030] HOST ALERT: ADRODC-LIV;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 7.42 ms
[1469834033] SERVICE ALERT: honorspk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 51% - OK!
[1469834033] SERVICE ALERT: hallegroundcolor;Device Status;OK;HARD;3;Sleep mode on
[1469834033] SERVICE ALERT: crossroadspk;Device Status;OK;HARD;3;Sleep mode on
[1469834033] SERVICE ALERT: hallefirstcolor;Image Fuser Kit HP 110V-CE246A;OK;HARD;3;Image Fuser Kit HP 110V-CE246A, 220V-CE247A is at 100% - OK!
[1469834033] SERVICE ALERT: studentcenterpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1469834037] HOST ALERT: honorspk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 6.37 ms
[1469834037] HOST ALERT: crossroadspk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1469834037] HOST ALERT: hallefirstcolor;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.36 ms
[1469834042] SERVICE ALERT: hallefirstcolor;Device Status;OK;HARD;3;Sleep mode on
[1469834042] SERVICE ALERT: owen101labcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1469834043] SERVICE ALERT: owen101lab1;Device Status;OK;HARD;3;Sleep Mode on
[1469834043] SERVICE ALERT: owen101labcolor;Cyan Cartridge HP CE261A;OK;HARD;3;Cyan Cartridge HP CE261A is at 35% - OK!
[1469834044] SERVICE ALERT: owen101lab2;Device Status;OK;HARD;3;Sleep Mode on
[1469834046] HOST ALERT: owen101labcolor;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1469834047] HOST ALERT: owen101lab1;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.52 ms
[1469834048] HOST ALERT: owen101lab2;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.46 ms
[1469834049] SERVICE ALERT: BBTS_DB;Size: TS_DATA02.DBF;OK;HARD;3;OK: Size is  13924.0078125  Mb
[1469834063] SERVICE ALERT: markjeffersonpk;Device Status;OK;HARD;3;ORDER CARTRIDGE LESS THAN 1400 PAGES 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1469834072] SERVICE ALERT: crossroadspk;Black Cartridge HP CC364X;OK;HARD;3;Black Cartridge HP CC364X is at 85% - OK!
[1469834072] HOST ALERT: prayharroldpk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.55 ms
[1469834083] SERVICE ALERT: rooseveltpk;Device Status;OK;HARD;3;Sleep Mode on
[1469834083] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;OK;HARD;3;Image Transfer Kit HP CE249A is at 90% - OK!
[1469834086] SERVICE ALERT: owen101lab1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1469834087] HOST ALERT: rooseveltpk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1469834087] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 86% - OK!
[1469834088] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;WARNING;SOFT;1;(No output on stdout) stderr: Could not resolve hostname ADWU1: Name or service not known
[1469834089] SERVICE ALERT: BBTS_DB;Service: Windows Time;OK;HARD;3;W32Time: Started
[1469834090] SERVICE ALERT: BBTS_DB;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1469834090] SERVICE ALERT: ADRODC-WEMU;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1469834090] SERVICE ALERT: BBTS;Service: WMI;OK;HARD;3;Winmgmt: Started
[1469834091] SERVICE ALERT: BBTS;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1469834091] SERVICE ALERT: BBTS_DB;Service: RPC;OK;HARD;3;RpcSs: Started
[1469834091] SERVICE ALERT: halleground1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 50% - OK!
[1469834093] SERVICE ALERT: alexanderpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 75% - OK!
[1469834094] HOST ALERT: ADRODC-WEMU;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1469834095] HOST ALERT: halleground1;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.49 ms
[1469834097] HOST ALERT: alexanderpk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.54 ms
[1469834103] SERVICE ALERT: hallesouthcommons;Black Cartridge HP CC364X;OK;HARD;3;Black Cartridge HP CC364X is at 72% - OK!
[1469834107] HOST ALERT: hallesouthcommons;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1469834113] SERVICE ALERT: hallefirst1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1469834113] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 8% - OK!
[1469834114] SERVICE ALERT: honorspk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1469834114] HOST ALERT: mckennypk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.53 ms
[1469834114] SERVICE ALERT: halleground3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 74% - OK!
[1469834117] HOST ALERT: hallefirst1;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1469834118] HOST ALERT: halleground3;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1469834122] HOST ALERT: sillpk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.49 ms
[1469834123] SERVICE ALERT: owen101lab2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1469834133] SERVICE ALERT: halleground2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1469834137] SERVICE ALERT: BBTS_DB;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1469834142] SERVICE ALERT: porterpk;Device Status;OK;HARD;3;Ready
[1469834143] HOST ALERT: hallefirst3;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.49 ms
[1469834149] SERVICE ALERT: ADRODC-WEMU;Uptime;OK;HARD;3;System Uptime - 6 day(s) 23 hour(s) 59 minute(s)
[1469834149] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;OK;HARD;3;Magenta Cartridge HP CE263A is at 79% - OK!
[1469834150] SERVICE ALERT: ADRODC-WEMU;Service: Key Distribution Center;OK;HARD;3;KDC: Started
[1469834152] SERVICE ALERT: markjeffersonpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 5% - OK!
[1469834157] SERVICE ALERT: BBTS;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1469834162] HOST ALERT: eateriespk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.43 ms
[1469834163] SERVICE ALERT: mckennypk;Device Status;OK;HARD;3;Sleep Mode on
[1469834163] SERVICE ALERT: halleground1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1469834163] SERVICE ALERT: porterpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 45% - OK!
[1469834165] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1469834173] HOST ALERT: hallefirst2;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.44 ms
[1469834173] SERVICE ALERT: hallesouthcommons;Device Status;OK;HARD;3;Sleep mode on
[1469834173] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;OK;HARD;3;Black Cartridge HP CE260X is at 43% - OK!
[1469834180] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;OK;SOFT;2;OK: No entries found
[1469834182] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;OK;HARD;3;Yellow Cartridge HP CE262A is at 75% - OK!
[1469834183] SERVICE ALERT: ADRODC-WEMU;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1469834183] SERVICE ALERT: ADRODC-WEMU;Service: AD Domain Service;OK;HARD;3;NTDS: Started
[1469834183] SERVICE ALERT: ADRODC-WEMU;CPU Load;OK;HARD;3;CPU Load 1% (15 min average)
[1469834183] SERVICE ALERT: BBTS_DB;Uptime;OK;HARD;3;System Uptime - 1 day(s) 15 hour(s) 13 minute(s)
[1469834183] SERVICE ALERT: BBTS_DB;C:\ Drive Space;OK;HARD;3;c:\ - total: 144.98 Gb - used: 84.59 Gb (58%) - free 60.39 Gb (42%)
[1469834183] SERVICE ALERT: ADRODC-WEMU;Service: WMI;OK;HARD;3;Winmgmt: Started
[1469834183] HOST ALERT: hallesecondsouth;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.45 ms
[1469834188] SERVICE ALERT: ADRODC-WEMU;Service: Intersite Messaging;OK;HARD;3;IsmServ: Started
[1469834192] HOST ALERT: strongpk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.52 ms
[1469834192] SERVICE ALERT: ADRODC-WEMU;Memory Usage;OK;HARD;3;Memory usage: total:21654.72 MB - used: 1980.84 MB (9%) - free: 19673.88 MB (91%)
[1469834193] SERVICE ALERT: ADRODC-WEMU;Service: Windows Time;OK;HARD;3;W32Time: Started
[1469834193] SERVICE ALERT: ADRODC-LIV;Service: Remote Management;OK;HARD;3;WinRM: Started
[1469834193] SERVICE ALERT: ADRODC-WEMU;Service: DNS Service;OK;HARD;3;DNS: Started
[1469834194] SERVICE ALERT: ADRODC-WEMU;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1469834196] SERVICE ALERT: ADRODC-WEMU;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1469834197] SERVICE ALERT: ADRODC-WEMU;Service: Distributed File System Replication;OK;HARD;3;DFSR: Started
[1469834199] SERVICE ALERT: ADRODC-LIV;Service: AD Web Service;OK;HARD;3;ADWS: Started
[1469834200] SERVICE ALERT: ADRODC-WEMU;Service: RPC;OK;HARD;3;RpcSs: Started
[1469834201] SERVICE ALERT: ADRODC-WEMU;C:\ Drive Space;OK;HARD;3;c:\ - total: 97.31 Gb - used: 12.97 Gb (13%) - free 84.34 Gb (87%)
[1469834201] SERVICE ALERT: ADRODC-WEMU;Service: Remote Management;OK;HARD;3;WinRM: Started
[1469834201] SERVICE ALERT: ADRODC-WEMU;Service: DFS Namespace;OK;HARD;3;Dfs: Started
[1469834202] SERVICE ALERT: ADRODC-WEMU;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1469834204] SERVICE ALERT: ADRODC-LIV;Service: DNS Service;OK;HARD;3;DNS: Started
[1469834204] SERVICE ALERT: ADRODC-DET;Service: DNS Service;OK;HARD;3;DNS: Started
[1469834204] SERVICE ALERT: ADRODC-DET;Service: Key Distribution Center;OK;HARD;3;KDC: Started
[1469834208] SERVICE ALERT: ADRODC-DET;Uptime;OK;HARD;3;System Uptime - 7 day(s) 0 hour(s) 16 minute(s)
[1469834208] SERVICE ALERT: ADRODC-DET;Service: Intersite Messaging;OK;HARD;3;IsmServ: Started
[1469834208] HOST ALERT: ADRODC-DET;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 8.94 ms
[1469834210] SERVICE ALERT: BBTS;Memory Usage;OK;HARD;3;Memory usage: total:28669.54 MB - used: 3216.64 MB (11%) - free: 25452.90 MB (89%)
[1469834212] SERVICE ALERT: rackhampk;Black Cartridge HP CC364X;OK;HARD;3;Black Cartridge HP CC364X is at 4% - OK!
[1469834212] SERVICE ALERT: hallefirstcolor;Toner Collection Unit HP CE265A;OK;HARD;3;
[1469834212] SERVICE ALERT: owen101labcolor;Image Fuser Kit HP 110V-CE246A;OK;HARD;3;Image Fuser Kit HP 110V-CE246A, 220V-CE247A is at 36% - OK!
[1469834213] SERVICE ALERT: BBTS;Uptime;OK;HARD;3;System Uptime - 1 day(s) 14 hour(s) 59 minute(s)
[1469834213] SERVICE ALERT: ADRODC-DET;Service: DFS Namespace;OK;HARD;3;Dfs: Started
[1469834213] SERVICE ALERT: sillpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 47% - OK!
[1469834214] SERVICE ALERT: eateriespk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1469834216] HOST ALERT: rackhampk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1469834222] SERVICE ALERT: owen101labcolor;Black Cartridge HP CE260X;OK;HARD;3;Black Cartridge HP CE260X is at 31% - OK!
[1469834231] SERVICE ALERT: BBTS;VMWare Tools;OK;HARD;3;VMTools: Started
[1469834237] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 53% - OK!
[1469834237] SERVICE ALERT: hallesouthcommons;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1469834240] SERVICE ALERT: BBTS;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1469834242] SERVICE ALERT: hallegroundcolor;Yellow Cartridge HP CE262A;OK;HARD;3;Yellow Cartridge HP CE262A is at 97% - OK!
[1469834242] SERVICE ALERT: hallegroundcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1469834242] SERVICE ALERT: marshallpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1469834243] HOST ALERT: marshallpk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.52 ms
[1469834243] SERVICE ALERT: alexanderpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1469834254] SERVICE ALERT: BBTS;C:\ Drive Space;OK;HARD;3;c:\ - total: 246.90 Gb - used: 59.86 Gb (24%) - free 187.03 Gb (76%)
[1469834267] SERVICE ALERT: BBTS;Service: Windows Time;OK;HARD;3;W32Time: Started
[1469834271] SERVICE ALERT: BBTS_DB;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1469834271] SERVICE ALERT: strongpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 35% - OK!
[1469834273] SERVICE ALERT: owen101labcolor;Device Status;OK;HARD;3;Sleep mode on
[1469834274] SERVICE ALERT: studentcenterpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 82% - OK!
[1469834275] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1469834282] SERVICE ALERT: sillpk;Device Status;OK;HARD;3;Sleep Mode on
[1469834285] SERVICE ALERT: BBTS_DB;Service: WMI;OK;HARD;3;Winmgmt: Started
[1469834285] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1469834286] SERVICE ALERT: BBTS_DB;Memory Usage;OK;HARD;3;Memory usage: total:49149.71 MB - used: 7394.91 MB (15%) - free: 41754.80 MB (85%)
[1469834299] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 21% - OK!
[1469834304] SERVICE ALERT: prayharroldpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1469834304] SERVICE ALERT: ADRODC-LIV;Memory Usage;OK;HARD;3;Memory usage: total:37586.86 MB - used: 3477.42 MB (9%) - free: 34109.45 MB (91%)
[1469834306] SERVICE ALERT: ADRODC-LIV;Uptime;OK;HARD;3;System Uptime - 6 day(s) 23 hour(s) 57 minute(s)
[1469834307] SERVICE ALERT: BBTS;CPU Load;OK;HARD;3;CPU Load 2% (15 min average)
[1469834312] SERVICE ALERT: hallegroundcolor;Cyan Cartridge HP CE261A;OK;HARD;3;Cyan Cartridge HP CE261A is at 97% - OK!
[1469834313] SERVICE ALERT: ADRODC-DET;Service: PXE Boot;OK;HARD;3;WDSServer: Started
[1469834320] SERVICE ALERT: ADRODC-LIV;Service: MDT Monitor;OK;HARD;3;MDT_Monitor: Started
[1469834328] SERVICE ALERT: eateriespk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 34% - OK!
[1469834332] SERVICE ALERT: hallesecondsouth;Device Status;OK;HARD;3;Sleep Mode on
[1469834332] SERVICE ALERT: prayharroldpk;Device Status;OK;HARD;3;Ready
[1469834341] SERVICE ALERT: BBTS_DB;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1469834342] SERVICE ALERT: owen101labcolor;Magenta Cartridge HP CE263A;OK;HARD;3;Magenta Cartridge HP CE263A is at 55% - OK!
[1469834343] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1089 found since last run.
[1469834343] SERVICE ALERT: ADRODC-LIV;Service: Key Distribution Center;OK;HARD;3;KDC: Started
[1469834352] SERVICE ALERT: hallesecondsouth;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1469834353] SERVICE ALERT: ADRODC-LIV;CPU Load;OK;HARD;3;CPU Load 0% (15 min average)
[1469834354] SERVICE ALERT: ADRODC-LIV;Service: Distributed File System Replication;OK;HARD;3;DFSR: Started
[1469834356] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 64% - OK!
[1469834357] SERVICE ALERT: ADRODC-LIV;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1469834359] SERVICE ALERT: owen101labcolor;Toner Collection Unit HP CE265A;OK;HARD;3;
[1469834359] SERVICE ALERT: hallefirst3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1469834363] SERVICE ALERT: ADRODC-DET;Service: RPC;OK;HARD;3;RpcSs: Started
[1469834363] SERVICE ALERT: owen101labcolor;Image Transfer Kit HP CE249A;OK;HARD;3;Image Transfer Kit HP CE249A is at 37% - OK!
[1469834363] SERVICE ALERT: crossroadspk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 2 Trays Still Full
[1469834364] SERVICE ALERT: halleground3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1469834370] SERVICE ALERT: BBTS_DB;CPU Load;OK;HARD;3;CPU Load 3% (15 min average)
[1469834372] SERVICE ALERT: hallegroundcolor;Toner Collection Unit HP CE265A;OK;HARD;3;
[1469834380] SERVICE ALERT: BBTS;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1469834382] SERVICE ALERT: sillpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1469834390] SERVICE ALERT: ADRODC-LIV;C:\ Drive Space;OK;HARD;3;c:\ - total: 1115.16 Gb - used: 558.35 Gb (50%) - free 556.81 Gb (50%)
[1469834392] SERVICE ALERT: BBTS;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1469834392] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 54% - OK!
[1469834393] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Sleep Mode on
[1469834402] SERVICE ALERT: strongpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1469834403] SERVICE ALERT: hallefirst3;Device Status;OK;HARD;3;TRAY 2 EMPTY PLAIN LETTER   Ready
[1469834403] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 89% - OK!
[1469834405] SERVICE ALERT: mckennypk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1469834405] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1469834408] SERVICE ALERT: BBTS_DB;Size: IFDATA03.DBF;OK;HARD;3;OK: Size is  1956.0078125  Mb
[1469834413] SERVICE ALERT: ADRODC-DET;Service: Distributed File System Replication;OK;HARD;3;DFSR: Started
[1469834413] SERVICE ALERT: hallefirst1;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 40% - OK!
[1469834422] SERVICE ALERT: eateriespk;Device Status;OK;HARD;3;TRAY 3 OPEN  46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1469834423] SERVICE ALERT: ADRODC-DET;Memory Usage;OK;HARD;3;Memory usage: total:37586.86 MB - used: 3426.03 MB (9%) - free: 34160.84 MB (91%)
[1469834427] SERVICE ALERT: ADRODC-LIV;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1469834427] SERVICE ALERT: ADRODC-DET;CPU Load;OK;HARD;3;CPU Load 0% (15 min average)
[1469834429] SERVICE ALERT: ADRODC-LIV;Service: RPC;OK;HARD;3;RpcSs: Started
[1469834432] SERVICE ALERT: BBTS_DB;Size: TS_INDEX02.DBF;OK;HARD;3;OK: Size is  1380.0078125  Mb
[1469834433] SERVICE ALERT: ADRODC-DET;Service: MDT Monitor;OK;HARD;3;MDT_Monitor: Started
[1469834433] SERVICE ALERT: hallefirst2;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1469834435] SERVICE ALERT: ADRODC-DET;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1469834440] SERVICE ALERT: ADRODC-WEMU;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1469834443] SERVICE ALERT: ADRODC-LIV;Service: Windows Time;OK;HARD;3;W32Time: Started
[1469834443] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 58% - OK!
[1469834447] SERVICE ALERT: ADRODC-DET;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1469834451] SERVICE ALERT: ADRODC-LIV;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1469834452] SERVICE ALERT: strongpk;Device Status;OK;HARD;3;Sleep Mode on
[1469834454] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;OK;HARD;3;Cyan Cartridge HP CE261A is at 75% - OK!
[1469834456] SERVICE ALERT: ADRODC-DET;Service: WMI;OK;HARD;3;Winmgmt: Started
[1469834462] SERVICE ALERT: halleground3;Device Status;OK;HARD;3;Sleep Mode on
[1469834463] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469834463] SERVICE ALERT: ADRODC-DET;C:\ Drive Space;OK;HARD;3;c:\ - total: 1115.16 Gb - used: 279.46 Gb (25%) - free 835.70 Gb (75%)
[1469834463] SERVICE ALERT: ADRODC-LIV;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1469834463] SERVICE ALERT: ADRODC-LIV;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1469834463] SERVICE ALERT: ADRODC-DET;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1469834463] SERVICE ALERT: ADRODC-DET;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1469834463] SERVICE ALERT: ADRODC-DET;Service: Remote Management;OK;HARD;3;WinRM: Started
[1469834463] SERVICE ALERT: ADRODC-DET;Service: Windows Time;OK;HARD;3;W32Time: Started
[1469834463] SERVICE ALERT: ADRODC-DET;Service: AD Domain Service;OK;HARD;3;NTDS: Started
[1469834463] SERVICE ALERT: ADRODC-DET;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1469834463] SERVICE ALERT: ADRODC-DET;Service: AD Web Service;OK;HARD;3;ADWS: Started
[1469834468] SERVICE ALERT: ADRODC-DET;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1469834468] SERVICE ALERT: ADRODC-LIV;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1469834472] SERVICE ALERT: rooseveltpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1469834475] SERVICE ALERT: ADRODC-LIV;Service: PXE Boot;OK;HARD;3;WDSServer: Started
[1469834478] SERVICE ALERT: ADRODC-LIV;Service: DFS Namespace;OK;HARD;3;Dfs: Started
[1469834481] SERVICE ALERT: BBTS_DB;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1469834481] SERVICE ALERT: ADRODC-LIV;Service: Intersite Messaging;OK;HARD;3;IsmServ: Started
[1469834483] SERVICE ALERT: halleground1;Device Status;OK;HARD;3;Sleep Mode on
[1469834493] SERVICE ALERT: owen101labcolor;Yellow Cartridge HP CE262A;OK;HARD;3;Yellow Cartridge HP CE262A is at 52% - OK!
[1469834493] SERVICE ALERT: rooseveltpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 86% - OK!
[1469834499] SERVICE ALERT: ADRODC-DET;AD Replication;OK;HARD;3;OK: replication is up and running
[1469834508] SERVICE ALERT: ADRODC-LIV;Service: AD Domain Service;OK;HARD;3;NTDS: Started
[1469834510] SERVICE ALERT: marshallpk;Device Status;OK;HARD;3;Ready
[1469834511] SERVICE ALERT: BBTS_DB;VMWare Tools;OK;HARD;3;VMTools: Started
[1469834522] SERVICE ALERT: hallefirst1;Device Status;OK;HARD;3;Sleep Mode on
[1469834523] SERVICE ALERT: BBTS_DB;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1469834524] SERVICE ALERT: hallegroundcolor;Black Cartridge HP CE260X;OK;HARD;3;Black Cartridge HP CE260X is at 99% - OK!
[1469834533] SERVICE ALERT: marshallpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 56% - OK!
[1469834535] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1469834547] SERVICE ALERT: ADRODC-LIV;Service: WMI;OK;HARD;3;Winmgmt: Started
[1469834553] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Sleep mode on
[1469834572] SERVICE ALERT: rackhampk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1469834573] SERVICE ALERT: honorspk;Device Status;OK;HARD;3;Sleep Mode on
[1469834573] SERVICE ALERT: hallegroundcolor;Image Fuser Kit HP 110V-CE246A;OK;HARD;3;Image Fuser Kit HP 110V-CE246A, 220V-CE247A is at 60% - OK!
[1469834590] SERVICE ALERT: hallegroundcolor;Image Transfer Kit HP CE249A;OK;HARD;3;Image Transfer Kit HP CE249A is at 60% - OK!
[1469834593] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1014 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 1380 found since last run.
[1469834593] SERVICE ALERT: alexanderpk;Device Status;OK;HARD;3;Sleep Mode on
[1469834596] SERVICE ALERT: porterpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1469834597] SERVICE ALERT: BBTS;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1469834600] SERVICE ALERT: BBTS;Service: RPC;OK;HARD;3;RpcSs: Started
[1469834602] SERVICE ALERT: studentcenterpk;Device Status;OK;HARD;3;Sleep Mode on
[1469834602] SERVICE ALERT: halleground2;Device Status;OK;HARD;3;Sleep Mode on
[1469834603] SERVICE ALERT: ADRODC-LIV;AD Replication;OK;HARD;3;OK: replication is up and running
[1469834613] SERVICE ALERT: hallegroundcolor;Magenta Cartridge HP CE263A;OK;HARD;3;Magenta Cartridge HP CE263A is at 98% - OK!
[1469834713] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469835005] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 70% (15 min average)
[1469835005] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1469835063] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 1022 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1469835183] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469835313] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 532 found since last run.
[1469835433] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469835605] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1469835725] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1469835783] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 512 found since last run.
[1469835845] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1469835903] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469836033] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 604 found since last run.
[1469836153] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469836445] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 49% (15 min average)
[1469836753] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 555 found since last run.
[1469836873] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469837103] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 507 found since last run.
[1469837223] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469837645] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1469837765] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 89% (15 min average)
[1469837823] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 523 found since last run.
[1469837943] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469838073] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 587 found since last run.
[1469838193] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469838365] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1469838485] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1469838543] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 569 found since last run.
[1469838605] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 72% (15 min average)
[1469838663] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469838793] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 547 found since last run.
[1469838913] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469839205] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1469839325] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1469839445] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 95% (15 min average)
[1469839863] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 501 found since last run.
[1469839983] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469840045] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (15 min average)
[1469840113] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 503 found since last run.
[1469840233] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469840583] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 536 found since last run.
[1469840703] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469840833] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 532 found since last run.
[1469840953] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469841245] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1469841365] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1469841485] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 86% (15 min average)
[1469841553] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 592 found since last run.
[1469841673] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469841903] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 653 found since last run.
[1469842023] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469842085] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1469842205] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 93% (15 min average)
[1469842273] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 652 found since last run.
[1469842393] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469842623] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 547 found since last run.
[1469842743] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469842993] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 592 found since last run.
[1469843113] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469843343] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 534 found since last run.
[1469843405] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1469843463] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469843525] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1469843645] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1469843713] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 523 found since last run.
[1469843833] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469844433] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 548 found since last run.
[1469844553] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469844845] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1469845153] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 524 found since last run.
[1469845263] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 515 found since last run.
[1469845273] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469845383] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469845873] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 678 found since last run.
[1469845983] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 590 found since last run.
[1469845993] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469846103] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469846593] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 580 found since last run.
[1469846703] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 589 found since last run.
[1469846713] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469846823] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469847313] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 541 found since last run.
[1469847433] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469847845] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1469847965] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1469848023] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 506 found since last run.
[1469848033] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 569 found since last run.
[1469848085] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 94% (15 min average)
[1469848143] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469848153] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469848685] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1469848743] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 527 found since last run.
[1469848753] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 573 found since last run.
[1469848805] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1469848863] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469848873] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469848925] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1469849463] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 540 found since last run.
[1469849473] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 629 found since last run.
[1469849583] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469849593] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469850125] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 58% (15 min average)
[1469850183] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 502 found since last run.
[1469850193] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 554 found since last run.
[1469850303] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469850313] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1469850725] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1469850845] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1469850903] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 515 found since last run.
[1469850965] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 87% (15 min average)
[1469851023] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.