[1475467243] HOST ALERT: hallefirstcolor;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1475467269] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475467299] HOST ALERT: hallefirstcolor;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1475467329] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475467359] HOST ALERT: hallefirstcolor;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1475467383] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 502 found since last run.
[1475467389] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475467398] SERVICE ALERT: hallefirstcolor;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475467399] SERVICE ALERT: hallefirstcolor;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475467419] HOST ALERT: hallefirstcolor;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1475467419] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475467419] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475467449] HOST ALERT: hallefirstcolor;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1475467457] SERVICE ALERT: hallefirstcolor;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475467457] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475467459] SERVICE ALERT: hallefirstcolor;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475467463] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 3085 found since last run. oneLevelSearchOps OK. searchOps is too high! 3109 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 514 found since last run.
[1475467466] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475467479] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475467480] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475467487] HOST ALERT: hallefirstcolor;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1475467503] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475467509] SERVICE ALERT: hallefirstcolor;Toner Collection Unit HP CE265A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475467517] SERVICE ALERT: hallefirstcolor;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475467517] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475467519] SERVICE ALERT: hallefirstcolor;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475467526] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475467539] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475467539] HOST ALERT: hallefirstcolor;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1475467539] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475467569] SERVICE ALERT: hallefirstcolor;Toner Collection Unit HP CE265A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475467569] HOST ALERT: hallefirstcolor;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1475467577] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475467579] SERVICE ALERT: hallefirstcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475467583] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475467586] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475467607] HOST ALERT: hallefirstcolor;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1475467629] SERVICE ALERT: hallefirstcolor;Toner Collection Unit HP CE265A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475467639] SERVICE ALERT: hallefirstcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475467659] HOST ALERT: hallefirstcolor;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1475467699] SERVICE ALERT: hallefirstcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475468103] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 584 found since last run.
[1475468183] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 566 found since last run.
[1475468223] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475468303] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475468823] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 2503 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 532 found since last run.
[1475468943] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475469061] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 54% (15 min average)
[1475469503] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 541 found since last run.
[1475469543] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 503 found since last run.
[1475469623] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475469661] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1475469663] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475469781] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 89% (15 min average)
[1475470223] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 587 found since last run.
[1475470263] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 555 found since last run.
[1475470343] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475470381] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1475470383] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475470501] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1475470621] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1475470943] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 548 found since last run.
[1475470983] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 547 found since last run.
[1475471063] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475471103] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475471663] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 501 found since last run.
[1475471703] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 507 found since last run.
[1475471703] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;SOFT;1;Warning:  idm1  are more than 120 seconds behind!
[1475471783] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475471823] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475471824] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;SOFT;2;OK. All replicas within 120 seconds.
[1475472383] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 2790 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 526 found since last run.
[1475472423] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;1;Critcal!  idm1  are more than 240 seconds behind!
[1475472503] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475472544] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;SOFT;2;Critcal!  idm1  are more than 240 seconds behind!
[1475472600] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service has entered a period of scheduled downtime
[1475472666] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1475473103] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 539 found since last run.
[1475473223] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475473263] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1475473621] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 91% (15 min average)
[1475473823] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 537 found since last run.
[1475473943] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475474463] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1475474821] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1475474823] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 512 found since last run.
[1475474941] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1475474943] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475475061] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1475475063] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1475475063] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1475475143] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 513 found since last run.
[1475475263] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475475543] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 512 found since last run.
[1475475663] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475475863] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 528 found since last run.
[1475475983] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475476863] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 134 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 659 found since last run.
[1475476983] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475477183] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 531 found since last run.
[1475477303] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475477903] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 598 found since last run.
[1475478023] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475478623] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 113 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 650 found since last run.
[1475478743] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475479343] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 112 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 661 found since last run.
[1475479463] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475479861] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 86% (15 min average)
[1475479983] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 539 found since last run.
[1475480063] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 622 found since last run.
[1475480103] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475480183] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475480463] SERVICE ALERT: idm2.emich.edu;People.campus Replication;WARNING;HARD;3;Warning:  idm1  are more than 120 seconds behind!
[1475480626] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1475480745] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1475480783] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 623 found since last run.
[1475480865] SERVICE ALERT: epprogsql2;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1475480903] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475481063] SERVICE ALERT: idm2.emich.edu;People.campus Replication;CRITICAL;HARD;3;Critcal!  idm1  are more than 240 seconds behind!
[1475481303] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 105 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 587 found since last run.
[1475481423] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475481465] SERVICE ALERT: epprogsql2;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1475481661] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1475481663] SERVICE ALERT: idm2.emich.edu;People.campus Replication;OK;HARD;3;OK. All replicas within 120 seconds.
[1475481781] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1475481901] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1475482501] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 41% (15 min average)
[1475483400] SERVICE DOWNTIME ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service has exited from a period of scheduled downtime
[1475483701] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1475483821] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 49% (15 min average)
[1475483865] SERVICE ALERT: epprogsql2;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1475483903] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 582 found since last run.
[1475484023] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475484421] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1475484465] SERVICE ALERT: epprogsql2;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1475484465] SERVICE FLAPPING ALERT: epprogsql2;CPU Load;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1475484541] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1475484661] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1475485623] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 518 found since last run.
[1475485743] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475485861] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1475486461] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1475486581] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 56% (15 min average)
[1475487623] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 504 found since last run.
[1475487743] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475488943] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 527 found since last run.
[1475489063] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475489581] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1475489663] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 614 found since last run.
[1475489701] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1475489783] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475489821] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 83% (15 min average)
[1475490383] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 511 found since last run.
[1475490421] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1475490503] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475490541] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1475490543] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 508 found since last run.
[1475490661] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1475490663] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475490784] SERVICE ALERT: eateriespk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1475490844] SERVICE ALERT: eateriespk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1475490904] SERVICE ALERT: eateriespk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1475491103] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 549 found since last run.
[1475491223] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475491262] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1475491263] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 539 found since last run.
[1475491383] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475491823] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 562 found since last run.
[1475491943] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475492460] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1475492543] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 872 found since last run.
[1475492583] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 819 found since last run.
[1475492663] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475492703] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475493060] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1475493063] SERVICE FLAPPING ALERT: idm2.emich.edu;People.campus Replication;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1475493263] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1460 found since last run.
[1475493303] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1494 found since last run.
[1475493383] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475493423] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475493983] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1685 found since last run.
[1475494023] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1703 found since last run.
[1475494103] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475494143] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475494260] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 90% (15 min average)
[1475494703] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1807 found since last run.
[1475494743] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1763 found since last run.
[1475494823] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475494860] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1475494863] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475494980] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 68% (15 min average)
[1475495383] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1475495423] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1757 found since last run.
[1475495443] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1475495463] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1838 found since last run.
[1475495503] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1475495543] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475495580] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1475495583] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475495700] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1475495820] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1475495865] SERVICE FLAPPING ALERT: epprogsql2;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1475496103] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1475496143] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 1409 found since last run.
[1475496183] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 3234 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 1450 found since last run.
[1475496263] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475496303] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475496703] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1475496763] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1475496824] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1475496863] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 625 found since last run.
[1475496903] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5695 found since last run. oneLevelSearchOps OK. searchOps is too high! 5856 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 589 found since last run.
[1475496983] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475497020] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1475497023] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475497423] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Clearing paper path
[1475497423] SERVICE FLAPPING ALERT: rackhampk;Device Status;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1475497583] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 658 found since last run.
[1475497620] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1475497623] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5856 found since last run. oneLevelSearchOps OK. searchOps is too high! 5969 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 627 found since last run.
[1475497703] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475497740] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1475497743] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475497860] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 81% (15 min average)
[1475498023] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1475498083] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1475498111] SERVICE ALERT: hallefirstcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1475498113] SERVICE ALERT: hallefirstcolor;Device Status;OK;HARD;3;Tray 2 ANY TYPE LETTER To change size or type press OK Tray 2 ANY TYPE LETTER To accept settings 70 72 65 73 73 20 12
[1475498115] HOST ALERT: hallefirstcolor;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.36 ms
[1475498131] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;OK;HARD;3;Black Cartridge HP CE260X is at 89% - OK!
[1475498134] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;OK;HARD;3;Yellow Cartridge HP CE262A is at 75% - OK!
[1475498143] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1475498167] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;OK;HARD;3;Image Transfer Kit HP CE249A is at 80% - OK!
[1475498180] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;OK;HARD;3;Cyan Cartridge HP CE261A is at 40% - OK!
[1475498223] SERVICE ALERT: hallefirstcolor;Toner Collection Unit HP CE265A;OK;HARD;3;
[1475498293] SERVICE ALERT: hallefirstcolor;Image Fuser Kit HP 110V-CE246A;OK;HARD;3;Image Fuser Kit HP 110V-CE246A, 220V-CE247A is at 100% - OK!
[1475498303] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 3476 found since last run. oneLevelSearchOps OK. searchOps is too high! 3644 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 625 found since last run.
[1475498343] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 4934 found since last run. oneLevelSearchOps OK. searchOps is too high! 5103 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 614 found since last run.
[1475498423] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475498460] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1475498463] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475498580] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 79% (15 min average)
[1475498583] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;OK;HARD;3;Magenta Cartridge HP CE263A is at 70% - OK!
[1475498743] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1475499023] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 4325 found since last run. oneLevelSearchOps OK. searchOps is too high! 4415 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 629 found since last run.
[1475499063] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5321 found since last run. oneLevelSearchOps OK. searchOps is too high! 5440 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 653 found since last run.
[1475499143] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475499183] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475499343] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1475499403] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1475499464] SERVICE ALERT: rackhampk;Device Status;OK;SOFT;3;Open tray indicated.
[1475499743] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5466 found since last run. oneLevelSearchOps OK. searchOps is too high! 5578 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 1514 found since last run.
[1475499783] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6190 found since last run. oneLevelSearchOps OK. searchOps is too high! 6333 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 1358 found since last run.
[1475499863] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 602 found since last run.
[1475499903] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475499983] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475500063] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1475500123] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1475500183] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1475500503] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 4900 found since last run. oneLevelSearchOps OK. searchOps is too high! 5030 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 677 found since last run.
[1475500583] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 4478 found since last run. oneLevelSearchOps OK. searchOps is too high! 4641 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 802 found since last run.
[1475500623] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475500703] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475500783] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1475500980] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1475501100] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 88% (15 min average)
[1475501223] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 108 found since last run. wholeSubtreeSearchOps is too high! 8948 found since last run. oneLevelSearchOps OK. searchOps is too high! 9123 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 994 found since last run.
[1475501303] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 118 found since last run. wholeSubtreeSearchOps is too high! 7080 found since last run. oneLevelSearchOps OK. searchOps is too high! 7377 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 885 found since last run.
[1475501343] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475501423] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475501700] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1475501820] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1475501940] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1475501943] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6155 found since last run. oneLevelSearchOps OK. searchOps is too high! 6341 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 871 found since last run.
[1475501983] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1475502023] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 10290 found since last run. oneLevelSearchOps OK. searchOps is too high! 10547 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 779 found since last run.
[1475502043] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1475502063] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475502104] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1475502143] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475502663] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 107 found since last run. wholeSubtreeSearchOps is too high! 4296 found since last run. oneLevelSearchOps OK. searchOps is too high! 4455 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 858 found since last run.
[1475502703] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1475502743] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 105 found since last run. wholeSubtreeSearchOps is too high! 7515 found since last run. oneLevelSearchOps OK. searchOps is too high! 7674 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 760 found since last run.
[1475502783] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475502863] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475503140] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1475503303] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1475503363] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1475503383] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 108 found since last run. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 3112 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 812 found since last run.
[1475503423] SERVICE ALERT: rackhampk;Device Status;OK;SOFT;3;Open tray indicated.
[1475503463] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8376 found since last run. oneLevelSearchOps OK. searchOps is too high! 8543 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 809 found since last run.
[1475503503] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475503583] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475504023] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1475504083] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1475504103] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 3036 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 805 found since last run.
[1475504144] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1475504183] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6150 found since last run. oneLevelSearchOps OK. searchOps is too high! 6395 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 764 found since last run.
[1475504223] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475504303] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475504340] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1475504460] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1475504580] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1475504743] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1475504823] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7764 found since last run. oneLevelSearchOps OK. searchOps is too high! 7913 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 2331 found since last run.
[1475504903] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors is too high! 116 found since last run. wholeSubtreeSearchOps is too high! 6876 found since last run. oneLevelSearchOps OK. searchOps is too high! 7014 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 2920 found since last run.
[1475504943] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475505023] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475505543] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6362 found since last run. oneLevelSearchOps OK. searchOps is too high! 6536 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 998 found since last run.
[1475505623] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8201 found since last run. oneLevelSearchOps OK. searchOps is too high! 8249 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 652 found since last run.
[1475505663] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475505743] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475505780] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 90% (15 min average)
[1475505787] SERVICE ALERT: HDBACKUPS;F:\ Drive Space;WARNING;SOFT;1;f:\ - total: 3315.22 Gb - used: 2653.54 Gb (80%) - free 661.69 Gb (20%)
[1475505907] SERVICE ALERT: HDBACKUPS;F:\ Drive Space;WARNING;SOFT;2;f:\ - total: 3315.22 Gb - used: 2654.42 Gb (80%) - free 660.81 Gb (20%)
[1475505942] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1475506002] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1475506027] SERVICE ALERT: HDBACKUPS;F:\ Drive Space;WARNING;HARD;3;f:\ - total: 3315.22 Gb - used: 2655.30 Gb (80%) - free 659.93 Gb (20%)
[1475506062] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1475506263] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 13183 found since last run. oneLevelSearchOps OK. searchOps is too high! 13354 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 934 found since last run.
[1475506343] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8338 found since last run. oneLevelSearchOps OK. searchOps is too high! 8367 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 855 found since last run.
[1475506380] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1475506383] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 2723 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1475506463] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475506500] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1475506503] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475506620] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1475506627] SERVICE ALERT: HDBACKUPS;F:\ Drive Space;OK;HARD;3;f:\ - total: 3315.22 Gb - used: 2594.53 Gb (78%) - free 720.70 Gb (22%)
[1475506662] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1475506703] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1475506763] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1475506823] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1475507064] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7869 found since last run. oneLevelSearchOps OK. searchOps is too high! 7907 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 846 found since last run.
[1475507103] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9257 found since last run. oneLevelSearchOps OK. searchOps is too high! 9448 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 843 found since last run.
[1475507183] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475507223] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475507704] SERVICE ALERT: eateriespk;Device Status;OK;HARD;3;Clearing paper path
[1475507783] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6939 found since last run. oneLevelSearchOps OK. searchOps is too high! 6952 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 773 found since last run.
[1475507823] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6127 found since last run. oneLevelSearchOps OK. searchOps is too high! 6289 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 660 found since last run.
[1475507903] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475507943] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475508304] SERVICE ALERT: eateriespk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1475508364] SERVICE ALERT: eateriespk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1475508425] SERVICE ALERT: eateriespk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1475508503] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7602 found since last run. oneLevelSearchOps OK. searchOps is too high! 7611 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 765 found since last run.
[1475508543] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 4392 found since last run. oneLevelSearchOps OK. searchOps is too high! 4649 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 852 found since last run.
[1475508623] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475508624] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1475508664] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475508686] SERVICE ALERT: prayharroldpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1475508746] SERVICE ALERT: prayharroldpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1475508806] SERVICE ALERT: prayharroldpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1475509020] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 81% (15 min average)
[1475509223] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5361 found since last run. oneLevelSearchOps OK. searchOps is too high! 5368 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 767 found since last run.
[1475509263] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6217 found since last run. oneLevelSearchOps OK. searchOps is too high! 6305 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 670 found since last run.
[1475509343] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475509383] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475509406] SERVICE ALERT: prayharroldpk;Device Status;OK;HARD;3;Genuine HP cartridge installed
[1475509620] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1475509740] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1475509943] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7234 found since last run. oneLevelSearchOps OK. searchOps is too high! 7241 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 678 found since last run.
[1475509983] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6925 found since last run. oneLevelSearchOps OK. searchOps is too high! 7321 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 668 found since last run.
[1475510062] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1475510063] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475510103] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475510123] SERVICE ALERT: crossroadspk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1475510182] SERVICE ALERT: crossroadspk;Device Status;OK;SOFT;3;Open tray indicated.
[1475510610] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1475510663] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7128 found since last run. oneLevelSearchOps OK. searchOps is too high! 7135 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 786 found since last run.
[1475510670] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 25% - OK!
[1475510703] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8588 found since last run. oneLevelSearchOps OK. searchOps is too high! 9084 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 757 found since last run.
[1475510783] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475510823] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475510823] SERVICE ALERT: eateriespk;Device Status;OK;HARD;3;MENUS  RETRIEVE JOB  INFORMATION  PAPER HANDLING
[1475510940] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1475511060] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1475511180] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1475511383] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 6593 found since last run. oneLevelSearchOps OK. searchOps is too high! 6607 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 708 found since last run.
[1475511423] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 13666 found since last run. oneLevelSearchOps OK. searchOps is too high! 13945 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 682 found since last run.
[1475511503] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475511543] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 3425 found since last run. oneLevelSearchOps OK. searchOps is too high! 3450 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1475511663] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 3425 found since last run. oneLevelSearchOps OK. searchOps is too high! 3457 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1475512103] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 9345 found since last run. oneLevelSearchOps OK. searchOps is too high! 9351 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 795 found since last run.
[1475512223] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475512317] SERVICE ALERT: AXIOM-2012;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1475512437] SERVICE ALERT: AXIOM-2012;CPU Load;OK;SOFT;2;CPU Load 61% (15 min average)
[1475512661] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1475512722] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1475512781] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1475512823] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7538 found since last run. oneLevelSearchOps OK. searchOps is too high! 7547 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 752 found since last run.
[1475512943] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475513543] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7025 found since last run. oneLevelSearchOps OK. searchOps is too high! 7031 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 739 found since last run.
[1475513580] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1475513663] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475513980] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1475514180] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1475514263] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 8133 found since last run. oneLevelSearchOps OK. searchOps is too high! 8142 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 743 found since last run.
[1475514383] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475514580] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1475514640] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1475514700] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1475514983] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 5695 found since last run. oneLevelSearchOps OK. searchOps is too high! 5700 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 717 found since last run.
[1475515103] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475515300] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1475515438] HOST ALERT: hallefirst3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1475515509] SERVICE ALERT: hallefirst3;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1475515527] HOST ALERT: hallefirst3;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1475515569] SERVICE ALERT: hallefirst3;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1475515599] HOST ALERT: hallefirst3;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1475515609] SERVICE ALERT: hallefirst3;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1475515629] SERVICE ALERT: hallefirst3;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1475515639] HOST ALERT: hallefirst3;DOWN;SOFT;4;(Host check timed out after 30.00 seconds)
[1475515669] SERVICE ALERT: hallefirst3;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1475515699] HOST ALERT: hallefirst3;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1475515703] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 7416 found since last run. oneLevelSearchOps OK. searchOps is too high! 7423 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 697 found since last run.
[1475515713] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1475515729] SERVICE ALERT: hallefirst3;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1475515743] HOST ALERT: hallefirst3;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1475515773] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1475515803] HOST ALERT: hallefirst3;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1475515823] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475515833] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.180. Is the printer off?  Is the network down?
[1475515863] HOST ALERT: hallefirst3;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1475515901] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1475515953] HOST ALERT: hallefirst3;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1475515960] SERVICE ALERT: rackhampk;Device Status;OK;SOFT;2;Open tray indicated.
[1475516043] HOST ALERT: hallefirst3;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1475516423] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 4739 found since last run. oneLevelSearchOps OK. searchOps is too high! 4745 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 712 found since last run.
[1475516543] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475516559] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1475516619] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1475516679] SERVICE ALERT: rackhampk;Device Status;OK;SOFT;3;Open tray indicated.
[1475517143] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 12523 found since last run. oneLevelSearchOps OK. searchOps is too high! 12532 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 740 found since last run.
[1475517263] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 2631 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1475517279] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1475517339] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1475517367] HOST ALERT: hallefirst3;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.50 ms
[1475517383] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps is too high! 2723 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1475517399] SERVICE ALERT: rackhampk;Device Status;OK;SOFT;3;Open tray indicated.
[1475517423] SERVICE ALERT: hallefirst3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 2 of 3 Trays Still Full
[1475517454] SERVICE ALERT: strongpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1475517515] SERVICE ALERT: strongpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1475517523] SERVICE ALERT: hallefirst3;Device Status;OK;HARD;3;TRAY 2 EMPTY PLAIN LETTER   Ready
[1475517574] SERVICE ALERT: strongpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1475517627] SERVICE ALERT: hallefirst3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 22% - OK!
[1475518000] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1475518059] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1475518119] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1475518153] SERVICE ALERT: strongpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1475518173] SERVICE ALERT: strongpk;Device Status;OK;HARD;3;50 72 65 73 73 20 17 20 74 6F 20 73 65 65 20 73  74 65 70  by step information
[1475518213] SERVICE ALERT: strongpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1475518273] SERVICE ALERT: strongpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1475518719] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1475519319] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1475519379] SERVICE ALERT: rackhampk;Device Status;OK;SOFT;2;Open tray indicated.
[1475519580] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 91% (15 min average)
[1475519979] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1475520039] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1475520100] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1475520180] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1475520300] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1475520353] SERVICE ALERT: hallegroundcolor;Magenta Cartridge HP CE263A;CRITICAL;SOFT;1;Magenta Cartridge HP CE263A is at CRITICAL level!
[1475520414] SERVICE ALERT: hallegroundcolor;Magenta Cartridge HP CE263A;CRITICAL;SOFT;2;Magenta Cartridge HP CE263A is at CRITICAL level!
[1475520420] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1475520473] SERVICE ALERT: hallegroundcolor;Magenta Cartridge HP CE263A;CRITICAL;HARD;3;Magenta Cartridge HP CE263A is at CRITICAL level!
[1475520699] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1475521299] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1475521359] SERVICE ALERT: rackhampk;Device Status;OK;SOFT;2;Open tray indicated.
[1475521620] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 91% (15 min average)
[1475521873] SERVICE ALERT: strongpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 10% - OK!
[1475521959] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1475522019] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1475522079] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1475522220] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1475522340] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1475522373] SERVICE ALERT: strongpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1475522433] SERVICE ALERT: strongpk;Device Status;OK;SOFT;2;Clearing paper path
[1475522460] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1475522513] SERVICE ALERT: porterpk;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1475522574] SERVICE ALERT: porterpk;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1475522633] SERVICE ALERT: porterpk;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at 2% - WARNING!
[1475522680] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1475523279] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1475523339] SERVICE ALERT: rackhampk;Device Status;OK;SOFT;2;Open tray indicated.
[1475523939] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1475523999] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1475524059] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1475524260] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1475524513] SERVICE ALERT: hallegroundcolor;Yellow Cartridge HP CE262A;OK;HARD;3;Yellow Cartridge HP CE262A is at 100% - OK!
[1475524659] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1475524663] SERVICE ALERT: hallegroundcolor;Cyan Cartridge HP CE261A;OK;HARD;3;Cyan Cartridge HP CE261A is at 100% - OK!
[1475524673] SERVICE ALERT: hallegroundcolor;Magenta Cartridge HP CE263A;OK;HARD;3;Magenta Cartridge HP CE263A is at 100% - OK!
[1475524860] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1475525033] SERVICE ALERT: porterpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1475525259] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1475525320] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1475525379] SERVICE ALERT: rackhampk;Device Status;OK;SOFT;3;Open tray indicated.
[1475525460] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1475525634] SERVICE ALERT: porterpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1475525979] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1475526039] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1475526099] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1475526699] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1475527260] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1475527299] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1475527359] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1475527420] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1475527860] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1475528460] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1475528619] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1475529060] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1475529219] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1475529279] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1475529339] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1475530260] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1475530539] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1475530680] SERVICE ALERT: hallegroundcolor;Magenta Cartridge HP CE263A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1475530733] SERVICE ALERT: hallegroundcolor;Magenta Cartridge HP CE263A;OK;SOFT;2;Magenta Cartridge HP CE263A is at 100% - OK!
[1475530860] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1475531140] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1475531199] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1475531259] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1475532280] SERVICE ALERT: owen101lab1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1475532340] SERVICE ALERT: owen101lab1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1475532400] SERVICE ALERT: owen101lab1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1475532459] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1475533000] SERVICE ALERT: owen101lab1;Device Status;OK;HARD;3;Checking printer
[1475533059] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1475533119] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1475533179] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1475533260] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1475533600] SERVICE ALERT: owen101lab1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1475533660] SERVICE ALERT: owen101lab1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1475533720] SERVICE ALERT: owen101lab1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1475533780] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1475533860] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 94% (15 min average)
[1475534463] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1475534524] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1475534584] SERVICE ALERT: marshallpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1475534979] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1475535039] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1475535060] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1475535099] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1475535180] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1475535300] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 96% (15 min average)
[1475535699] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1475536014] SERVICE ALERT: hallefirst1;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1475536073] SERVICE ALERT: hallefirst1;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! All Trays Still Full
[1475536299] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1475536359] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1475536420] SERVICE ALERT: rackhampk;Device Status;OK;SOFT;3;Open tray indicated.
[1475537019] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1475537079] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1475537100] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 91% (15 min average)
[1475537139] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1475537700] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1475537739] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1475537820] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1475537940] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1475538939] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1475539000] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1475539059] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1475539740] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 74% (15 min average)
[1475540258] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1475540340] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1475540460] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1475540580] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1475540858] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1475540918] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1475540978] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1475541578] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1475542510] SERVICE ALERT: hallefirstcolor;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1475542571] SERVICE ALERT: hallefirstcolor;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1475542630] SERVICE ALERT: hallefirstcolor;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1475542721] SERVICE ALERT: owen101lab1;Device Status;OK;HARD;3;Clearing paper path
[1475542779] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1475542838] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1475542898] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1475543230] SERVICE ALERT: hallefirstcolor;Device Status;OK;HARD;3;Open tray indicated.
[1475543320] SERVICE ALERT: owen101lab1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1475543379] SERVICE ALERT: owen101lab1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1475543439] SERVICE ALERT: owen101lab1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1475543439] SERVICE FLAPPING ALERT: owen101lab1;Device Status;STARTED; Service appears to have started flapping (24.4% change >= 20.0% threshold)
[1475543498] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1475543612] SERVICE ALERT: halleground1;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1475543672] SERVICE ALERT: halleground1;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1475543732] SERVICE ALERT: halleground1;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1475544698] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1475544758] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1475544777] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1475544818] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1475545377] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1475545419] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1475546577] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1475546618] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1475546678] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1475546738] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1475547177] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1475547338] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1475547777] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 84% (15 min average)
[1475547932] SERVICE ALERT: halleground1;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1475547938] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1475547998] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1475548059] SERVICE ALERT: rackhampk;Device Status;OK;SOFT;3;Open tray indicated.
[1475548377] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1475548497] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1475548617] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1475548658] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1475548718] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1475548778] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1475549378] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1475549817] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 87% (15 min average)
[1475550417] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1475550537] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1475550578] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1475550639] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1475550657] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1475550698] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1475551257] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1475551298] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1475552457] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1475552498] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1475552558] SERVICE ALERT: rackhampk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1475552618] SERVICE ALERT: rackhampk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1475552712] SERVICE ALERT: hallefirstcolor;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475552736] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475552739] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475552742] HOST ALERT: hallefirstcolor;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1475552769] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475552772] SERVICE ALERT: hallefirstcolor;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475552783] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1475552784] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475552796] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475552799] HOST ALERT: hallefirstcolor;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1475552800] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475552829] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475552829] SERVICE ALERT: hallefirstcolor;Toner Collection Unit HP CE265A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475552830] HOST ALERT: hallefirstcolor;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1475552832] SERVICE ALERT: hallefirstcolor;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475552836] SERVICE ALERT: hallefirstcolor;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475552844] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475552856] SERVICE ALERT: hallefirstcolor;Black Cartridge HP CE260X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475552859] SERVICE ALERT: hallefirstcolor;Yellow Cartridge HP CE262A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475552862] HOST ALERT: hallefirstcolor;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1475552889] SERVICE ALERT: hallefirstcolor;Image Transfer Kit HP CE249A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475552889] SERVICE ALERT: hallefirstcolor;Toner Collection Unit HP CE265A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475552897] SERVICE ALERT: hallefirstcolor;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475552898] SERVICE ALERT: hallefirstcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475552904] SERVICE ALERT: hallefirstcolor;Cyan Cartridge HP CE261A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475552919] HOST ALERT: hallefirstcolor;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1475552949] SERVICE ALERT: hallefirstcolor;Toner Collection Unit HP CE265A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475552956] SERVICE ALERT: hallefirstcolor;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475552958] SERVICE ALERT: hallefirstcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475552979] HOST ALERT: hallefirstcolor;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1475553018] SERVICE ALERT: hallefirstcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475553048] HOST ALERT: hallefirstcolor;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1475553138] HOST ALERT: hallefirstcolor;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1475553189] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475553218] SERVICE ALERT: rackhampk;Device Status;OK;HARD;3;Open tray indicated.
[1475553219] HOST ALERT: hallefirstcolor;DOWN;SOFT;9;(Host check timed out after 30.00 seconds)
[1475553250] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475553280] HOST ALERT: hallefirstcolor;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1475553309] SERVICE ALERT: hallefirstcolor;Magenta Cartridge HP CE263A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.104.150. Is the printer off?  Is the network down?
[1475553383] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 500 found since last run.
[1475553503] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.