[1459483568] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 46% (5 min average)
[1459484168] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1459484288] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1459484408] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1459490482] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1459490602] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 90% (5 min average)
[1459490722] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 97% (5 min average)
[1459491322] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 76% (5 min average)
[1459491322] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1459491556] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 84.90 Gb - used: 79.24 Gb (93%) - free 5.66 Gb (7%)
[1459491676] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 84.90 Gb - used: 79.24 Gb (93%) - free 5.66 Gb (7%)
[1459491796] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 84.90 Gb - used: 79.24 Gb (93%) - free 5.66 Gb (7%)
[1459492522] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1459492642] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 95% (5 min average)
[1459492762] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 96% (5 min average)
[1459493362] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 89% (5 min average)
[1459495162] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1459495282] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 92% (5 min average)
[1459495402] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 82% (5 min average)
[1459496002] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 91% (5 min average)
[1459496122] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 91% (5 min average)
[1459496242] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 78% (5 min average)
[1459497008] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 93% (5 min average)
[1459497608] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1459499408] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 58% (5 min average)
[1459500608] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1459500728] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 46% (5 min average)
[1459501328] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1459501448] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1459501568] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1459501568] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1459502044] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;SOFT;1;WARNING: Replication with ADDC5 failed!
[1459502164] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;SOFT;2;WARNING: Replication with ADDC5 failed!
[1459502168] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 68% (5 min average)
[1459502283] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;HARD;3;WARNING: Replication with ADDC5 failed!
[1459502768] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1459502888] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 83% (5 min average)
[1459503488] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 94% (5 min average)
[1459503608] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 85% (5 min average)
[1459505242] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1459505408] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459505528] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1459505648] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1459506248] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 82% (5 min average)
[1459506848] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1459506968] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 90% (5 min average)
[1459507088] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1459507688] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 93% (5 min average)
[1459508288] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 36% (5 min average)
[1459508888] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1459509008] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 79% (5 min average)
[1459509608] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1459509728] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1459509848] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1459511648] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 54% (5 min average)
[1459512248] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1459512368] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1459512488] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 96% (5 min average)
[1459512963] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1459513088] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 71% (5 min average)
[1459513089] SERVICE ALERT: ADRODC-LIV;AD Replication;OK;HARD;3;OK: replication is up and running
[1459513995] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;CRITICAL;HARD;3;c:\ - total: 84.90 Gb - used: 80.78 Gb (95%) - free 4.12 Gb (5%)
[1459514763] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1459514823] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1459514883] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1459516088] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1459516208] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 74% (5 min average)
[1459516808] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1459516928] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 36% (5 min average)
[1459517139] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459517199] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459517260] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459517528] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1459517648] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 38% (5 min average)
[1459517672] SERVICE ALERT: rooseveltpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459517732] SERVICE ALERT: rooseveltpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459517792] SERVICE ALERT: rooseveltpk;Device Status;OK;SOFT;3;Open tray indicated.
[1459518195] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 84.90 Gb - used: 80.40 Gb (95%) - free 4.50 Gb (5%)
[1459518992] SERVICE ALERT: rooseveltpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459519052] SERVICE ALERT: rooseveltpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459519113] SERVICE ALERT: rooseveltpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459519712] SERVICE ALERT: rooseveltpk;Device Status;OK;HARD;3;Open tray indicated.
[1459520312] SERVICE ALERT: rooseveltpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459520372] SERVICE ALERT: rooseveltpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459520432] SERVICE ALERT: rooseveltpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459520773] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1459520893] SERVICE ALERT: BOEAPPTEST;CPU Load;OK;SOFT;2;CPU Load 24% (5 min average)
[1459521632] SERVICE ALERT: rooseveltpk;Device Status;OK;HARD;3;Open tray indicated.
[1459521632] SERVICE FLAPPING ALERT: rooseveltpk;Device Status;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1459521795] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;OK;HARD;3;c:\ - total: 84.90 Gb - used: 75.96 Gb (89%) - free 8.94 Gb (11%)
[1459522833] SERVICE ALERT: rooseveltpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459522892] SERVICE ALERT: rooseveltpk;Device Status;OK;SOFT;2;Open tray indicated.
[1459523048] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 90% (5 min average)
[1459523168] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 96% (5 min average)
[1459523258] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Ready
[1459523288] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 83% (5 min average)
[1459523858] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459523918] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459523978] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459524092] SERVICE ALERT: rooseveltpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459524152] SERVICE ALERT: rooseveltpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459524212] SERVICE ALERT: rooseveltpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459524578] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Ready
[1459524578] SERVICE FLAPPING ALERT: hallefirst2;Device Status;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1459524812] SERVICE ALERT: rooseveltpk;Device Status;OK;HARD;3;Open tray indicated. TRAY 2
[1459525088] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1459525500] HOST DOWNTIME ALERT: AXLIC-DEVL;STARTED; Host has entered a period of scheduled downtime
[1459525500] HOST DOWNTIME ALERT: INTLTESTAPP;STARTED; Host has entered a period of scheduled downtime
[1459525500] HOST DOWNTIME ALERT: AX-QUAL;STARTED; Host has entered a period of scheduled downtime
[1459525500] HOST DOWNTIME ALERT: SANDBOX-APP;STARTED; Host has entered a period of scheduled downtime
[1459525500] HOST DOWNTIME ALERT: MAPS01TEST;STARTED; Host has entered a period of scheduled downtime
[1459525500] HOST DOWNTIME ALERT: INTLAPP-ISSMAP;STARTED; Host has entered a period of scheduled downtime
[1459525500] HOST DOWNTIME ALERT: ADDC4;STARTED; Host has entered a period of scheduled downtime
[1459525500] HOST DOWNTIME ALERT: ITRFTEST;STARTED; Host has entered a period of scheduled downtime
[1459525500] HOST DOWNTIME ALERT: BOEAPPTEST;STARTED; Host has entered a period of scheduled downtime
[1459525500] HOST DOWNTIME ALERT: R25-DELTA2012;STARTED; Host has entered a period of scheduled downtime
[1459525500] HOST DOWNTIME ALERT: ADDC2;STARTED; Host has entered a period of scheduled downtime
[1459525500] HOST DOWNTIME ALERT: ADDC3;STARTED; Host has entered a period of scheduled downtime
[1459525500] HOST DOWNTIME ALERT: AXIOM-TEST-2012;STARTED; Host has entered a period of scheduled downtime
[1459525500] HOST DOWNTIME ALERT: EVENTSTEST-2012;STARTED; Host has entered a period of scheduled downtime
[1459525500] HOST DOWNTIME ALERT: INTLTESTDB;STARTED; Host has entered a period of scheduled downtime
[1459525500] HOST DOWNTIME ALERT: R25-ZETA2012;STARTED; Host has entered a period of scheduled downtime
[1459525500] HOST DOWNTIME ALERT: R25-ALPHA2012;STARTED; Host has entered a period of scheduled downtime
[1459525500] HOST DOWNTIME ALERT: ADDC5;STARTED; Host has entered a period of scheduled downtime
[1459525500] HOST DOWNTIME ALERT: TARGETXTEST2012;STARTED; Host has entered a period of scheduled downtime
[1459525682] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1459525862] SERVICE ALERT: C4SAS-SE;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459525982] SERVICE ALERT: C4SAS-SE;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1459526012] SERVICE ALERT: rooseveltpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459526072] SERVICE ALERT: rooseveltpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459526132] SERVICE ALERT: rooseveltpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459526287] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1459526407] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1459526527] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1459526732] SERVICE ALERT: rooseveltpk;Device Status;OK;HARD;3;Open top cover.
[1459527603] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 65% - OK!
[1459528327] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 38% (5 min average)
[1459529527] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1459529647] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1459529722] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 100% - OK!
[1459529767] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 92% (5 min average)
[1459530367] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1459530367] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1459530493] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459530613] SERVICE ALERT: BOEAPPTEST;CPU Load;OK;SOFT;2;CPU Load 77% (5 min average)
[1459531213] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459531333] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1459531347] SERVICE ALERT: hallegroundcolor;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 1 Trays Empty
[1459531395] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;UNKNOWN;SOFT;1;NSClient - ERROR: Invalid return from command: check_drivesize
[1459531406] SERVICE ALERT: hallegroundcolor;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 1 Trays Empty
[1459531454] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1459531466] SERVICE ALERT: hallegroundcolor;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 1 Trays Empty
[1459531515] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 84.90 Gb - used: 76.46 Gb (90%) - free 8.44 Gb (10%)
[1459531635] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;OK;SOFT;3;c:\ - total: 84.90 Gb - used: 76.39 Gb (90%) - free 8.51 Gb (10%)
[1459532053] SERVICE ALERT: BOEAPPTEST;CPU Load;OK;HARD;3;CPU Load 56% (5 min average)
[1459532235] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 84.90 Gb - used: 76.62 Gb (90%) - free 8.28 Gb (10%)
[1459532357] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 84.90 Gb - used: 76.61 Gb (90%) - free 8.29 Gb (10%)
[1459532475] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 84.90 Gb - used: 76.61 Gb (90%) - free 8.29 Gb (10%)
[1459532653] SERVICE ALERT: BOEAPPTEST;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459532767] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 91% (5 min average)
[1459532773] SERVICE ALERT: BOEAPPTEST;CPU Load;OK;SOFT;2;CPU Load 87% (5 min average)
[1459533266] SERVICE ALERT: hallegroundcolor;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1459533367] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (5 min average)
[1459534567] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459534587] SERVICE ALERT: ADCA2;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1459534684] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;SOFT;1;WARNING: Replication with ADDC5 failed!
[1459534687] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 95% (5 min average)
[1459534706] SERVICE ALERT: ADCA2;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1459534769] SERVICE DOWNTIME ALERT: SCCM-PRIM01;CPU Load;STARTED; Service has entered a period of scheduled downtime
[1459534769] HOST DOWNTIME ALERT: BBTS_DB;STARTED; Host has entered a period of scheduled downtime
[1459534769] HOST DOWNTIME ALERT: ADDC2;STARTED; Host has entered a period of scheduled downtime
[1459534769] HOST DOWNTIME ALERT: AX-QUAL;STARTED; Host has entered a period of scheduled downtime
[1459534769] HOST DOWNTIME ALERT: ITRFTEST;STARTED; Host has entered a period of scheduled downtime
[1459534769] HOST DOWNTIME ALERT: MAPS01TEST;STARTED; Host has entered a period of scheduled downtime
[1459534769] HOST DOWNTIME ALERT: AXIOM-TEST-2012;STARTED; Host has entered a period of scheduled downtime
[1459534769] HOST DOWNTIME ALERT: R25-ALPHA2012;STARTED; Host has entered a period of scheduled downtime
[1459534769] HOST DOWNTIME ALERT: R25-DELTA2012;STARTED; Host has entered a period of scheduled downtime
[1459534769] HOST DOWNTIME ALERT: ADDC3;STARTED; Host has entered a period of scheduled downtime
[1459534769] HOST DOWNTIME ALERT: AXLIC-DEVL;STARTED; Host has entered a period of scheduled downtime
[1459534769] HOST DOWNTIME ALERT: R25-ZETA2012;STARTED; Host has entered a period of scheduled downtime
[1459534769] HOST DOWNTIME ALERT: SANDBOX-APP;STARTED; Host has entered a period of scheduled downtime
[1459534769] HOST DOWNTIME ALERT: BOEAPPTEST;STARTED; Host has entered a period of scheduled downtime
[1459534769] HOST DOWNTIME ALERT: TARGETXTEST2012;STARTED; Host has entered a period of scheduled downtime
[1459534769] SERVICE DOWNTIME ALERT: BOEAPPTEST;C:\ Drive Space;STARTED; Service has entered a period of scheduled downtime
[1459534769] HOST DOWNTIME ALERT: ADDC4;STARTED; Host has entered a period of scheduled downtime
[1459534769] HOST DOWNTIME ALERT: EVENTSTEST-2012;STARTED; Host has entered a period of scheduled downtime
[1459534769] HOST DOWNTIME ALERT: INTLAPP-ISSMAP;STARTED; Host has entered a period of scheduled downtime
[1459534769] HOST DOWNTIME ALERT: ADDC5;STARTED; Host has entered a period of scheduled downtime
[1459534769] HOST DOWNTIME ALERT: INTLTESTAPP;STARTED; Host has entered a period of scheduled downtime
[1459534769] HOST DOWNTIME ALERT: INTLTESTDB;STARTED; Host has entered a period of scheduled downtime
[1459534804] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;SOFT;2;WARNING: Replication with ADDC5 failed!
[1459534807] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 93% (5 min average)
[1459534807] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STARTED; Service appears to have started flapping (38.1% change >= 20.0% threshold)
[1459534924] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;HARD;3;WARNING: Replication with ADDC5 failed!
[1459536007] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1459536607] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 81% (5 min average)
[1459536731] SERVICE ALERT: NETIDADPROXY;Service: DNS Client;CRITICAL;SOFT;1;Dnscache: Stopped
[1459536851] SERVICE ALERT: NETIDADPROXY;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1459536978] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459537028] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459537038] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459537068] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459537082] SERVICE ALERT: EVENTSTEST-2012;CPU Load;WARNING;SOFT;1;CPU Load 91% (5 min average)
[1459537088] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459537098] SERVICE ALERT: prayharroldpk;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459537104] SERVICE ALERT: SANDBOX-APP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459537128] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459537148] SERVICE ALERT: prayharroldpk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459537158] HOST ALERT: prayharroldpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459537188] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.88.118. Is the printer off?  Is the network down?
[1459537196] HOST ALERT: prayharroldpk;UP;SOFT;2;PING OK - Packet loss = 44%, RTA = 0.58 ms
[1459537202] SERVICE ALERT: EVENTSTEST-2012;CPU Load;WARNING;SOFT;2;CPU Load 90% (5 min average)
[1459537207] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459537224] SERVICE ALERT: SANDBOX-APP;CPU Load;CRITICAL;SOFT;2;CPU Load 96% (5 min average)
[1459537227] SERVICE ALERT: ADDC2;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459537236] SERVICE ALERT: ADDC2;Service: AD Domain Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459537292] SERVICE ALERT: AXIOM-TEST-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459537322] SERVICE ALERT: EVENTSTEST-2012;CPU Load;WARNING;HARD;3;CPU Load 94% (5 min average)
[1459537327] SERVICE ALERT: ADDC2;Memory Usage;OK;SOFT;2;Memory usage: total:20479.66 MB - used: 976.32 MB (5%) - free: 19503.34 MB (95%)
[1459537327] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1459537330] SERVICE ALERT: ADDC3;Service: User Profile;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1459537330] SERVICE ALERT: ADDC3;Service: Dell Backup Agent;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1459537336] SERVICE ALERT: ADDC2;Service: AD Domain Service;OK;SOFT;2;NTDS: Started
[1459537344] SERVICE ALERT: SANDBOX-APP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1459537356] SERVICE ALERT: ADDC3;Service: AD Web Service;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459537359] SERVICE ALERT: ADDC3;CPU Load;CRITICAL;SOFT;1;connect to address 164.76.99.237 and port 12489: Connection refused
[1459537362] SERVICE ALERT: ADDC3;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.99.237 and port 12489: Connection refused
[1459537363] SERVICE ALERT: ADDC3;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459537412] SERVICE ALERT: AXIOM-TEST-2012;CPU Load;OK;SOFT;2;CPU Load 88% (5 min average)
[1459537447] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1459537449] SERVICE ALERT: ADDC3;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1459537450] SERVICE ALERT: ADDC3;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1459537450] SERVICE ALERT: ADDC3;Service: Dell Backup Agent;CRITICAL;SOFT;2;ErdAgent: Starting
[1459537456] SERVICE ALERT: ADDC3;Service: AD Web Service;OK;SOFT;2;ADWS: Started
[1459537463] SERVICE ALERT: ADDC3;Service: RPC;OK;SOFT;2;RpcSs: Started
[1459537473] SERVICE ALERT: ADDC3;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1459537476] SERVICE ALERT: ADDC3;CPU Load;OK;SOFT;2;CPU Load 1% (5 min average)
[1459537569] SERVICE ALERT: ADDC3;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1459537570] SERVICE ALERT: ADDC3;Service: Dell Backup Agent;OK;SOFT;3;ErdAgent: Started
[1459537682] SERVICE ALERT: ADDC4;AD Replication;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1459537682] SERVICE ALERT: ADDC4;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.99.238 and port 12489: Connection refused
[1459537692] SERVICE ALERT: prayharroldpk;Device Status;OK;HARD;3;Processing job from tray 4
[1459537704] SERVICE ALERT: ADDC4;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459537743] SERVICE ALERT: prayharroldpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1459537775] SERVICE ALERT: ADDC4;Service: Dell Backup Agent;CRITICAL;SOFT;1;ErdAgent: Starting
[1459537782] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 21% - OK!
[1459537802] SERVICE ALERT: ADDC4;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1459537804] SERVICE ALERT: ADDC4;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1459537813] SERVICE ALERT: ADDC4;AD Replication;OK;SOFT;2;OK: replication is up and running
[1459537895] SERVICE ALERT: ADDC4;Service: Dell Backup Agent;OK;SOFT;2;ErdAgent: Started
[1459537922] SERVICE ALERT: EVENTSTEST-2012;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1459538009] SERVICE ALERT: INTLAPP-ISSMAP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459538064] SERVICE ALERT: INTLTESTAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459538129] SERVICE ALERT: INTLAPP-ISSMAP;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1459538172] SERVICE ALERT: R25-ZETA2012;C:\ Drive Space;UNKNOWN;SOFT;1;NSClient - ERROR: Invalid return from command: check_drivesize
[1459538184] SERVICE ALERT: INTLTESTAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1459538222] SERVICE ALERT: R25-ALPHA2012;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.99.204 and port 12489: Connection refused
[1459538225] SERVICE ALERT: R25-ZETA2012;CPU Load;CRITICAL;SOFT;1;connect to address 164.76.99.206 and port 12489: Connection refused
[1459538249] SERVICE ALERT: INTLAPP-ISSMAP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1459538292] SERVICE ALERT: R25-ZETA2012;C:\ Drive Space;OK;SOFT;2;c:\ - total: 39.66 Gb - used: 17.12 Gb (43%) - free 22.53 Gb (57%)
[1459538304] SERVICE ALERT: INTLTESTAPP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1459538342] SERVICE ALERT: R25-ZETA2012;CPU Load;OK;SOFT;2;CPU Load 68% (5 min average)
[1459538342] SERVICE ALERT: R25-ALPHA2012;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1459538480] SERVICE ALERT: SANDBOX-APP;Service: Firewall;CRITICAL;SOFT;1;connect to address 164.76.99.175 and port 12489: Connection refused
[1459538510] HOST ALERT: SANDBOX-APP;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459538522] SERVICE ALERT: EVENTSTEST-2012;CPU Load;OK;HARD;3;CPU Load 74% (5 min average)
[1459538534] HOST ALERT: SANDBOX-APP;UP;SOFT;2;PING OK - Packet loss = 58%, RTA = 0.35 ms
[1459538544] SERVICE ALERT: SANDBOX-APP;CPU Load;OK;HARD;3;CPU Load 59% (5 min average)
[1459538600] SERVICE ALERT: SANDBOX-APP;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1459538641] SERVICE ALERT: AXIOM-TEST-2012;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459538741] SERVICE ALERT: AXIOM-TEST-2012;C:\ Drive Space;OK;SOFT;2;c:\ - total: 39.66 Gb - used: 24.99 Gb (63%) - free 14.67 Gb (37%)
[1459538882] SERVICE ALERT: INTLTESTDB;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459539002] SERVICE ALERT: INTLTESTDB;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1459539113] SERVICE ALERT: EVENTSTEST-2012;Service: RPC;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1459539122] SERVICE ALERT: EVENTSTEST-2012;Service: Windows Time;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1459539122] SERVICE ALERT: INTLTESTDB;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1459539155] SERVICE ALERT: EVENTSTEST-2012;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459539182] SERVICE ALERT: INTLTESTAPP;Service: Firewall;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1459539212] HOST ALERT: INTLTESTAPP;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1459539233] SERVICE ALERT: EVENTSTEST-2012;Service: RPC;OK;SOFT;2;RpcSs: Started
[1459539242] SERVICE ALERT: EVENTSTEST-2012;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1459539247] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 90% (5 min average)
[1459539254] SERVICE ALERT: INTLAPP-ISSMAP;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.99.188 and port 12489: Connection refused
[1459539254] SERVICE ALERT: INTLTESTDB;Service: RPC;CRITICAL;SOFT;1;connect to address 164.76.7.82 and port 12489: Connection refused
[1459539255] SERVICE ALERT: EVENTSTEST-2012;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 2 minute(s)
[1459539270] HOST ALERT: ITRFTEST;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459539276] HOST ALERT: INTLTESTAPP;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1459539302] SERVICE ALERT: INTLTESTAPP;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1459539334] HOST ALERT: ITRFTEST;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1459539374] SERVICE ALERT: INTLAPP-ISSMAP;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1459539374] SERVICE ALERT: INTLTESTDB;Service: RPC;OK;SOFT;2;RpcSs: Started
[1459539400] SERVICE ALERT: R25-GAMMA2012;VMWare Tools;CRITICAL;SOFT;1;connect to address 164.76.99.209 and port 12489: Connection refused
[1459539490] SERVICE ALERT: R25-EPSILON2012;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1459539504] SERVICE ALERT: INTLTESTAPP;CPU Load;OK;HARD;3;CPU Load 89% (5 min average)
[1459539517] SERVICE ALERT: R25-GAMMA2012;VMWare Tools;OK;SOFT;2;VMTools: Started
[1459539585] SERVICE ALERT: TARGETXTEST2012;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1459539610] SERVICE ALERT: R25-EPSILON2012;CPU Load;OK;SOFT;2;CPU Load 52% (5 min average)
[1459539705] SERVICE ALERT: TARGETXTEST2012;CPU Load;OK;SOFT;2;CPU Load 27% (5 min average)
[1459539722] SERVICE ALERT: EVENTSTEST-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459539823] SERVICE ALERT: ADDC5;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1459539826] SERVICE ALERT: ADDC5;Service: Dell Backup Agent;CRITICAL;SOFT;1;ErdAgent: Starting
[1459539842] SERVICE ALERT: EVENTSTEST-2012;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1459539847] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 85% (5 min average)
[1459539943] SERVICE ALERT: ADDC5;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1459539946] SERVICE ALERT: ADDC5;Service: Dell Backup Agent;OK;SOFT;2;ErdAgent: Started
[1459539962] SERVICE ALERT: EVENTSTEST-2012;CPU Load;OK;SOFT;3;CPU Load 67% (5 min average)
[1459540049] SERVICE ALERT: INTLAPP-ISSMAP;CPU Load;OK;HARD;3;CPU Load 3% (5 min average)
[1459540079] SERVICE ALERT: VC;Memory Usage;WARNING;SOFT;1;No data was received from host!
[1459540092] SERVICE ALERT: sillpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1459540104] SERVICE ALERT: INTLTESTAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459540153] SERVICE ALERT: sillpk;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! All Trays Still Full
[1459540180] SERVICE ALERT: VC;Memory Usage;OK;SOFT;2;Memory usage: total:28669.44 MB - used: 6910.10 MB (24%) - free: 21759.34 MB (76%)
[1459540224] SERVICE ALERT: INTLTESTAPP;CPU Load;OK;SOFT;2;CPU Load 28% (5 min average)
[1459540275] SERVICE ALERT: BOEAPPTEST;C:\ Drive Space;OK;HARD;3;c:\ - total: 99.90 Gb - used: 76.79 Gb (77%) - free 23.11 Gb (23%)
[1459540305] SERVICE ALERT: TARGETXTEST2012;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459540322] SERVICE ALERT: INTLTESTDB;CPU Load;OK;HARD;3;CPU Load 6% (5 min average)
[1459540425] SERVICE ALERT: TARGETXTEST2012;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1459540545] SERVICE ALERT: TARGETXTEST2012;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1459540824] SERVICE ALERT: INTLTESTAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1459540944] SERVICE ALERT: INTLTESTAPP;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1459541145] SERVICE ALERT: TARGETXTEST2012;CPU Load;WARNING;HARD;3;CPU Load 94% (5 min average)
[1459541244] SERVICE ALERT: TARGETXTEST2012;Service: WMI;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1459541266] SERVICE ALERT: TARGETXTEST2012;Service: Group Policy;CRITICAL;SOFT;1;connect to address 164.76.99.153 and port 12489: Connection refused
[1459541291] SERVICE ALERT: TARGETXTEST2012;Service: Windows Time;CRITICAL;SOFT;1;W32Time: Stopped
[1459541321] HOST ALERT: TARGETXTEST2012;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459541364] SERVICE ALERT: TARGETXTEST2012;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1459541368] HOST ALERT: TARGETXTEST2012;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1459541383] SERVICE ALERT: TARGETXTEST2012;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1459541402] SERVICE ALERT: TARGETXTEST2012;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1459541527] SERVICE ALERT: VUM;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459541632] SERVICE ALERT: VCENTERSSO;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459541647] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1459541647] SERVICE ALERT: VUM;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1459541745] SERVICE ALERT: TARGETXTEST2012;CPU Load;OK;HARD;3;CPU Load 0% (5 min average)
[1459541752] SERVICE ALERT: VCENTERSSO;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1459541767] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1459541767] SERVICE ALERT: VUM;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1459541872] SERVICE ALERT: VCENTERSSO;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1459541887] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 84% (5 min average)
[1459542024] SERVICE ALERT: VC;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459542144] SERVICE ALERT: VC;CPU Load;CRITICAL;SOFT;2;CPU Load 97% (5 min average)
[1459542264] SERVICE ALERT: VC;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1459542487] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1459542607] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1459542727] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 85% (5 min average)
[1459542864] SERVICE ALERT: VC;CPU Load;OK;HARD;3;CPU Load 63% (5 min average)
[1459543072] SERVICE ALERT: VCENTERSSO;CPU Load;OK;HARD;3;CPU Load 3% (5 min average)
[1459543095] SERVICE ALERT: VUM;VMWare Tools;CRITICAL;SOFT;1;connect to address 164.76.2.88 and port 12489: Connection refused
[1459543102] SERVICE ALERT: VUM;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459543106] SERVICE ALERT: VUM;Uptime;CRITICAL;SOFT;1;connect to address 164.76.2.88 and port 12489: Connection refused
[1459543125] HOST ALERT: VUM;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459543149] HOST ALERT: VUM;UP;SOFT;2;PING OK - Packet loss = 77%, RTA = 0.44 ms
[1459543202] SERVICE ALERT: VUM;Service: Group Policy;OK;SOFT;2;gpsvc: Started
[1459543212] SERVICE ALERT: VUM;VMWare Tools;OK;SOFT;2;VMTools: Started
[1459543226] SERVICE ALERT: VUM;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 2 minute(s)
[1459543327] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1459543447] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1459543458] SERVICE ALERT: VC;Uptime;CRITICAL;SOFT;1;connect to address 164.76.158.200 and port 12489: Connection refused
[1459543464] SERVICE ALERT: VC;CPU Load;CRITICAL;SOFT;1;connect to address 164.76.158.200 and port 12489: Connection refused
[1459543488] HOST ALERT: VC;DOWN;SOFT;1;(Host check timed out after 30.00 seconds)
[1459543500] HOST DOWNTIME ALERT: AXIOM-TEST-2012;STOPPED; Host has exited from a period of scheduled downtime
[1459543500] HOST DOWNTIME ALERT: ADDC2;STOPPED; Host has exited from a period of scheduled downtime
[1459543500] HOST DOWNTIME ALERT: AX-QUAL;STOPPED; Host has exited from a period of scheduled downtime
[1459543500] HOST DOWNTIME ALERT: MAPS01TEST;STOPPED; Host has exited from a period of scheduled downtime
[1459543500] HOST DOWNTIME ALERT: ITRFTEST;STOPPED; Host has exited from a period of scheduled downtime
[1459543500] HOST DOWNTIME ALERT: SANDBOX-APP;STOPPED; Host has exited from a period of scheduled downtime
[1459543500] HOST DOWNTIME ALERT: R25-ALPHA2012;STOPPED; Host has exited from a period of scheduled downtime
[1459543500] HOST DOWNTIME ALERT: R25-DELTA2012;STOPPED; Host has exited from a period of scheduled downtime
[1459543500] HOST DOWNTIME ALERT: ADDC3;STOPPED; Host has exited from a period of scheduled downtime
[1459543500] HOST DOWNTIME ALERT: BOEAPPTEST;STOPPED; Host has exited from a period of scheduled downtime
[1459543500] HOST DOWNTIME ALERT: AXLIC-DEVL;STOPPED; Host has exited from a period of scheduled downtime
[1459543500] HOST DOWNTIME ALERT: R25-ZETA2012;STOPPED; Host has exited from a period of scheduled downtime
[1459543500] HOST DOWNTIME ALERT: INTLTESTAPP;STOPPED; Host has exited from a period of scheduled downtime
[1459543500] HOST DOWNTIME ALERT: ADDC4;STOPPED; Host has exited from a period of scheduled downtime
[1459543500] HOST DOWNTIME ALERT: INTLTESTDB;STOPPED; Host has exited from a period of scheduled downtime
[1459543500] HOST DOWNTIME ALERT: TARGETXTEST2012;STOPPED; Host has exited from a period of scheduled downtime
[1459543500] HOST DOWNTIME ALERT: INTLAPP-ISSMAP;STOPPED; Host has exited from a period of scheduled downtime
[1459543500] HOST DOWNTIME ALERT: EVENTSTEST-2012;STOPPED; Host has exited from a period of scheduled downtime
[1459543500] HOST DOWNTIME ALERT: ADDC5;STOPPED; Host has exited from a period of scheduled downtime
[1459543552] HOST ALERT: VC;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.36 ms
[1459543567] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1459543579] SERVICE ALERT: VC;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 2 minute(s)
[1459543584] SERVICE ALERT: VC;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1459543672] SERVICE ALERT: VCENTERSSO;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459543704] SERVICE ALERT: VC;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1459543792] SERVICE ALERT: VCENTERSSO;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1459544304] SERVICE ALERT: VC;CPU Load;OK;HARD;3;CPU Load 57% (5 min average)
[1459544304] SERVICE FLAPPING ALERT: VC;CPU Load;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1459544767] SERVICE ALERT: VUM;CPU Load;OK;HARD;3;CPU Load 0% (5 min average)
[1459545578] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1459545638] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1459545698] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1459545729] SERVICE ALERT: ADRODC-LIV;AD Replication;OK;HARD;3;OK: replication is up and running
[1459547451] HOST ALERT: SNAP;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459547541] HOST ALERT: SNAP;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1459547612] HOST ALERT: SNAP;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1459547659] HOST ALERT: SNAP;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1459547700] HOST ALERT: SNAP;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1459547790] HOST ALERT: SNAP;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1459547880] HOST ALERT: SNAP;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1459547970] HOST ALERT: SNAP;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1459548046] SERVICE DOWNTIME ALERT: BOEAPPTEST;C:\ Drive Space;STOPPED; Service has exited from a period of scheduled downtime
[1459548060] HOST ALERT: SNAP;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1459548150] HOST ALERT: SNAP;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1459548967] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 79% (5 min average)
[1459549567] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459549687] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1459549807] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1459549942] SERVICE ALERT: SNAP;NSClient++ Version;CRITICAL;SOFT;1;connect to address 164.76.102.40 and port 12489: No route to host
[1459549987] HOST ALERT: SNAP;UP;HARD;10;PING OK - Packet loss = 16%, RTA = 0.50 ms
[1459550070] SERVICE ALERT: SNAP;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1459550090] SERVICE ALERT: SNAP;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459550170] SERVICE ALERT: SNAP;NSClient++ Version;OK;SOFT;3;NSClient++ 0.4.4.15 2015-11-25
[1459550190] SERVICE ALERT: SNAP;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1459551336] HOST ALERT: SNAP;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1459551400] HOST ALERT: SNAP;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1459551490] HOST ALERT: SNAP;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1459551580] HOST ALERT: SNAP;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1459551664] HOST ALERT: SNAP;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1459551754] HOST ALERT: SNAP;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1459551812] HOST ALERT: SNAP;DOWN;SOFT;7;(Host check timed out after 30.01 seconds)
[1459551859] HOST ALERT: SNAP;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1459551949] HOST ALERT: SNAP;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1459552000] HOST ALERT: SNAP;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1459553407] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 86% (5 min average)
[1459554007] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1459554127] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 33% (5 min average)
[1459554726] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1459554846] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1459554966] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 98% (5 min average)
[1459555702] SERVICE FLAPPING ALERT: VC;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1459556166] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 74% (5 min average)
[1459556523] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;SOFT;1;WARNING: Replication with ADDC5 failed!
[1459556643] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;SOFT;2;WARNING: Replication with ADDC5 failed!
[1459556763] SERVICE ALERT: ADRODC-LIV;AD Replication;WARNING;HARD;3;WARNING: Replication with ADDC5 failed!
[1459556766] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 93% (5 min average)
[1459556886] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1459557006] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1459559231] SERVICE ALERT: TARGETX-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459559351] SERVICE ALERT: TARGETX-2012;CPU Load;CRITICAL;SOFT;2;CPU Load 95% (5 min average)
[1459559471] SERVICE ALERT: TARGETX-2012;CPU Load;CRITICAL;HARD;3;CPU Load 96% (5 min average)
[1459560071] SERVICE ALERT: TARGETX-2012;CPU Load;OK;HARD;3;CPU Load 85% (5 min average)
[1459563440] SERVICE ALERT: VUM;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459563540] SERVICE ALERT: VUM;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1459563592] SERVICE ALERT: MBAMDB;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459563692] SERVICE ALERT: MBAMDB;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1459563956] SERVICE ALERT: CBORD;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1459564056] SERVICE ALERT: CBORD;VMWare Tools;OK;SOFT;2;VMTools: Started
[1459566439] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1459566559] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1459566679] SERVICE ALERT: BOEAPP;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1459566967] SERVICE ALERT: ADRODC-LIV;AD Replication;OK;HARD;3;OK: replication is up and running
[1459567879] SERVICE ALERT: BOEAPP;CPU Load;WARNING;HARD;3;CPU Load 93% (5 min average)
[1459568406] SERVICE FLAPPING ALERT: SCCM-PRIM01;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1459568479] SERVICE ALERT: BOEAPP;CPU Load;OK;HARD;3;CPU Load 79% (5 min average)