[1454552767] HOST DOWNTIME ALERT: AD-RMAD;STARTED; Host has entered a period of scheduled downtime
[1454552767] HOST DOWNTIME ALERT: EMU-VALUE01;STARTED; Host has entered a period of scheduled downtime
[1454552767] HOST DOWNTIME ALERT: EMUSCRIPT;STARTED; Host has entered a period of scheduled downtime
[1454552767] HOST DOWNTIME ALERT: EMUSQLEXPRESS;STARTED; Host has entered a period of scheduled downtime
[1454552767] HOST DOWNTIME ALERT: HDBACKUPS;STARTED; Host has entered a period of scheduled downtime
[1454552767] HOST DOWNTIME ALERT: IMAGES;STARTED; Host has entered a period of scheduled downtime
[1454552767] HOST DOWNTIME ALERT: LANSWEEPER;STARTED; Host has entered a period of scheduled downtime
[1454552767] HOST DOWNTIME ALERT: MBAM;STARTED; Host has entered a period of scheduled downtime
[1454552767] HOST DOWNTIME ALERT: MBAMDB;STARTED; Host has entered a period of scheduled downtime
[1454552767] HOST DOWNTIME ALERT: OFFICEKMS;STARTED; Host has entered a period of scheduled downtime
[1454552767] HOST DOWNTIME ALERT: REFWIN;STARTED; Host has entered a period of scheduled downtime
[1454552767] HOST DOWNTIME ALERT: SANDBOX-DB;STARTED; Host has entered a period of scheduled downtime
[1454552872] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;SOFT;1;e:\ - total: 39.87 Gb - used: 34.86 Gb (87%) - free 5.02 Gb (13%)
[1454552874] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1454552904] HOST ALERT: owen101lab3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1454552931] SERVICE ALERT: owen101lab3;Device Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1454552934] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1454552941] SERVICE ALERT: DB01-2012;CPU Load;WARNING;SOFT;1;CPU Load 88% (5 min average)
[1454552956] HOST ALERT: hallethirdsouth;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1454552961] HOST ALERT: owen101lab3;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1454552988] SERVICE ALERT: owen101lab3;Paper Tray Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1454552991] SERVICE ALERT: owen101lab3;Device Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1454552992] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;SOFT;2;e:\ - total: 39.87 Gb - used: 34.86 Gb (87%) - free 5.02 Gb (13%)
[1454552994] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;WARNING;HARD;3;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1454553018] HOST ALERT: owen101lab3;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1454553031] SERVICE ALERT: SANDBOX-DB;C:\ Drive Space;WARNING;SOFT;1;c:\ - total: 39.66 Gb - used: 32.54 Gb (82%) - free 7.11 Gb (18%)
[1454553046] HOST ALERT: hallethirdsouth;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1454553048] SERVICE ALERT: owen101lab3;Paper Tray Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1454553051] SERVICE ALERT: owen101lab3;Device Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1454553061] SERVICE ALERT: DB01-2012;CPU Load;OK;SOFT;2;CPU Load 76% (5 min average)
[1454553078] HOST ALERT: owen101lab3;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1454553098] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at 4% - WARNING!
[1454553108] SERVICE ALERT: owen101lab3;Paper Tray Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.203.103! Make sure host is up and SNMP is configured properly.
[1454553112] SERVICE ALERT: AD-RMAD;E:\ Drive Space;WARNING;HARD;3;e:\ - total: 39.87 Gb - used: 34.86 Gb (87%) - free 5.02 Gb (13%)
[1454553130] SERVICE ALERT: DB01-2012;CPU Load;WARNING;SOFT;1;CPU Load 81% (5 min average)
[1454553136] HOST ALERT: hallethirdsouth;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1454553138] HOST ALERT: owen101lab3;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1454553151] SERVICE ALERT: SANDBOX-DB;C:\ Drive Space;WARNING;SOFT;2;c:\ - total: 39.66 Gb - used: 32.54 Gb (82%) - free 7.11 Gb (18%)
[1454553158] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at 4% - WARNING!
[1454553199] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at WARNING level!
[1454553216] SERVICE ALERT: hallethirdsouth;Device Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.107.213! Make sure host is up and SNMP is configured properly.
[1454553219] SERVICE ALERT: prayharroldpk;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at 4% - WARNING!
[1454553224] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 3 Trays Empty
[1454553226] HOST ALERT: hallethirdsouth;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1454553228] HOST ALERT: owen101lab3;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1454553250] SERVICE ALERT: DB01-2012;CPU Load;CRITICAL;SOFT;2;CPU Load 90% (5 min average)
[1454553259] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at WARNING level!
[1454553271] SERVICE ALERT: SANDBOX-DB;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 39.66 Gb - used: 32.54 Gb (82%) - free 7.11 Gb (18%)
[1454553271] SERVICE ALERT: strongpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1454553273] SERVICE ALERT: hallethirdsouth;Paper Tray Status;WARNING;SOFT;1;WARNING: No SNMP response from 164.76.107.213! Make sure host is up and SNMP is configured properly.
[1454553276] SERVICE ALERT: hallethirdsouth;Device Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.107.213! Make sure host is up and SNMP is configured properly.
[1454553284] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 3 Trays Empty
[1454553303] HOST ALERT: hallethirdsouth;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1454553318] HOST ALERT: owen101lab3;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1454553320] SERVICE ALERT: hallefirst2;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at WARNING level!
[1454553332] SERVICE ALERT: strongpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1454553333] SERVICE ALERT: hallethirdsouth;Paper Tray Status;WARNING;SOFT;2;WARNING: No SNMP response from 164.76.107.213! Make sure host is up and SNMP is configured properly.
[1454553336] SERVICE ALERT: hallethirdsouth;Device Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.107.213! Make sure host is up and SNMP is configured properly.
[1454553344] SERVICE ALERT: markjeffersonpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 3 Trays Empty
[1454553363] HOST ALERT: hallethirdsouth;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1454553370] SERVICE ALERT: DB01-2012;CPU Load;WARNING;HARD;3;CPU Load 83% (5 min average)
[1454553391] SERVICE ALERT: strongpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1454553393] SERVICE ALERT: hallethirdsouth;Paper Tray Status;WARNING;HARD;3;WARNING: No SNMP response from 164.76.107.213! Make sure host is up and SNMP is configured properly.
[1454553408] HOST ALERT: owen101lab3;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1454553423] HOST ALERT: hallethirdsouth;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1454553498] HOST ALERT: owen101lab3;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1454553513] HOST ALERT: hallethirdsouth;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1454553588] HOST ALERT: owen101lab3;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1454553603] HOST ALERT: hallethirdsouth;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1454553693] HOST ALERT: hallethirdsouth;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1454553970] SERVICE ALERT: DB01-2012;CPU Load;OK;HARD;3;CPU Load 19% (5 min average)
[1454555757] SERVICE ALERT: VC;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1454555857] SERVICE ALERT: VC;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1454556600] HOST DOWNTIME ALERT: LANSWEEPER;STARTED; Host has entered a period of scheduled downtime
[1454556600] HOST DOWNTIME ALERT: ADWSUSCLIENTS;STARTED; Host has entered a period of scheduled downtime
[1454556600] HOST DOWNTIME ALERT: EMUSCRIPT;STARTED; Host has entered a period of scheduled downtime
[1454556600] HOST DOWNTIME ALERT: MBAM;STARTED; Host has entered a period of scheduled downtime
[1454556600] HOST DOWNTIME ALERT: OFFICEKMS;STARTED; Host has entered a period of scheduled downtime
[1454556600] HOST DOWNTIME ALERT: MBAMDB;STARTED; Host has entered a period of scheduled downtime
[1454556600] HOST DOWNTIME ALERT: EMUSQLEXPRESS;STARTED; Host has entered a period of scheduled downtime
[1454556600] HOST DOWNTIME ALERT: REFWIN;STARTED; Host has entered a period of scheduled downtime
[1454556600] HOST DOWNTIME ALERT: HDBACKUPS;STARTED; Host has entered a period of scheduled downtime
[1454556600] HOST DOWNTIME ALERT: VC;STARTED; Host has entered a period of scheduled downtime
[1454556600] HOST DOWNTIME ALERT: IMAGES;STARTED; Host has entered a period of scheduled downtime
[1454556600] HOST DOWNTIME ALERT: AD-RMAD;STARTED; Host has entered a period of scheduled downtime
[1454556600] HOST DOWNTIME ALERT: ADWSUSSERVERS;STARTED; Host has entered a period of scheduled downtime
[1454558424] SERVICE ALERT: LANSWEEPER;CPU Load;CRITICAL;SOFT;1;CPU Load 92% (5 min average)
[1454558544] SERVICE ALERT: LANSWEEPER;CPU Load;WARNING;SOFT;2;CPU Load 85% (5 min average)
[1454558664] SERVICE ALERT: LANSWEEPER;CPU Load;CRITICAL;HARD;3;CPU Load 96% (5 min average)
[1454558736] SERVICE ALERT: LANSWEEPER;CPU Load;CRITICAL;SOFT;1;CPU Load 92% (5 min average)
[1454558856] SERVICE ALERT: LANSWEEPER;CPU Load;CRITICAL;SOFT;2;CPU Load 91% (5 min average)
[1454558976] SERVICE ALERT: LANSWEEPER;CPU Load;CRITICAL;HARD;3;CPU Load 92% (5 min average)
[1454559150] SERVICE ALERT: EMU-VALUE01;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1454559151] SERVICE ALERT: EMU-VALUE01;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1454559152] SERVICE ALERT: EMU-VALUE01;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1454559180] HOST ALERT: EMU-VALUE01;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1454559182] HOST ALERT: EMU-VALUE01;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1454559208] SERVICE ALERT: EMU-VALUE01;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1454559209] SERVICE ALERT: EMU-VALUE01;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1454559212] SERVICE ALERT: EMU-VALUE01;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1454559215] SERVICE ALERT: EMU-VALUE01;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1454559231] SERVICE ALERT: EMU-VALUE01;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1454559238] HOST ALERT: EMU-VALUE01;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1454559242] HOST ALERT: EMU-VALUE01;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1454559256] SERVICE ALERT: EMU-VALUE01;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1454559264] SERVICE ALERT: LANSWEEPER;CPU Load;OK;HARD;3;CPU Load 62% (5 min average)
[1454559265] SERVICE ALERT: EMU-VALUE01;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1454559266] SERVICE ALERT: EMU-VALUE01;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1454559270] SERVICE ALERT: EMU-VALUE01;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1454559271] SERVICE ALERT: EMU-VALUE01;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1454559272] SERVICE ALERT: EMU-VALUE01;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1454559278] SERVICE ALERT: EMU-VALUE01;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1454559286] HOST ALERT: EMU-VALUE01;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1454559300] HOST ALERT: EMU-VALUE01;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1454559328] SERVICE ALERT: EMU-VALUE01;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1454559329] SERVICE ALERT: EMU-VALUE01;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1454559332] SERVICE ALERT: EMU-VALUE01;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1454559335] SERVICE ALERT: EMU-VALUE01;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1454559340] SERVICE ALERT: EMU-VALUE01;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1454559342] SERVICE ALERT: EMU-VALUE01;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1454559355] HOST ALERT: EMU-VALUE01;UP;SOFT;4;PING OK - Packet loss = 79%, RTA = 0.45 ms
[1454559355] HOST ALERT: EMU-VALUE01;UP;SOFT;4;PING WARNING - Packet loss = 82%, RTA = 0.37 ms
[1454559356] SERVICE ALERT: EMU-VALUE01;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1454559365] SERVICE ALERT: EMU-VALUE01;CPU Load;OK;SOFT;2;CPU Load 6% (5 min average)
[1454559366] SERVICE ALERT: EMU-VALUE01;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 3 minute(s)
[1454559370] SERVICE ALERT: EMU-VALUE01;CPU Load;OK;SOFT;3;CPU Load 8% (5 min average)
[1454559371] SERVICE ALERT: EMU-VALUE01;Service: Windows Time;OK;SOFT;3;W32Time: Started
[1454559372] SERVICE ALERT: EMU-VALUE01;Service: WMI;OK;SOFT;3;Winmgmt: Started
[1454559377] SERVICE ALERT: EMU-VALUE01;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1454559428] SERVICE ALERT: EMU-VALUE01;C:\ Drive Space;OK;SOFT;3;c:\ - total: 465.42 Gb - used: 14.79 Gb (3%) - free 450.63 Gb (97%)
[1454559429] SERVICE ALERT: EMU-VALUE01;Service: Windows Time;OK;SOFT;3;W32Time: Started
[1454559432] SERVICE ALERT: EMU-VALUE01;Memory Usage;OK;SOFT;3;Memory usage: total:9444.43 MB - used: 2096.24 MB (22%) - free: 7348.19 MB (78%)
[1454559435] SERVICE ALERT: EMU-VALUE01;Uptime;OK;SOFT;3;System Uptime - 0 day(s) 0 hour(s) 4 minute(s)
[1454559442] SERVICE ALERT: EMU-VALUE01;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1454559575] SERVICE ALERT: LANSWEEPER;CPU Load;OK;HARD;3;CPU Load 0% (5 min average)
[1454559653] SERVICE ALERT: AD-RMAD;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.102.135 and port 12489: Connection refused
[1454559655] SERVICE ALERT: AD-RMAD;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1454559755] SERVICE ALERT: AD-RMAD;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1454559764] SERVICE ALERT: AD-RMAD;Service: DNS Client;OK;SOFT;2;Dnscache: Started