[1455599106] SERVICE ALERT: halleground1;Device Status;OK;HARD;3;Clearing paper path
[1455599106] SERVICE FLAPPING ALERT: halleground1;Device Status;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1455599130] SERVICE ALERT: halleground1;Device Status;OK;HARD;3;Clearing paper path
[1455602053] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;1;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1455602173] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;2;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1455602293] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;HARD;3;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1455609493] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;OK;HARD;3;OK: No entries found
[1455609906] SERVICE FLAPPING ALERT: halleground1;Device Status;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1455624009] SERVICE ALERT: DB01-2012;CPU Load;WARNING;SOFT;1;CPU Load 86% (5 min average)
[1455624129] SERVICE ALERT: DB01-2012;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1455627384] SERVICE ALERT: halleground2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1455627405] SERVICE ALERT: halleground2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1455627414] HOST ALERT: halleground2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1455627438] SERVICE ALERT: halleground2;Device Status;OK;SOFT;2;ORDER CARTRIDGE LESS THAN 1700 PAGES 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Processing job
[1455627442] HOST ALERT: halleground2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.58 ms
[1455627459] SERVICE ALERT: halleground2;Device Status;OK;SOFT;2;ORDER CARTRIDGE LESS THAN 1700 PAGES 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F   Ready
[1455629546] SERVICE ALERT: MBAM;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1455629666] SERVICE ALERT: MBAM;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1455632829] SERVICE ALERT: VC;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1455632859] HOST ALERT: VC;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1455632859] HOST ALERT: VC;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1455632923] HOST ALERT: VC;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.68 ms
[1455632923] HOST ALERT: VC;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.36 ms
[1455632929] SERVICE ALERT: VC;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1455633350] SERVICE ALERT: VC;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1455633394] SERVICE ALERT: VC;CPU Load;WARNING;SOFT;1;CPU Load 92% (5 min average)
[1455633470] SERVICE ALERT: VC;CPU Load;OK;SOFT;2;CPU Load 2% (5 min average)
[1455633514] SERVICE ALERT: VC;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1455643032] HOST ALERT: markjeffersonpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1455643096] HOST ALERT: markjeffersonpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.62 ms
[1455644425] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at 5% - WARNING!
[1455644432] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at 5% - WARNING!
[1455644485] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at 5% - WARNING!
[1455644492] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at 5% - WARNING!
[1455644545] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at 5% - WARNING!
[1455644552] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at 5% - WARNING!
[1455646251] SERVICE ALERT: ADCA2;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1455646371] SERVICE ALERT: ADCA2;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1455647148] SERVICE ALERT: porterpk;Black Cartridge HP CC364X;OK;HARD;3;Black Cartridge HP CC364X is at 100% - OK!
[1455647326] SERVICE ALERT: porterpk;Black Cartridge HP CC364X;OK;HARD;3;Black Cartridge HP CC364X is at 100% - OK!
[1455650309] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at 5% - WARNING!
[1455650368] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at 5% - WARNING!
[1455650427] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at 5% - WARNING!
[1455650548] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;WARNING;SOFT;1;Black Print Cartridge HP C8543X is at 5% - WARNING!
[1455650609] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;WARNING;SOFT;2;Black Print Cartridge HP C8543X is at 5% - WARNING!
[1455650668] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;WARNING;HARD;3;Black Print Cartridge HP C8543X is at 5% - WARNING!
[1455654152] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 5% - OK!
[1455654267] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 5% - OK!
[1455657319] SERVICE ALERT: hallesecondsouth;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1455657378] SERVICE ALERT: hallesecondsouth;Device Status;OK;SOFT;2;RETRIEVE JOB  NO STORED JOBS
[1455658157] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1455658167] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1455658216] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1455658227] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1455658276] SERVICE ALERT: hallefirst2;Device Status;OK;SOFT;3;MENUS  RETRIEVE JOB  INFORMATION  PAPER HANDLING
[1455658287] SERVICE ALERT: hallefirst2;Device Status;OK;SOFT;3;50 72 65 73 73 20 17 20 74 6F 20 73 65 65 20 73  74 65 70  by step information
[1455658876] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1455658888] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1455658936] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1455658947] SERVICE ALERT: hallefirst2;Device Status;OK;SOFT;2;RETRIEVE JOB  NO STORED JOBS
[1455658996] SERVICE ALERT: hallefirst2;Device Status;OK;SOFT;3;RETRIEVE JOB  NO STORED JOBS
[1455659547] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1455659580] SERVICE ALERT: SANDBOX-DB;Service: DNS Client;UNKNOWN;SOFT;1;(No output on stdout) stderr:
[1455659580] HOST ALERT: SANDBOX-DB;DOWN;SOFT;1;(No output on stdout) stderr:
[1455659596] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1455659607] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1455659633] HOST ALERT: SANDBOX-DB;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.41 ms
[1455659657] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1455659667] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1455659700] SERVICE ALERT: SANDBOX-DB;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1455659716] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1455661468] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Clearing paper path
[1455661516] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Clearing paper path
[1455663037] HOST ALERT: MBAMDB;DOWN;SOFT;1;CRITICAL - Network Unreachable (164.76.7.97)
[1455663037] SERVICE ALERT: OFFICEKMS;Service: DNS Client;CRITICAL;SOFT;1;connect to address 164.76.158.141 and port 12489: Network is unreachable
[1455663037] HOST ALERT: OFFICEKMS;DOWN;SOFT;1;CRITICAL - Network Unreachable (164.76.158.141)
[1455663099] HOST ALERT: MBAMDB;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.40 ms
[1455663100] SERVICE ALERT: CASPER;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1455663100] HOST ALERT: OFFICEKMS;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.47 ms
[1455663145] SERVICE ALERT: halleground2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 5% - OK!
[1455663157] SERVICE ALERT: OFFICEKMS;Service: DNS Client;OK;SOFT;2;Dnscache: Started
[1455663200] SERVICE ALERT: CASPER;Service: NetLogin;OK;SOFT;2;Netlogon: Started
[1455663627] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 5% - OK!
[1455667197] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 2 Trays Empty
[1455667257] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 2 Trays Empty
[1455667316] SERVICE ALERT: alexanderpk;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 2 Trays Empty
[1455670061] SERVICE ALERT: owen101lab1;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1455670121] SERVICE ALERT: owen101lab1;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1455670181] SERVICE ALERT: owen101lab1;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1455670182] SERVICE ALERT: owen101labcolor;Yellow Cartridge HP CE262A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1455670212] HOST ALERT: owen101labcolor;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1455670233] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1455670242] SERVICE ALERT: owen101labcolor;Yellow Cartridge HP CE262A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1455670272] HOST ALERT: owen101labcolor;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1455670293] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1455670302] SERVICE ALERT: owen101labcolor;Yellow Cartridge HP CE262A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1455670312] HOST ALERT: owen101labcolor;UP;SOFT;3;PING OK - Packet loss = 54%, RTA = 0.49 ms
[1455670344] SERVICE ALERT: owen101lab2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1455670350] SERVICE ALERT: owen101lab2;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1455670353] SERVICE ALERT: owen101lab2;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1455670398] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1455670404] SERVICE ALERT: owen101lab2;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1455670410] SERVICE ALERT: owen101lab2;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1455670423] SERVICE ALERT: owen101labcolor;Image Transfer Kit HP CE249A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1455670454] SERVICE ALERT: owen101labcolor;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1455670458] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1455670464] SERVICE ALERT: owen101lab2;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1455670470] SERVICE ALERT: owen101lab2;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.102. Is the printer off?  Is the network down?
[1455670475] SERVICE ALERT: owen101labcolor;Toner Collection Unit HP CE265A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1455670483] SERVICE ALERT: owen101labcolor;Image Transfer Kit HP CE249A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1455670514] SERVICE ALERT: owen101labcolor;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1455670518] SERVICE ALERT: owen101lab1;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1455670520] SERVICE ALERT: owen101lab1;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1455670535] SERVICE ALERT: owen101labcolor;Toner Collection Unit HP CE265A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1455670543] SERVICE ALERT: owen101labcolor;Image Transfer Kit HP CE249A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1455670574] SERVICE ALERT: owen101labcolor;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1455670580] SERVICE ALERT: owen101lab1;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1455670593] SERVICE ALERT: owen101labcolor;Black Cartridge HP CE260X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1455670595] SERVICE ALERT: owen101labcolor;Toner Collection Unit HP CE265A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1455670604] HOST ALERT: owen101labcolor;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1455670641] SERVICE ALERT: owen101lab1;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.101. Is the printer off?  Is the network down?
[1455670644] SERVICE ALERT: owen101labcolor;Magenta Cartridge HP CE263A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1455670648] HOST ALERT: owen101labcolor;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.45 ms
[1455670653] SERVICE ALERT: owen101labcolor;Black Cartridge HP CE260X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1455670675] SERVICE ALERT: owen101labcolor;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1455670695] SERVICE ALERT: owen101labcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1455670704] SERVICE ALERT: owen101labcolor;Magenta Cartridge HP CE263A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1455670713] SERVICE ALERT: owen101labcolor;Black Cartridge HP CE260X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1455670735] SERVICE ALERT: owen101labcolor;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1455670755] SERVICE ALERT: owen101labcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1455670764] SERVICE ALERT: owen101labcolor;Magenta Cartridge HP CE263A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1455670786] HOST ALERT: owen101lab3;UP;HARD;10;PING OK - Packet loss = 75%, RTA = 0.74 ms
[1455670795] SERVICE ALERT: owen101labcolor;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1455670815] SERVICE ALERT: owen101labcolor;Image Fuser Kit HP 110V-CE246A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1455670825] HOST ALERT: owen101labcolor;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1455670915] HOST ALERT: owen101labcolor;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1455670921] SERVICE ALERT: owen101labcolor;Cyan Cartridge HP CE261A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1455670951] HOST ALERT: owen101labcolor;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1455670981] SERVICE ALERT: owen101labcolor;Cyan Cartridge HP CE261A;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1455671011] HOST ALERT: owen101labcolor;UP;SOFT;4;PING WARNING - Packet loss = 96%, RTA = 0.39 ms
[1455671041] SERVICE ALERT: owen101labcolor;Cyan Cartridge HP CE261A;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.232.173. Is the printer off?  Is the network down?
[1455671173] HOST ALERT: owen101labcolor;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1455671204] HOST ALERT: owen101labcolor;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1455671255] SERVICE ALERT: owen101lab3;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1455671259] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 30% - OK!
[1455671268] HOST ALERT: owen101labcolor;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1455671268] HOST FLAPPING ALERT: owen101labcolor;STARTED; Host appears to have started flapping (21.6% change > 20.0% threshold)
[1455671361] SERVICE ALERT: owen101lab3;Device Status;OK;HARD;3;RETRIEVE JOB  NO STORED JOBS
[1455671671] HOST ALERT: owen101labcolor;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1455671761] HOST ALERT: owen101labcolor;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1455671804] HOST ALERT: owen101labcolor;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1455671894] HOST ALERT: owen101labcolor;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1455671943] HOST ALERT: owen101labcolor;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1455671961] SERVICE ALERT: owen101lab3;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1455671994] HOST ALERT: owen101labcolor;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1455671994] HOST FLAPPING ALERT: owen101labcolor;STOPPED; Host appears to have stopped flapping (4.9% change < 5.0% threshold)
[1455672021] SERVICE ALERT: owen101lab3;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1455672025] HOST ALERT: owen101labcolor;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1455672081] SERVICE ALERT: owen101lab3;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1455672115] HOST ALERT: owen101labcolor;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1455672205] HOST ALERT: owen101labcolor;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1455672271] HOST ALERT: owen101labcolor;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1455672415] HOST ALERT: owen101lab3;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1455672460] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1455672465] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1455672490] HOST ALERT: owen101lab3;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1455672517] HOST ALERT: owen101labcolor;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.47 ms
[1455672520] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1455672525] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1455672550] HOST ALERT: owen101lab3;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1455672580] SERVICE ALERT: owen101lab3;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1455672585] SERVICE ALERT: owen101lab3;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1455672594] HOST ALERT: owen101labcolor;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1455672610] HOST ALERT: owen101lab3;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1455672625] HOST ALERT: owen101labcolor;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1455672687] SERVICE ALERT: owen101lab3;Device Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.203.103. Is the printer off?  Is the network down?
[1455672689] HOST ALERT: owen101labcolor;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.54 ms
[1455672700] HOST ALERT: owen101lab3;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1455672790] HOST ALERT: owen101lab3;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1455672880] HOST ALERT: owen101lab3;DOWN;SOFT;7;PING CRITICAL - Packet loss = 100%
[1455672970] HOST ALERT: owen101lab3;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1455673060] HOST ALERT: owen101lab3;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1455673150] HOST ALERT: owen101lab3;DOWN;HARD;10;(Host check timed out after 30.00 seconds)
[1455675124] SERVICE ALERT: DB01-2012;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1455675244] SERVICE ALERT: DB01-2012;CPU Load;OK;SOFT;2;CPU Load 80% (5 min average)
[1455675603] SERVICE ALERT: SANDBOX-DB;CPU Load;WARNING;SOFT;1;CPU Load 94% (5 min average)
[1455675723] SERVICE ALERT: SANDBOX-DB;CPU Load;OK;SOFT;2;CPU Load 14% (5 min average)
[1455680333] SERVICE ALERT: ADWSUSSERVERS;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1455680345] HOST ALERT: ADWSUSSERVERS;DOWN;SOFT;1;CRITICAL - Host Unreachable (164.76.102.133)
[1455680348] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10042;CRITICAL;SOFT;1;(Return code of 255 is out of bounds)
[1455680352] HOST ALERT: ADWSUSSERVERS;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1455680433] SERVICE ALERT: ADWSUSSERVERS;Service: Firewall;OK;SOFT;2;MpsSvc: Started
[1455680456] SERVICE ALERT: ADWSUSSERVERS;Windows Update Services Event 10042;OK;SOFT;2;Eventlog check ok
[1455681515] SERVICE ALERT: eateriespk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1455681575] SERVICE ALERT: eateriespk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1455681635] SERVICE ALERT: eateriespk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1455682235] SERVICE ALERT: eateriespk;Device Status;OK;HARD;3;TRAY 3 SIZE=  ANY SIZE *LETTER  LEGAL