[1462680560] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462680680] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462680800] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 85% (15 min average)
[1462681400] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462681520] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462681640] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462682561] SERVICE ALERT: ADWU1;Windows Update Services Event 364;UNKNOWN;SOFT;1;Failed to process command : EvtQuery failed: 1717: The interface is unknown.
[1462682681] SERVICE ALERT: ADWU1;Windows Update Services Event 364;OK;SOFT;2;OK: No entries found
[1462682823] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;1;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1462682943] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;2;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1462683063] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;HARD;3;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1462683308] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1462683428] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1462683548] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 91% (15 min average)
[1462685348] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462685468] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 92% (15 min average)
[1462688840] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462690040] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462690263] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;OK;HARD;3;OK: No entries found
[1462690868] SERVICE FLAPPING ALERT: BOEAPP;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1462691239] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462691837] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462695436] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 90% (15 min average)
[1462696037] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462696156] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462696276] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1462696876] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1462701076] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462701196] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462701316] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462702516] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462703116] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462703715] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 87% (15 min average)
[1462704315] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1462704435] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462704555] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 66% (15 min average)
[1462705155] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462705268] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1462705275] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462705388] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 92% (15 min average)
[1462705395] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462708988] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1462709108] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 79% (15 min average)
[1462709593] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 82% (15 min average)
[1462710193] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462710313] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 82% (15 min average)
[1462711507] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1462711627] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1462711747] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 76% (15 min average)
[1462714513] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462714633] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1462714753] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1462715953] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1462716553] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462716673] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1462716793] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 85% (15 min average)
[1462717393] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462717513] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462717633] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462721233] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 74% (15 min average)
[1462722433] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462722553] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 86% (15 min average)
[1462723153] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462723273] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462723393] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 83% (15 min average)
[1462723993] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462724113] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462724233] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462725433] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 64% (15 min average)
[1462726033] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1462726153] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462726273] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1462727947] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462728067] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1462728640] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462728700] SERVICE ALERT: hallefirst2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462728760] SERVICE ALERT: hallefirst2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462729961] SERVICE ALERT: hallefirst2;Device Status;OK;HARD;3;Clearing paper path
[1462731073] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 76% (15 min average)
[1462731667] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462731787] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 74% (15 min average)
[1462733473] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462733593] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 98% (15 min average)
[1462733713] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462733973] SERVICE ALERT: halleground2;Device Status;OK;HARD;3;Clearing paper path
[1462734574] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462734633] SERVICE ALERT: halleground2;Device Status;OK;SOFT;2;INPUT PAPER PATH OPEN 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1462735233] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462735293] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462735353] SERVICE ALERT: halleground2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462735987] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1462736107] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 96% (15 min average)
[1462736227] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 80% (15 min average)
[1462736827] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1462736947] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;2;CPU Load 94% (15 min average)
[1462741513] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 92% (15 min average)
[1462742713] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1462742833] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462742953] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462745353] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462745953] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462747753] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)
[1462748353] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462748473] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462748593] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462752313] SERVICE ALERT: PNCDB-2012;D:\ Drive Space;CRITICAL;HARD;3;d:\ - total: 350.00 Gb - used: 319.21 Gb (91%) - free 30.78 Gb (9%)
[1462753953] SERVICE ALERT: halleground2;Device Status;OK;HARD;3;Checking printer
[1462754554] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462754613] SERVICE ALERT: halleground2;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462754673] SERVICE ALERT: halleground2;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462755283] SERVICE ALERT: halleground1;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1462755343] SERVICE ALERT: halleground1;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1462755403] SERVICE ALERT: halleground1;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1462756391] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 79% (15 min average)
[1462756991] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1462757111] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1462757231] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462757831] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462759631] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462760231] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462761878] SERVICE ALERT: BBTS_DB;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462761923] SERVICE ALERT: BBTS_DB;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462761943] SERVICE ALERT: BBTS_DB;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462761943] SERVICE ALERT: BBTS_DB;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1462761978] SERVICE ALERT: BBTS_DB;Service: Windows Time;OK;SOFT;2;W32Time: Started
[1462762023] SERVICE ALERT: BBTS_DB;Uptime;OK;SOFT;2;System Uptime - 32 day(s) 22 hour(s) 24 minute(s)
[1462762043] SERVICE ALERT: BBTS_DB;VMWare Tools;OK;SOFT;2;VMTools: Started
[1462762043] SERVICE ALERT: BBTS_DB;C:\ Drive Space;OK;SOFT;2;c:\ - total: 144.98 Gb - used: 79.51 Gb (55%) - free 65.47 Gb (45%)
[1462762631] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 91% (15 min average)
[1462763231] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1462763345] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1462763351] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1462763465] SERVICE ALERT: BOEAPP;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1462763471] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1462763585] SERVICE ALERT: BOEAPP;CPU Load;OK;SOFT;3;CPU Load 92% (15 min average)
[1462764071] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1462765871] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 93% (15 min average)