[1464667771] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 433 found since last run. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464667891] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464668491] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 363 found since last run. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464668611] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464668750] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 85% (15 min average)
[1464669211] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 560 found since last run. oneLevelSearchOps OK. searchOps is too high! 590 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464669331] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464669350] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464669470] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464669590] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464669931] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 615 found since last run. oneLevelSearchOps OK. searchOps is too high! 640 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464670051] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464670651] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 1096 found since last run. oneLevelSearchOps OK. searchOps is too high! 1117 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464670771] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464671371] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 462 found since last run. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464671491] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464672091] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 347 found since last run. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464672211] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464672811] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 410 found since last run. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464672931] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464673531] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 554 found since last run. oneLevelSearchOps OK. searchOps is too high! 629 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464673651] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464674251] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 334 found since last run. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 795 found since last run.
[1464674371] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464674971] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 678 found since last run. oneLevelSearchOps OK. searchOps is too high! 707 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464675091] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464675589] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1464675691] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 1236 found since last run. oneLevelSearchOps OK. searchOps is too high! 1258 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 716 found since last run.
[1464675811] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464676189] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464676411] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 323 found since last run. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 547 found since last run.
[1464676531] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464677131] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 475 found since last run. oneLevelSearchOps OK. searchOps is too high! 500 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464677251] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464677851] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 492 found since last run. oneLevelSearchOps OK. searchOps is too high! 514 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464677971] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464677985] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464678571] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 454 found since last run. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464678585] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 220 found since last run. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464678691] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464678705] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464679189] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1464679291] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 326 found since last run. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464679305] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 385 found since last run. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464679412] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464679425] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464679789] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1464680011] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 268 found since last run. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464680025] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 315 found since last run. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464680130] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464680145] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464680730] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 232 found since last run. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464680850] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464681345] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 246 found since last run. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464681465] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464682050] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 246 found since last run. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464682065] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 288 found since last run. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 537 found since last run.
[1464682170] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464682185] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464682189] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 83% (15 min average)
[1464682770] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 334 found since last run. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464682785] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 409 found since last run. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 529 found since last run.
[1464682890] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464682905] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464683490] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 361 found since last run. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464683505] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 273 found since last run. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 534 found since last run.
[1464683610] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464683625] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464684210] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 203 found since last run. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464684225] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 332 found since last run. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 566 found since last run.
[1464684330] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464684345] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464684589] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 95% (15 min average)
[1464684709] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464684829] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 94% (15 min average)
[1464684946] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 300 found since last run. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464685065] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464685530] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 200 found since last run. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464685650] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464685665] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 294 found since last run. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464685785] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464686250] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 378 found since last run. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464686370] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464686385] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 371 found since last run. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 531 found since last run.
[1464686505] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464687105] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 472 found since last run. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 513 found since last run.
[1464687225] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464687570] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 338 found since last run. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464687690] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464687825] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 383 found since last run. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464687829] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1464687945] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464687949] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 38% (15 min average)
[1464688290] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 411 found since last run. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464688410] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464688468] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;1;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1464688545] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 386 found since last run. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464688588] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;SOFT;2;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1464688665] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464688708] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;CRITICAL;HARD;3;Windows Server Update Services, The WSUS content directory is not accessible. System.Net.WebException: The remote server returned an error: (503) Server Unavailable.   at System.Net.HttpWebRequest.GetResponse()   at Microsoft.UpdateServices.Internal.HealthMonitoring.HmtWebServices.CheckContentDirWebAccess(EventLoggingType type, HealthEventLogger logger)
[1464689010] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 436 found since last run. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464689130] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464689149] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (15 min average)
[1464689265] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 442 found since last run. oneLevelSearchOps OK. searchOps is too high! 523 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464689269] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1464689385] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464689389] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464689730] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 547 found since last run. oneLevelSearchOps OK. searchOps is too high! 616 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464689850] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464689985] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 502 found since last run.
[1464690105] SERVICE ALERT: idm3.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464690450] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 550 found since last run. oneLevelSearchOps OK. searchOps is too high! 594 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464690570] SERVICE ALERT: idm4.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464690589] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1464690705] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 1299 found since last run. oneLevelSearchOps OK. searchOps is too high! 1327 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 1484 found since last run.
[1464690825] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 271 found since last run. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds is too high! 562 found since last run.
[1464690945] SERVICE ALERT: idm3.emich.edu;LDAP Stats;WARNING;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 507 found since last run. oneLevelSearchOps OK. searchOps is too high! 517 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 820 found since last run.
[1464691170] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 2042 found since last run. oneLevelSearchOps OK. searchOps is too high! 2085 found since last run. strongAuthBinds OK. simpleAuthBinds is too high! 3010 found since last run.
[1464691189] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 80% (15 min average)
[1464691290] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 421 found since last run. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464691410] SERVICE ALERT: idm4.emich.edu;LDAP Stats;WARNING;HARD;3;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 692 found since last run. oneLevelSearchOps OK. searchOps is too high! 708 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464691789] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1464691909] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464692029] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1464693229] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 87% (15 min average)
[1464695629] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1464695749] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 95% (15 min average)
[1464695869] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464695908] SERVICE ALERT: ADWU1;Windows Update Services Event 12072;OK;HARD;3;OK: No entries found
[1464696008] HOST ALERT: porterpk;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1464696469] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 80% (15 min average)
[1464697669] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1464697789] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1464697892] HOST ALERT: porterpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1464697909] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1464697912] HOST ALERT: porterpk;UP;SOFT;2;PING OK - Packet loss = 73%, RTA = 0.96 ms
[1464698492] HOST ALERT: porterpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1464698524] HOST ALERT: porterpk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1464698588] HOST ALERT: porterpk;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.54 ms
[1464698588] HOST FLAPPING ALERT: porterpk;STARTED; Host appears to have started flapping (22.2% change > 20.0% threshold)
[1464698668] SERVICE ALERT: porterpk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! All Trays Still Full
[1464699056] SERVICE ALERT: porterpk;Device Status;OK;HARD;3;Ready
[1464699088] SERVICE ALERT: porterpk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 44% - OK!
[1464699109] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 88% (15 min average)
[1464700309] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1464700429] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 70% (15 min average)
[1464700715] HOST FLAPPING ALERT: porterpk;STOPPED; Host appears to have stopped flapping (4.0% change < 5.0% threshold)
[1464701029] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 96% (15 min average)
[1464701149] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 77% (15 min average)
[1464701749] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464701869] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 38% (15 min average)
[1464701869] SERVICE ALERT: PNCDB-2012;D:\ Drive Space;OK;HARD;3;d:\ - total: 450.00 Gb - used: 114.96 Gb (26%) - free 335.04 Gb (74%)
[1464704269] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464704389] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464704509] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1464705109] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464706309] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 88% (15 min average)
[1464706648] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 747 found since last run. oneLevelSearchOps OK. searchOps is too high! 748 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464706768] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 204 found since last run. oneLevelSearchOps OK. searchOps is too high! 205 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464706888] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;3;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464706888] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1464706949] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1464707008] SERVICE ALERT: mckennypk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1464707509] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464707628] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464707748] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464709548] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 37% (15 min average)
[1464710748] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464710868] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;2;CPU Load 68% (15 min average)
[1464711209] SERVICE ALERT: mckennypk;Device Status;OK;HARD;3;MENUS  RETRIEVE JOB  INFORMATION  PAPER HANDLING
[1464711468] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464711588] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464711708] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464712408] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1464712468] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1464712528] SERVICE ALERT: mckennypk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1464714159] SERVICE ALERT: ADDC4;Service: WMI;UNKNOWN;SOFT;1;Failed to open service manager: 1115: A system shutdown is in progress.
[1464714279] SERVICE ALERT: ADDC4;Service: WMI;OK;SOFT;2;Winmgmt: Started
[1464715898] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;1;WARNING! Printer Needs Paper! 1 Trays Empty
[1464715908] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 77% (15 min average)
[1464715958] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;SOFT;2;WARNING! Printer Needs Paper! 1 Trays Empty
[1464716019] SERVICE ALERT: hallefirstcolor;Paper Tray Status;WARNING;HARD;3;WARNING! Printer Needs Paper! 1 Trays Empty
[1464716488] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1464716508] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1464716548] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1464716608] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1464716628] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464716728] SERVICE ALERT: mckennypk;Device Status;OK;HARD;3;CLOSE FRONT OR LEFT DOOR 46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1464716728] SERVICE FLAPPING ALERT: mckennypk;Device Status;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1464716748] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 98% (15 min average)
[1464717114] SERVICE ALERT: mckennypk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1464717144] HOST ALERT: mckennypk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1464717174] SERVICE ALERT: mckennypk;Paper Tray Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1464717204] HOST ALERT: mckennypk;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1464717234] SERVICE ALERT: mckennypk;Paper Tray Status;UNKNOWN;HARD;3;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1464717244] HOST ALERT: mckennypk;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1464717334] SERVICE ALERT: mckennypk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.152.93. Is the printer off?  Is the network down?
[1464717334] HOST ALERT: mckennypk;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1464717348] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 84% (15 min average)
[1464717388] SERVICE ALERT: mckennypk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1464717392] HOST ALERT: mckennypk;UP;SOFT;5;PING OK - Packet loss = 0%, RTA = 0.52 ms
[1464717448] SERVICE ALERT: mckennypk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1464717808] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;OK;HARD;3;Black Print Cartridge HP C8543X is at 92% - OK!
[1464717829] SERVICE ALERT: mckennypk;Paper Tray Status;OK;HARD;3;Printer Paper Is OK! 1 of 3 Trays Still Full
[1464717948] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464718068] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464718188] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464719608] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1464719668] SERVICE ALERT: mckennypk;Black Print Cartridge HP C8543X;OK;SOFT;2;Black Print Cartridge HP C8543X is at 99% - OK!
[1464720088] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 657 found since last run. oneLevelSearchOps OK. searchOps is too high! 658 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464720208] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464722832] SERVICE ALERT: hallegroundcolor;Image Transfer Kit HP CE249A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1464722862] HOST ALERT: hallegroundcolor;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1464722864] SERVICE ALERT: hallegroundcolor;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1464722879] SERVICE ALERT: hallegroundcolor;Magenta Cartridge HP CE263A;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.54. Is the printer off?  Is the network down?
[1464722882] HOST ALERT: hallegroundcolor;UP;SOFT;2;PING OK - Packet loss = 73%, RTA = 0.37 ms
[1464722885] SERVICE ALERT: hallegroundcolor;Image Transfer Kit HP CE249A;OK;SOFT;2;Image Transfer Kit HP CE249A is at 63% - OK!
[1464722919] SERVICE ALERT: hallegroundcolor;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! All Trays Still Full
[1464722933] SERVICE ALERT: hallegroundcolor;Magenta Cartridge HP CE263A;OK;SOFT;2;
[1464723748] HOST ALERT: INTLTESTDB;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1464723788] SERVICE ALERT: INTLTESTDB;Service: Group Policy;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464723818] HOST ALERT: INTLTESTDB;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1464723888] SERVICE ALERT: INTLTESTDB;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464723888] SERVICE ALERT: INTLTESTDB;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464723908] HOST ALERT: INTLTESTDB;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1464723908] SERVICE ALERT: INTLTESTDB;Service: Group Policy;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464723938] HOST ALERT: INTLTESTDB;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1464724008] SERVICE ALERT: INTLTESTDB;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464724008] SERVICE ALERT: INTLTESTDB;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464724028] HOST ALERT: INTLTESTDB;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1464724028] SERVICE ALERT: INTLTESTDB;Service: Group Policy;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464724042] SERVICE ALERT: INTLTESTDB;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464724045] SERVICE ALERT: INTLTESTDB;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464724058] HOST ALERT: INTLTESTDB;DOWN;SOFT;6;PING CRITICAL - Packet loss = 100%
[1464724128] SERVICE ALERT: INTLTESTDB;Service: NetLogin;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464724128] SERVICE ALERT: INTLTESTDB;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464724129] SERVICE ALERT: INTLTESTDB;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464724134] SERVICE ALERT: INTLTESTDB;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464724148] HOST ALERT: INTLTESTDB;DOWN;SOFT;7;(Host check timed out after 30.00 seconds)
[1464724162] SERVICE ALERT: INTLTESTDB;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464724165] SERVICE ALERT: INTLTESTDB;Service: User Profile;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464724192] HOST ALERT: INTLTESTDB;DOWN;SOFT;8;PING CRITICAL - Packet loss = 100%
[1464724209] SERVICE ALERT: INTLTESTDB;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464724214] SERVICE ALERT: INTLTESTDB;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464724239] SERVICE ALERT: INTLTESTDB;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464724239] HOST ALERT: INTLTESTDB;DOWN;SOFT;9;PING CRITICAL - Packet loss = 100%
[1464724249] SERVICE ALERT: INTLTESTDB;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464724254] SERVICE ALERT: INTLTESTDB;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464724259] SERVICE ALERT: INTLTESTDB;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464724279] HOST ALERT: INTLTESTDB;DOWN;HARD;10;PING CRITICAL - Packet loss = 100%
[1464724282] SERVICE ALERT: INTLTESTDB;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464724285] SERVICE ALERT: INTLTESTDB;Service: User Profile;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464724308] SERVICE ALERT: INTLTESTDB;VMWare Tools;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464724308] SERVICE ALERT: INTLTESTDB;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464724318] SERVICE ALERT: INTLTESTDB;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464724329] SERVICE ALERT: INTLTESTDB;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464724334] SERVICE ALERT: INTLTESTDB;Service: Windows Time;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464724359] SERVICE ALERT: INTLTESTDB;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464724369] SERVICE ALERT: INTLTESTDB;Service: WMI;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464724373] SERVICE ALERT: INTLTESTDB;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464724379] SERVICE ALERT: INTLTESTDB;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464724428] SERVICE ALERT: INTLTESTDB;VMWare Tools;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464724428] SERVICE ALERT: INTLTESTDB;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464724438] SERVICE ALERT: INTLTESTDB;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464724449] SERVICE ALERT: INTLTESTDB;Service: DNS Client;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464724454] SERVICE ALERT: INTLTESTDB;Service: Windows Time;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464724462] SERVICE ALERT: INTLTESTDB;Uptime;CRITICAL;HARD;3;connect to address 164.76.7.82 and port 12489: Connection refused
[1464724481] SERVICE ALERT: INTLTESTDB;Service: RPC;OK;SOFT;3;RpcSs: Started
[1464724528] SERVICE ALERT: INTLTESTDB;VMWare Tools;OK;SOFT;3;VMTools: Started
[1464724528] SERVICE ALERT: INTLTESTDB;Service: Firewall;OK;SOFT;3;MpsSvc: Started
[1464724538] SERVICE ALERT: INTLTESTDB;CPU Load;OK;SOFT;3;CPU Load 2% (15 min average)
[1464724608] SERVICE ALERT: INTLTESTDB;Service: Group Policy;OK;HARD;3;gpsvc: Started
[1464724708] SERVICE ALERT: INTLTESTDB;C:\ Drive Space;OK;HARD;3;c:\ - total: 84.22 Gb - used: 71.59 Gb (85%) - free 12.63 Gb (15%)
[1464725099] SERVICE ALERT: INTLTESTDB;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464725132] HOST ALERT: INTLTESTDB;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1464725199] SERVICE ALERT: INTLTESTDB;Service: RPC;OK;SOFT;2;RpcSs: Started
[1464725308] SERVICE ALERT: INTLTESTDB;Service: NetLogin;OK;HARD;3;Netlogon: Started
[1464725462] SERVICE ALERT: INTLTESTDB;Memory Usage;OK;HARD;3;Memory usage: total:14311.42 MB - used: 1183.42 MB (8%) - free: 13128.00 MB (92%)
[1464725465] SERVICE ALERT: INTLTESTDB;Service: User Profile;OK;HARD;3;ProfSvc: Started
[1464725549] SERVICE ALERT: INTLTESTDB;Service: WMI;OK;HARD;3;Winmgmt: Started
[1464725552] SERVICE ALERT: INTLTESTDB;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.15 2015-11-25
[1464725629] SERVICE ALERT: INTLTESTDB;Service: DNS Client;OK;HARD;3;Dnscache: Started
[1464725634] SERVICE ALERT: INTLTESTDB;Service: Windows Time;OK;HARD;3;W32Time: Started
[1464725659] SERVICE ALERT: INTLTESTDB;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 9 minute(s)
[1464726808] SERVICE ALERT: idm1.emich.edu;LDAP Stats;WARNING;SOFT;1;errors OK. securityErrors OK. wholeSubtreeSearchOps is too high! 297 found since last run. oneLevelSearchOps OK. searchOps is too high! 298 found since last run. strongAuthBinds OK. simpleAuthBinds OK.
[1464726928] SERVICE ALERT: idm1.emich.edu;LDAP Stats;OK;SOFT;2;errors OK. securityErrors OK. wholeSubtreeSearchOps OK. oneLevelSearchOps OK. searchOps OK. strongAuthBinds OK. simpleAuthBinds OK.
[1464728388] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 73% (15 min average)
[1464728848] SERVICE FLAPPING ALERT: mckennypk;Device Status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1464728855] SERVICE ALERT: halleground2;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.106.52. Is the printer off?  Is the network down?
[1464728885] HOST ALERT: halleground2;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1464728911] SERVICE ALERT: halleground2;Device Status;OK;SOFT;2;57.03 ERROR  To continue turn off then on 57.03 ERROR  46 6F 72 20 68 65 6C 70 20 70 72 65 73 73 20 1F
[1464728915] HOST ALERT: halleground2;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.53 ms
[1464728988] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464729108] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1464729228] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 71% (15 min average)
[1464729823] SERVICE ALERT: emu-fndn-fe7;Service: Firewall;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464729828] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464729838] SERVICE ALERT: emu-fndn-fe7;Service: NetLogin;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464729853] HOST ALERT: emu-fndn-fe7;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1464729888] SERVICE ALERT: emu-fndn-fe7;Service: DNS Client;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464729897] SERVICE ALERT: emu-fndn-fe7;Service: WMI;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464729918] HOST ALERT: emu-fndn-fe7;DOWN;SOFT;2;PING CRITICAL - Packet loss = 100%
[1464729921] SERVICE ALERT: emu-fndn-fe7;Service: RPC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464729943] SERVICE ALERT: emu-fndn-fe7;Service: Firewall;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464729948] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464729948] SERVICE ALERT: emu-fndn-fe7;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464729948] SERVICE ALERT: emu-fndn-fe7;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464729951] HOST ALERT: emu-fndn-fe7;DOWN;SOFT;3;PING CRITICAL - Packet loss = 100%
[1464729958] SERVICE ALERT: emu-fndn-fe7;Service: NetLogin;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464729988] HOST ALERT: emu-fndn-fe7;DOWN;SOFT;4;PING CRITICAL - Packet loss = 100%
[1464730008] SERVICE ALERT: emu-fndn-fe7;Service: Windows Time;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464730008] SERVICE ALERT: emu-fndn-fe7;Service: DNS Client;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464730017] SERVICE ALERT: emu-fndn-fe7;Service: WMI;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464730038] HOST ALERT: emu-fndn-fe7;DOWN;SOFT;5;PING CRITICAL - Packet loss = 100%
[1464730041] SERVICE ALERT: emu-fndn-fe7;Service: RPC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464730063] SERVICE ALERT: emu-fndn-fe7;Service: Firewall;CRITICAL;HARD;3;CRITICAL - Socket timeout after 20 seconds
[1464730066] HOST ALERT: emu-fndn-fe7;UP;SOFT;6;PING WARNING - Packet loss = 80%, RTA = 0.39 ms
[1464730067] SERVICE ALERT: emu-fndn-fe7;Service: NetLogin;OK;SOFT;3;Netlogon: Started
[1464730068] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 88% (15 min average)
[1464730068] SERVICE ALERT: emu-fndn-fe7;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464730068] SERVICE ALERT: emu-fndn-fe7;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 20 seconds
[1464730071] SERVICE ALERT: emu-fndn-fe7;Service: User Profile;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 20 seconds
[1464730108] SERVICE ALERT: emu-fndn-fe7;Service: DNS Client;OK;SOFT;3;Dnscache: Started
[1464730108] SERVICE ALERT: emu-fndn-fe7;Service: Windows Time;CRITICAL;SOFT;2;W32Time: Stopped
[1464730117] SERVICE ALERT: emu-fndn-fe7;Service: WMI;OK;SOFT;3;Winmgmt: Started
[1464730141] SERVICE ALERT: emu-fndn-fe7;Service: RPC;OK;SOFT;3;RpcSs: Started
[1464730168] SERVICE ALERT: emu-fndn-fe7;Memory Usage;OK;SOFT;3;Memory usage: total:18781.43 MB - used: 1624.78 MB (9%) - free: 17156.65 MB (91%)
[1464730168] SERVICE ALERT: emu-fndn-fe7;C:\ Drive Space;OK;SOFT;3;c:\ - total: 698.53 Gb - used: 42.67 Gb (6%) - free 655.86 Gb (94%)
[1464730171] SERVICE ALERT: emu-fndn-fe7;Service: User Profile;OK;SOFT;2;ProfSvc: Started
[1464730228] SERVICE ALERT: emu-fndn-fe7;Service: Windows Time;OK;SOFT;3;W32Time: Started
[1464730554] SERVICE ALERT: markjeffersonpk;Device Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1464730584] HOST ALERT: markjeffersonpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1464730614] SERVICE ALERT: markjeffersonpk;Device Status;UNKNOWN;SOFT;2;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1464730618] HOST ALERT: markjeffersonpk;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.95 ms
[1464730643] SERVICE ALERT: emu-fndn-fe7;Service: Firewall;OK;HARD;3;MpsSvc: Started
[1464730668] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464730668] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1464730787] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464730907] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;SOFT;3;CPU Load 85% (15 min average)
[1464731508] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 98% (15 min average)
[1464731627] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 97% (15 min average)
[1464731746] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464736668] SERVICE ALERT: markjeffersonpk;Device Status;OK;HARD;3;Open  tray 4. Remove any  jammed
[1464737010] HOST ALERT: markjeffersonpk;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1464737024] SERVICE ALERT: markjeffersonpk;Paper Tray Status;UNKNOWN;SOFT;1;Unknown: No SNMP response from 164.76.81.90. Is the printer off?  Is the network down?
[1464737035] HOST ALERT: markjeffersonpk;UP;SOFT;2;PING OK - Packet loss = 58%, RTA = 1.24 ms
[1464737079] SERVICE ALERT: markjeffersonpk;Paper Tray Status;OK;SOFT;2;Printer Paper Is OK! All Trays Still Full
[1464737146] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1464737746] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464737869] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!
[1464737928] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;SOFT;2;Possible Paper Jam Detected!
[1464737988] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;HARD;3;Possible Paper Jam Detected!
[1464739432] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;1;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1464739492] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;UNKNOWN;SOFT;2;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1464739552] SERVICE ALERT: hallesecondsouth;Black Print Cartridge HP C8543X;UNKNOWN;HARD;3;UNKNOWN - No match found for 'Black Print Cartridge HP C8543X'! Use the CONSUM TEST option to determine which consumables may be monitored.
[1464741346] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 77% (15 min average)
[1464741946] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464742066] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464742186] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 97% (15 min average)
[1464742786] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 89% (15 min average)
[1464743386] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 97% (15 min average)
[1464743506] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;2;CPU Load 99% (15 min average)
[1464743626] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464750226] SERVICE ALERT: SCCM-PRIM01;CPU Load;OK;HARD;3;CPU Load 90% (15 min average)
[1464750826] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;SOFT;1;CPU Load 99% (15 min average)
[1464750946] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (15 min average)
[1464751066] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464751666] SERVICE ALERT: SCCM-PRIM01;CPU Load;CRITICAL;HARD;3;CPU Load 100% (15 min average)
[1464752266] SERVICE ALERT: SCCM-PRIM01;CPU Load;WARNING;HARD;3;CPU Load 99% (15 min average)
[1464752988] SERVICE ALERT: markjeffersonpk;Device Status;OK;HARD;3;Warming up
[1464753588] SERVICE ALERT: markjeffersonpk;Device Status;WARNING;SOFT;1;Possible Paper Jam Detected!